<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>1.2.1: Compare Two Datasets</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The heights of maize plants of two different varieties V1 and V2 are measured in centimetres. The dot plots below show the heights of samples of plants from each variety.</p>

<p>Describe the differences between the height of plants for the two varieties.</p>

[[jsxgraph]] 
// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});

//Create points of Data Set 2
var lx2 = {#dorx2#};
var ly2 = {#dory2#};
var i = 0;
for (i = 0; i <= lx2.length - 1; i++) {
  p[i] = board.create('point', [lx2[i], ly2[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 2
var txt = board.create('text', [{#xmin#}+1, {#maxfreq2#}/2+2, "V2"], {
  fontSize: 30,fixed: true
});

//Create points of Data Set 1
var lx1 = {#dorx1#};
var ly1 = {#dory1#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 1
var txt = board.create('text', [{#xmin#}+1, {#maxfreq1#}/2+{#maxfreq2#} + 5, "V1"], {
  fontSize: 30,fixed: true
});
[[/jsxgraph]]
</p>
<ol class="HELM_parts">
<li><p> How do the heights compare?</p></li>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]

<li><p> How does the data's centre compare?</p></li>
[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]

<li><p> How does the variability of heights compare?</p></li>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>

<li><p> How does the data's spread compare?</p></li>
<p>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[mean1:rand(101)+100;
aa:rand([-1,0,1]);
mean2:mean1+(rand(11)+10)*aa;
sd1:rand(5)+5;
bb:rand([1/3,1/2,1,2,3]);
sd2:sd1*bb
ds1:rand(70)+30;
ds2:ds1;
dun1:makelist(round(float(quantile_normal(rand(1.0),mean1,sd1))),i,1,ds1);
mdun1:float(mean(dun1));
r_mdun1:round(mdun1);
sddun1:round(float(std1(dun1)));
dorx1:sort(dun1);
dun2:makelist(round(float(quantile_normal(rand(1.0),mean2,sd2))),i,1,ds2);
mdun2:float(mean(dun2));
mdiff:float(mdun1-mdun2);
dun2: if aa =0 then dun2+mdiff else dun2;
mdun2:float(mean(dun2));
r_mdun2:round(mdun2);
sddun2:round(float(std1(dun2)));
dorx2:sort(dun2);

maxfreq1:0;
k:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;
maxfreq2:0;
k:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (k:k+1, if maxfreq2 < k then maxfreq2:k) else k:1;

xmin:if dorx1[1]<=dorx2[1] then dorx1[1]-2 else dorx2[1]-2;
xmax:if dorx1[ds1]>=dorx2[ds2] then dorx1[ds1]+2+(dorx1[ds1]-dorx1[1])/1.5 else dorx2[ds2]+2+(dorx2[ds2]-dorx2[1])/1.5;
ymin:-4;
ymax: maxfreq1+maxfreq2+5;

j:1;
dory1:makelist(0,i,1,ds1);
dory1[1]:1+maxfreq2+3;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory1[i]:j+maxfreq2+3) else (j:1, dory1[i]:j+maxfreq2+3);
j:1;
dory2:makelist(0,i,1,ds2);
dory2[1]:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (j:j+1, dory2[i]:j) else (j:1, dory2[i]:j);

ranchoice1:random_permutation([[1,"Production line L1 tends to produce longer widgets than L2"],[2,"Production line L2 tends to produce longer widgets than L1"],[3,"Neither production line tends to produce longer widgets than the other one"]]);



ta1_c:if aa>0 then 2 elseif aa<0 then 1 else 3;
ta1:[[3, is(ta1_c=3), "Neither variety has consistently taller plants than the other one."],[1, is(ta1_c=1), "On average variety V1 has much taller plants than V2."], [2, is(ta1_c=2), "On average variety V2 has much taller plants than V1."]];

ta2_c:if aa>0 then 2 elseif aa<0 then 1 else 3;
ta2:[[3, is(ta2_c=3), "\\( \\text{centre}(V1) \\approx \\text{centre}(V2) \\) "],[1, is(ta2_c=1), "\\( \\text{centre}(V1) > \\text{centre}(V2) \\)"], [2, is(ta2_c=2), "\\( \\text{centre}(V1) < \\text{centre}(V2) \\)"]];

ma1:if mdun1<mdun2 then 2 elseif mdun1>mdun2 then 1;

ta3_c:if bb>1 then 2 elseif bb<1 then 1 else 3;
ta3:[[3, is(ta3_c=3), "There is not a large difference between the variability of the heights of both varieties."],[1, is(ta3_c=1), "Variety V1 has plants whose heights are more variable than V2."], [2, is(ta3_c=2), "Variety V2 has plants whose heights are more variable than V1."]];

ta4_c:if bb>1 then 2 elseif bb<1 then 1 else 3;
ta4:[[3, is(ta4_c=3), "\\( \\text{spread}(V1) \\approx \\text{spread}(V2) \\) "],[1, is(ta4_c=1), "\\( \\text{spread}(V1) > \\text{spread}(V2) \\)"], [2, is(ta4_c=2), "\\( \\text{spread}(V1) < \\text{spread}(V2) \\)"]];


sda1:if mdun1<mdun2 then 2 elseif mdun1>mdun2 then 1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>sd1 = {@sd1@}, sd2 = {@sd2@}.
mean1 = {@mean1@}, mean2 = {@mean2@}
ta1 = {@ta1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans3</name>
      <type>radio</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans4</name>
      <type>radio</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1_c</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, you have correctly interpreted the graph. The variety V1 has mean {@r_mdun1@} and V2 has mean {@r_mdun2@}.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ma1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done you have correctly identified the relative difference between the averages, however, the difference between the mean of V1 {@r_mdun1@} and the mean of V2 {@r_mdun2@} is relatively small and so these could have been considered roughly equal.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>You have not correctly interpreted the representation of the graph. The mean of V1 is {@r_mdun1@} and the mean of V2 is {@r_mdun2@}. This question is asking you to roughly compare these means by eye, it is a skill which can be acquired through practice.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2_c</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, you have correctly interpreted the graph. The variety V1 has mean {@r_mdun1@} and V2 has mean {@r_mdun2@}.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br> 
You have answered the first two parts inconsistently. Notice that they are asking the same thing, the only difference is the language used, the meaning remains the same.]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br> 
Well done you have answered the first two parts consistently. The only difference is the language used, the meaning remains the same.]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3_c</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, you have correctly interpreted the graph. The variety V1 has standard deviation {@sddun1@} and V2 has standard deviation {@sddun2@}.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>You have not correctly interpreted the representation of the graph. The standard deviation of V1 is {@sddun1@} and the standard deviation of V2 is {@sddun2@}. This question is asking you to roughly compare these standard deviations by eye, it is a skill which can be acquired through practice.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>sda1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You have not correctly interpreted the representation of the graph. The standard deviation of V1 is {@sddun1@} and the standard deviation of V2 is {@sddun2@}. This question is asking you to roughly compare these standard deviations by eye, it is a skill which can be acquired through practice.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4_c</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, you have correctly interpreted the graph. The variety V1 has standard deviation {@sddun1@} and V2 has standard deviation {@sddun2@}.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>sda1</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You have not correctly interpreted the representation of the graph. The standard deviation of V1 is {@sddun1@} and the standard deviation of V2 is {@sddun2@}. This question is asking you to roughly compare these standard deviations by eye, it is a skill which can be acquired through practice.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ans4</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You have answered the last two parts inconsistently. Notice that they are asking the same thing, the only difference is the language used, the meaning remains the same.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>599769975</deployedseed>
    <deployedseed>1841930269</deployedseed>
    <deployedseed>194246555</deployedseed>
    <deployedseed>492753055</deployedseed>
    <deployedseed>1059780349</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description></description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
