<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Arithmetic Mean and Other Stats Calculation for a Given Dataset 1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the following data set:</p>

<p>{@data@}</p>

<p><em>Note: For the following, submit your answer correct to two decimal places where appropriate.</em></p>

<p>a. Calculate its arithmetic mean: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p>b. Calculate its median: [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<p>c. Calculate its standard deviation: [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>

<p>d. Calculate its interquartile range: [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Arithmetic mean</strong></p>

<p>The arithmetic mean is given by the formula
\( \displaystyle \bar{x} = \frac{\sum_{i=1}^n x_i}{n}\).</p>

<p>In this case,
\( \displaystyle \sum_{i=1}^n x_i = {@sum(data[i],i,1,20)@}\)
and \(n = 20\). Therefore,
\(\bar{x} = {@disp_mean@}\).</p>


<p><strong>Median</strong></p>

<p>The median is the middle value of an ordered dataset. Since there are \(20\) observations, the median is the arithmetic mean of the 10th and 11th values after the data have been arranged in ascending order.</p>

<p>Therefore, the median of this dataset is
\({@disp_median@}\).</p>


<p><strong>Standard deviation</strong></p>

<p>The population standard deviation is given by the formula
\( \displaystyle \sigma = \sqrt{\frac{\sum_{i=1}^n (x_i-\bar{x})^2}{n}}\).</p>

<p>For this dataset, \(n=20\) and
\(\bar{x}={@disp_mean@}\).</p>

<p>Therefore, the standard deviation is
\(\sigma = {@disp_std@}\).</p>


<p><strong>Interquartile range</strong></p>

<p>The interquartile range is the difference between the third quartile and the first quartile:
\( \displaystyle \mathrm{IQR}=Q_3-Q_1\).</p>

<p>For this dataset,
\(Q_1={@disp_q1@}\) and
\(Q_3={@disp_q3@}\).</p>

<p>Therefore,
\(\mathrm{IQR}={@disp_q3@}-{@disp_q1@}
={@disp_iqr@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Generate a random dataset containing 20 positive integers */
data:makelist(1+rand(1+rand(500)),i,1,20);

/* Calculate the arithmetic mean of the dataset */
mean_data:mean(data);

/* Convert the mean to a floating-point value for the student answer */
ta1:float(mean_data);

/* Calculate the median of the dataset */
median_data:median(data);

/* Convert the median to a floating-point value for the student answer */
ta2:float(median_data);

/* Calculate the population standard deviation of the dataset */
std_data:std(data);

/* Convert the standard deviation to a floating-point value for the student answer */
ta3:float(std_data);

/* Calculate the first quartile */
q1:quantile(data,1/4);

/* Calculate the third quartile */
q3:quantile(data,3/4);

/* Calculate the interquartile range (Q3 - Q1) */
iqr_data:q3-q1;

/* Convert the interquartile range to a floating-point value for the student answer */
ta4:float(iqr_data);

/* Variables for display in general feedback */
/* Format the mean without unnecessary trailing zeros */
disp_mean:
    if integerp(mean_data) then mean_data
    elseif integerp(10*mean_data) then dispdp(float(mean_data),1)
    else dispdp(float(mean_data),2);

/* Format the median without unnecessary trailing zeros */
disp_median:
    if integerp(median_data) then median_data
    elseif integerp(10*median_data) then dispdp(float(median_data),1)
    else dispdp(float(median_data),2);

/* Format the standard deviation without unnecessary trailing zeros */
disp_std:
    if integerp(std_data) then std_data
    elseif integerp(10*std_data) then dispdp(float(std_data),1)
    else dispdp(float(std_data),2);

/* Format the first quartile without unnecessary trailing zeros */
disp_q1:
    if integerp(q1) then q1
    elseif integerp(10*q1) then dispdp(float(q1),1)
    else dispdp(float(q1),2);

/* Format the third quartile without unnecessary trailing zeros */
disp_q3:
    if integerp(q3) then q3
    elseif integerp(10*q3) then dispdp(float(q3),1)
    else dispdp(float(q3),2);

/* Format the interquartile range without unnecessary trailing zeros */
disp_iqr:
    if integerp(iqr_data) then iqr_data
    elseif integerp(10*iqr_data) then dispdp(float(iqr_data),1)
    else dispdp(float(iqr_data),2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Data: {@data@}</p>
<p>Mean={@ta1@}, median={@ta2@}, sd={@ta3@}, IQR={@ta4@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Santiago Borio</p>
<p id="concept">James Musyoka</p>
<p id="reviewer"></p>
<p id="description">This question asks students to calculate the arithmetic mean of a given dataset. The answer should be submitted rounded to two decimal places.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>dispdp(ta1,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispdp(ta2,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>dispdp(ta3,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>dispdp(ta4,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check if the mean is integer*/
int_ta: if integerp(mean_data) then 1 else 0;

/*Check if the mean has exactly 1dp*/
one_dp_ta: if integerp(10*mean_data) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>test correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check if answer is integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>int_ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if correct answer when integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mean_data</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Check if mean has exactly 1dp </description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>one_dp_ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Check correct answer when exactly 1dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>1</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check if the mean is integer*/
int_ta: if integerp(median_data) then 1 else 0;

/*Check if the mean has exactly 1dp*/
one_dp_ta: if integerp(10*median_data) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>test correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check if answer is integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>int_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if correct answer when integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Check if mean has exactly 1dp </description>
        <answertest>AlgEquiv</answertest>
        <sans>one_dp_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Check correct answer when exactly 1dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check if the standard deviation is integer*/
int_ta: if integerp(std_data) then 1 else 0;

/*Check if the standard deviation has exactly 1dp*/
one_dp_ta: if integerp(10*std_data) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>test correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check if answer is integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>int_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if correct answer when integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Check if standard deviation has exactly 1dp </description>
        <answertest>AlgEquiv</answertest>
        <sans>one_dp_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Check correct answer when exactly 1dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Check if the interquartile range is integer*/
int_ta: if integerp(iqr_data) then 1 else 0;

/*Check if the interquartile range has exactly 1dp*/
one_dp_ta: if integerp(10*iqr_data) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>test correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check if answer is integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>int_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if correct answer when integer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Check if interquartile range has exactly 1dp </description>
        <answertest>AlgEquiv</answertest>
        <sans>one_dp_ta</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt4-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Check correct answer when exactly 1dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>515232793</deployedseed>
    <deployedseed>538830914</deployedseed>
    <deployedseed>1482579168</deployedseed>
    <deployedseed>1481783857</deployedseed>
    <deployedseed>1139987239</deployedseed>
    <deployedseed>74458267</deployedseed>
    <deployedseed>378498156</deployedseed>
    <deployedseed>1656381872</deployedseed>
    <deployedseed>1537867514</deployedseed>
    <deployedseed>1146669806</deployedseed>
    <deployedseed>1922014844</deployedseed>
    <deployedseed>981705887</deployedseed>
    <deployedseed>764065018</deployedseed>
    <deployedseed>456916676</deployedseed>
    <deployedseed>229848150</deployedseed>
    <deployedseed>492617448</deployedseed>
    <deployedseed>1689793631</deployedseed>
    <deployedseed>950991960</deployedseed>
    <deployedseed>337046062</deployedseed>
    <deployedseed>51590523</deployedseed>
    <deployedseed>168200283</deployedseed>
    <deployedseed>711584981</deployedseed>
    <deployedseed>553634567</deployedseed>
    <deployedseed>723542397</deployedseed>
    <deployedseed>1165324100</deployedseed>
    <deployedseed>1190540010</deployedseed>
    <deployedseed>1699762144</deployedseed>
    <deployedseed>1719469404</deployedseed>
    <deployedseed>1597998938</deployedseed>
    <deployedseed>1188490533</deployedseed>
    <deployedseed>536068537</deployedseed>
    <deployedseed>81018141</deployedseed>
    <deployedseed>1015462199</deployedseed>
    <deployedseed>797974361</deployedseed>
    <deployedseed>1299110429</deployedseed>
    <deployedseed>1190137294</deployedseed>
    <deployedseed>1373756447</deployedseed>
    <deployedseed>2135648391</deployedseed>
    <deployedseed>728664280</deployedseed>
    <deployedseed>1055385759</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta1,2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(ta2,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispdp(ta3,2)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispdp(ta4,2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt2-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt3-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt4-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
