<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating Summary Statistics for Cystic Fibrosis Patients' PLmax Data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Maximal static inspiratory pressure (OImax) is an index of respiratory muscle strength. The following data show the measure of  PLmax (cm \(H_2O\) ) for \(15\) cystic fibrosis patients.</p>
<p>
<table border="solid">
  <tr>
    <th width="12.5%", style="text-align:center">\({@a1@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a2@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a3@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a4@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a5@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a6@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a7@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a8@}\)</th>
  </tr>
<tr>
    <th width="12.5%", style="text-align:center">\({@a9@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a10@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a11@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a12@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a13@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a14@}\)</th>
    <th width="12.5%", style="text-align:center">\({@a15@}\)</th>
    <th width="12.5%", style="text-align:center"></th>
  </tr>
</table>
<p>Find:</p>
<p>Lower quartile  \(Q_{1}=\)  [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>Median  \(Q_{2}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>Upper quartile  \(Q_{3}=\)   [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p>Interquartile range  \(R=\)  [[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Arranging the data in an ascending order</p>
<p>\({@sorted_list@}\)</p>
<p>To find the lower half median, you need to find the middle value of the lower half which in our case by inspection:</p>
<p> Lower quartile \(Q_1 = {@ta1@}\)</p>
<p>The median \( (Q_2) \) is the middle value in the data set.</p>
<p>For this data set, as it has an odd number of values \( (15) \), the median is the middle value:</p>
<p>median \(Q_2= {@ta2@}\)</p>
<p>To find the upper half median, you need to find the middle value of the upper half which in our case by inspection:</p>
<p>Upper quartile \(Q_3 = {@ta3@}\)</p>
<p>Interquartile range \(R = Q_3- Q_1= {@ta3@}-{@ta1@} ={@ta4@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a1:rand_with_step(40,120,5);
a2:rand_with_step(40,120,5);
a3:rand_with_step(40,120,5);
a4:rand_with_step(40,120,5);
a5:rand_with_step(40,120,5);
a6:rand_with_step(40,120,5);
a7:rand_with_step(40,120,5);
a8:rand_with_step(40,120,5);
a9:rand_with_step(40,120,5);
a10:rand_with_step(40,120,5);
a11:rand_with_step(40,120,5);
a12:rand_with_step(40,120,5);
a13:rand_with_step(40,120,5);
a14:rand_with_step(40,120,5);
a15:rand_with_step(40,120,5);
my_list: [a1, a2,a3,a4,a5,a6,a7,a8,a9,a10,a11,a12,a13,a14,a15];
sorted_list: sort(my_list);
ta1:sorted_list[4];
ta2:sorted_list[8];
ta3:sorted_list[12];
ta4:ta3-ta1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Arranging the data in an ascending order: \({@sorted_list@}\)</p>
<p>\(Q_1 = {@ta1@}\)</p>
<p>\(Q_2 = {@ta2@}\)</p>
<p>\(Q_3 = {@ta3@}\)</p>
<p>Interquartile range \(R= {@ta4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This question requires the student to compute the lower and upper quartiles, find the median, and determine the interquartile range for the PLmax (cm H2O) measurements of 15 cystic fibrosis patients.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1377723106</deployedseed>
    <deployedseed>325010306</deployedseed>
    <deployedseed>1607287397</deployedseed>
    <deployedseed>764871418</deployedseed>
    <deployedseed>475034114</deployedseed>
    <deployedseed>1558648206</deployedseed>
    <deployedseed>922383801</deployedseed>
    <deployedseed>460773369</deployedseed>
    <deployedseed>517322375</deployedseed>
    <deployedseed>426262423</deployedseed>
    <deployedseed>505938092</deployedseed>
    <deployedseed>34441045</deployedseed>
    <deployedseed>1012699754</deployedseed>
    <deployedseed>1531729084</deployedseed>
    <deployedseed>575959981</deployedseed>
    <deployedseed>1528702825</deployedseed>
    <deployedseed>420847332</deployedseed>
    <deployedseed>1889052817</deployedseed>
    <deployedseed>1191305439</deployedseed>
    <deployedseed>1680645586</deployedseed>
    <deployedseed>498685981</deployedseed>
    <deployedseed>116991156</deployedseed>
    <deployedseed>380942905</deployedseed>
    <deployedseed>370840487</deployedseed>
    <deployedseed>1935577940</deployedseed>
    <deployedseed>1618234435</deployedseed>
    <deployedseed>1365654249</deployedseed>
    <deployedseed>783956556</deployedseed>
    <deployedseed>152349707</deployedseed>
    <deployedseed>1269009819</deployedseed>
    <deployedseed>1853057122</deployedseed>
    <deployedseed>805437991</deployedseed>
    <deployedseed>1037643837</deployedseed>
    <deployedseed>1072188864</deployedseed>
    <deployedseed>65133315</deployedseed>
    <deployedseed>1412839218</deployedseed>
    <deployedseed>1828464578</deployedseed>
    <deployedseed>536435993</deployedseed>
    <deployedseed>1172144921</deployedseed>
    <deployedseed>1918110642</deployedseed>
    <deployedseed>888425633</deployedseed>
    <deployedseed>639392760</deployedseed>
    <deployedseed>312340641</deployedseed>
    <deployedseed>252126266</deployedseed>
    <deployedseed>1391895969</deployedseed>
    <deployedseed>862239082</deployedseed>
    <deployedseed>1598520102</deployedseed>
    <deployedseed>1214697790</deployedseed>
    <deployedseed>251599870</deployedseed>
    <deployedseed>2069478679</deployedseed>
    <deployedseed>677000751</deployedseed>
    <deployedseed>762120799</deployedseed>
    <deployedseed>1989895665</deployedseed>
    <deployedseed>1011788764</deployedseed>
    <deployedseed>2089946659</deployedseed>
    <deployedseed>1450317396</deployedseed>
    <deployedseed>1610115535</deployedseed>
    <deployedseed>992489351</deployedseed>
    <deployedseed>387694367</deployedseed>
    <deployedseed>1899147407</deployedseed>
    <deployedseed>826156208</deployedseed>
    <deployedseed>19218282</deployedseed>
    <deployedseed>1676501240</deployedseed>
    <deployedseed>1465152797</deployedseed>
    <deployedseed>1467872318</deployedseed>
    <deployedseed>613190938</deployedseed>
    <deployedseed>1810125949</deployedseed>
    <deployedseed>824417142</deployedseed>
    <deployedseed>1445162881</deployedseed>
    <deployedseed>233684862</deployedseed>
    <deployedseed>123935605</deployedseed>
    <deployedseed>1760921830</deployedseed>
    <deployedseed>2075733039</deployedseed>
    <deployedseed>68319303</deployedseed>
    <deployedseed>836423922</deployedseed>
    <deployedseed>1265033212</deployedseed>
    <deployedseed>780360283</deployedseed>
    <deployedseed>289948724</deployedseed>
    <deployedseed>1247611329</deployedseed>
    <deployedseed>120981125</deployedseed>
    <deployedseed>1034116952</deployedseed>
    <deployedseed>1660240253</deployedseed>
    <deployedseed>2106966691</deployedseed>
    <deployedseed>711825693</deployedseed>
    <deployedseed>548471696</deployedseed>
    <deployedseed>2069922515</deployedseed>
    <deployedseed>1006727530</deployedseed>
    <deployedseed>1774266143</deployedseed>
    <deployedseed>1829270519</deployedseed>
    <deployedseed>1650570088</deployedseed>
    <deployedseed>1103946721</deployedseed>
    <deployedseed>43066073</deployedseed>
    <deployedseed>1515807934</deployedseed>
    <deployedseed>2142985897</deployedseed>
    <deployedseed>61520460</deployedseed>
    <deployedseed>1772133223</deployedseed>
    <deployedseed>1914899272</deployedseed>
    <deployedseed>887643175</deployedseed>
    <deployedseed>1175151957</deployedseed>
    <deployedseed>597941009</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
