<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the combined mean and standard deviation for two datasets presented as dotplots</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>If two or more groups of values with different means and standard deviations are combined, you should be able to anticipate the effect on the mean and standard deviation.</p>
<p>Plastic containers are manufactured by two machines whose settings are often different. The stacked dot plots below describe component lengths from the three machines. What are the mean and standard deviation of the combined data set?</p>

<p>
[[jsxgraph]] 
// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true
});

//Create points of Data Set 2
var lx2 = {#dorx2#};
var ly2 = {#dory2#};
var i = 0;
for (i = 0; i <= lx2.length - 1; i++) {
  p[i] = board.create('point', [lx2[i], ly2[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 2
var txt = board.create('text', [{#xmin#}+1, {#maxfreq2#}/2+2, "Machine 2"], {
  fontSize: 20,fixed: true
});

//Create points of Data Set 1
var lx1 = {#dorx1#};
var ly1 = {#dory1#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 1
var txt = board.create('text', [{#xmin#}+1, {#maxfreq1#}/2+{#maxfreq2#} + 5, "Machine 1"], {
  fontSize: 20,fixed: true
});
[[/jsxgraph]]</p>


<p>What is the combined mean of the both datasets?
[[input:ans1]] [[validation:ans1]]</p>

<p>What is the combined standard deviation  of the both datasets
[[input:ans2]] [[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The outlier will increase the standard deviation, but not by this much. (Think about x̅ ± s and x̅ ± 2s.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[mean1:rand(50)+10;
aa:rand([-1,0,1]);
mean2:mean1+(rand(11)+5)*aa;
sd1:rand(5)+5;
bb:rand([1/3,1/2,1,2,3]);
sd2:sd1*bb
ds1:rand(70)+30;
ds2:ds1;
dun1:makelist(round(float(quantile_normal(rand(1.0),mean1,sd1))),i,1,ds1);
dorx1:sort(dun1);
dun2:makelist(round(float(quantile_normal(rand(1.0),mean2,sd2))),i,1,ds2);
dorx2:sort(dun2);
ta1:mean(dun1);
ta2:round(float(sqrt(dun1+dun2)));

maxfreq1:0;
k:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;
maxfreq2:0;
k:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (k:k+1, if maxfreq2 < k then maxfreq2:k) else k:1;

xmin:if dorx1[1]<=dorx2[1] then dorx1[1]-2 else dorx2[1]-2;
xmax:if dorx1[ds1]>=dorx2[ds2] then dorx1[ds1]+2 else dorx2[ds2]+2;
ymin:-4;
ymax: maxfreq1+maxfreq2+5;

j:1;
dory1:makelist(0,i,1,ds1);
dory1[1]:1+maxfreq2+3;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory1[i]:j+maxfreq2+3) else (j:1, dory1[i]:j+maxfreq2+3);
j:1;
dory2:makelist(0,i,1,ds2);
dory2[1]:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (j:j+1, dory2[i]:j) else (j:1, dory2[i]:j);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>sd1 = {@sd1@}, sd2 = {@sd2@}.
mean1 = {@mean1@}, mean2 = {@mean2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.5</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
