<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating Probability Mass Function Values for a Discrete Random Variable</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p></p><br>
<p>A discrete random variable \(X\) has probability mass function</p>
<table width="500" cellspacing="5" cellpadding="5" border="1">
    <tbody>
        <tr>
            <td valign="center">
                <p>\( \text{X} \)</p>
            </td>
            <td valign="top">
                <p>\(0\)</p>
            </td>
            <td valign="center">
                <p>\(1\)</p>
            </td>
            <td valign="top">
                <p>\(2\)</p>
            </td>
        </tr>
        <tr>
            <td valign="center">
<p>\( \text{f(x)=p(X=x)}\)</p>
            </td>
            <td valign="top">
                <p>\({@(aa-variable)@}\)</p>
            </td>
            <td valign="top">
                <p>\({@variable@}\)</p>
            </td>
            <td valign="top">
                <p>\({@(bb+variable)@}\)</p>
            </td>
        </tr>
    </tbody>
</table>
<p><br></p>

<p>Determine the value of \({@variable@}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Use the value of \({@variable@}\) to complete the table below:</p>
<p>
<table width="500" cellspacing="5" cellpadding="5" border="1">
    <tbody>
        <tr>
            <td valign="center">
                <p>\( \text{X} \)</p>
            </td>
            <td valign="top">
                <p>\(0\)</p>
            </td>
            <td valign="center">
                <p>\(1\)</p>
            </td>
            <td valign="top">
                <p>\(2\)</p>
            </td>
        </tr>
        <tr>
            <td valign="center">
<p>\( \text{f(x)=p(X=x)}\)</p>
            </td>
            <td valign="top">
                <p>\({@(aa-variable)@}=\)[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
            </td>
            <td valign="top">
                <p>\({@variable@}=\)[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
            </td>
            <td valign="top">
                <p>\({@(bb+variable)@}=\)[[input:ans4]] [[validation:ans4]] [[feedback:prt4]] </p>
            </td>
        </tr>
    </tbody>
</table>
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked Solution:</p>
<p>The probability mass function denoted as \(PMF\) is a probability that gives probabilities of the possible values for a random variable.</p>
<p>Therefore, in this case \(X\) is a discrete random variable and its range is countable which is equal to 1.</p>
<p><b>Properties of PMF</b></p>
<p>\( \bullet \quad 0 \leq P_{X}(x) \leq 1  \quad \text{for all x}\)</p>
<p>\( \bullet \quad \sum_{x \in R_{x}}  P_{X}(x)=1 \)</p>
<p>\( \bullet \quad \text{for any set} A \subset R_{x}, P(X \in A)- \sum_{x \in A} P_{X}(x)\)</p>
<p>The value of \({@variable@}\) will be calculated as follows;</p> 
<p>
\(
\begin{align}
f(x) &= p(X=x)=1 \\
f(x) &= p(X=0,1,2)=p(X=1)+p(X=2)+p(X=0)=1 \\
      &({@(aa-variable)@})+{@variable@}+({@(bb+variable)@})=1 \\
      &{@aa@}+{@bb@}-{@variable@}+{@variable@}+{@variable@}=1 \\
      &{@(aa+bb)@}+{@(-variable+variable+variable)@}=1 \\
      &{@(-variable+variable+variable)@}=1-{@(aa+bb)@} \\
      &{@(-variable+variable+variable)@}={@1-(aa+bb)@}
\end{align}
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Question variables*/
variable:rand([a,b,c,d,g,h,k]); /*variable*/
aa:(rand([5,6,7,8])/10);/*f(x)=x0-variable*/
bb:(rand([1,2])/10);/*f(x)=x2-variable*/
func:(aa-variable)+variable+(bb+variable); /*PMF*/
ans:solve(func=1)[1];  

/*Model answer*/
ta:rhs(ans);
ta1:aa-ta;
ta2:ta;
ta3:bb+ta;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@variable@}={@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checking for correct answer*/
/*node 2:mistake:Subtracting the {@variable @} from f(x2)*/
mistake_2:(aa-variable)+variable+(bb-variable); /*PMF*/
sans2:rhs(solve(mistake_2=1)[1]);
/*node 3:mistake:Adding the {@variable@} from f(x0)*/
mistake_3:(aa+variable)+variable+(bb+variable); /*PMF*/
sans3:rhs(solve(mistake_3=1)[1]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have subtracted the value of  \({@variable@}\) from the probability of \(f(x=2)\) instead you should have added.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have added the value of \({@variable@}\) to the probability of \(f(x=0)\) instead you should have subtracted..</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checking for correct answer*/
/*node 2:mistake:Finding a negative probability for f(x=0)*/
mistake_2:aa-ta;
sans2:-1*(mistake_2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your solution is less than \(0\). Recall that the probability of \(x\) must be satisfying the rule \( \quad 0 \leq P_{X}(x) \leq 1  \quad \text{for all x}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1770743276</deployedseed>
    <deployedseed>1750177612</deployedseed>
    <deployedseed>537344341</deployedseed>
    <deployedseed>2086478064</deployedseed>
    <deployedseed>204460767</deployedseed>
    <deployedseed>267589385</deployedseed>
    <deployedseed>390330861</deployedseed>
    <deployedseed>1345933998</deployedseed>
    <deployedseed>643065781</deployedseed>
    <deployedseed>289711453</deployedseed>
    <deployedseed>1605829798</deployedseed>
    <deployedseed>1135218151</deployedseed>
    <deployedseed>1111518662</deployedseed>
    <deployedseed>1208125230</deployedseed>
    <deployedseed>441458561</deployedseed>
    <deployedseed>1059751137</deployedseed>
    <deployedseed>553872651</deployedseed>
    <deployedseed>758984779</deployedseed>
    <deployedseed>789956160</deployedseed>
  </question>
</quiz>
