<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Counting Words with Specific Letter Constraints</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>How many <b>{@word1@}</b>-letter words can be formed from the letters of the word {@word@} if
the word must begin with \(A\) and end with \(A\) and no letter may be repeated.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>We need to form  a {@word1@} letter word from the word  {@word@} where the first and last letters are already selected </p>
<p>From the remaining {@wordlength1-2@} letters we choose the three middle letters:</p>
[[ if test='wordlength1=6' ]]
<p>\(\displaystyle 4 \times 3 \times 2 = \frac {4!}{2!}={@factorial(4)/factorial(2)@}\)</p>
[[ elif test='wordlength1=7 ' ]]
<p>\(\displaystyle 5 \times 4 \times 3 = \frac {5!}{2!}={@factorial(5)/factorial(2)@}\)</p>
[[ else ]]
<p>\(\displaystyle 6 \times 5 \times 4 = \frac {6!}{2!}={@factorial(6)/factorial(2)@}\)</p>
[[/ if]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[wordlength1:rand([6,7,8]);
word:if wordlength1 = 6 then rand(["AMOEBA", "ATHENA"]) elseif wordlength1=7 then rand(["ALGEBRA","ALBERTA","AGNOSIA"]) else "ANGELICA";
word1:if wordlength1=6  then "four" elseif wordlength1=7 then "five" else "six";
ta:if wordlength1=6 then factorial(4)/factorial(2) elseif wordlength1 =7 then factorial(5)/factorial(2) else factorial(6)/factorial(2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@word@}={@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Georg Osang</p>
<p id="reviewer"></p>
<p id="description">This question requires the student to explore the possibilities of forming different words from given sets of letters in the given words. Students are tasked with determining the total number of distinct words that can be created by rearranging the letters in each word.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>207754745</deployedseed>
    <deployedseed>566649067</deployedseed>
    <deployedseed>1449800472</deployedseed>
    <deployedseed>668279899</deployedseed>
    <deployedseed>583984060</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
