<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Counting Combinations of Toffees, Chestnuts, and Berries</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>In how many ways can a boy select \({@toffee@}\) toffees, \({@chestnut@}\) chestnuts and \({@berries@}\) berries from a jar containing \({@total_toffs@}\) toffees, \({@chestnuts@}\) chestnuts and \({@total_berries@}\) berries.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> To find the total number of ways to select toffees, chestnuts, and berries together, we multiply these individual counts:</p>
<p>\(^{@total_toffs@}C_{@toffee@} \times   {}^{@chestnuts@}C_{@chestnut@} \times  {}^{@total_berries@}C_{@berries@}\)</p>
<p>\(^{@total_toffs@}C_{@toffee@}\) represents the number of ways to choose \({@toffee@}\) toffees out of \({@total_toffs@}\) available toffees.</p>
<p>\(^{@chestnuts@}C_{@chestnut@}\) represents the number of ways to choose \({@chestnut@}\) chestnuts out of \({@chestnuts@}\) available chestnuts</p>
<p>\(^{@total_berries@}C_{@berries@}\) represents the number of ways to choose \({@berries@}\) berries out of \({@total_berries@}\) available berries.</p>
<p>Therefore,</p>
<p>\({}^{@total_toffs@}C_{@toffee@} \times   {}^{@chestnuts@}C_{@chestnut@} \times  {}^{@total_berries@}C_{@berries@}={@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>0</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*prob*/
toffee:rand(5)+2;
total_toffs:toffee+3;
chestnut:rand_with_step(1,5,1);
chestnuts:rand_with_prohib(5,9,[toffee,total_toffs]);
berries:rand_with_prohib(2,5,[toffee,chestnut]);
total_berries:rand_with_prohib(6,10,[chestnuts,total_toffs]);

/*worked solution variables*/
sol1:binomial(total_toffs,toffee);
sol2:binomial(chestnuts,chestnut);
sol3:binomial(total_berries,berries);
ta:sol1*sol2*sol3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@toffee@}\) toffees, \({@chestnut@}\) chestnuts and \({@berries@}\) berries from a jar containing \({@total_toffs@}\) toffees, \({@chestnuts@}\) chestnuts and \({@total_berries@}\) berries.</p>
<p>\(^{@total_toffs@}C_{@toffee@} \times {}^{@chestnuts@}C_{@chestnut@} \times {}^{@total_berries@}C_{@berries@}={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary</p>
<p id="concept"></p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question involves finding the number of ways that a given number of berries, chestnuts and toffees can be chosen from a given total number of berries,chestnuts and toffees respectfully.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>231475533</deployedseed>
    <deployedseed>475236232</deployedseed>
    <deployedseed>552027102</deployedseed>
    <deployedseed>533543206</deployedseed>
    <deployedseed>963877927</deployedseed>
    <deployedseed>2131589029</deployedseed>
    <deployedseed>838404649</deployedseed>
    <deployedseed>1201480683</deployedseed>
    <deployedseed>2047815426</deployedseed>
    <deployedseed>1183617895</deployedseed>
    <deployedseed>1879511093</deployedseed>
    <deployedseed>1609302652</deployedseed>
    <deployedseed>1522256386</deployedseed>
    <deployedseed>843001881</deployedseed>
    <deployedseed>1386907210</deployedseed>
    <deployedseed>221779885</deployedseed>
    <deployedseed>603152025</deployedseed>
    <deployedseed>1370028441</deployedseed>
    <deployedseed>2019611181</deployedseed>
    <deployedseed>1785697331</deployedseed>
    <deployedseed>1717387303</deployedseed>
    <deployedseed>1205909051</deployedseed>
    <deployedseed>1249405359</deployedseed>
    <deployedseed>188947510</deployedseed>
    <deployedseed>489212305</deployedseed>
    <deployedseed>552877212</deployedseed>
    <deployedseed>1826192751</deployedseed>
    <deployedseed>1084216289</deployedseed>
    <deployedseed>1313771789</deployedseed>
    <deployedseed>152955731</deployedseed>
    <deployedseed>1664765929</deployedseed>
    <deployedseed>660961639</deployedseed>
    <deployedseed>990562361</deployedseed>
    <deployedseed>1828520830</deployedseed>
    <deployedseed>1582290880</deployedseed>
    <deployedseed>50999941</deployedseed>
    <deployedseed>1479820180</deployedseed>
    <deployedseed>2042123437</deployedseed>
    <deployedseed>514710510</deployedseed>
    <deployedseed>1520667368</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
