<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Creating an Ordered Stem and Leaf Diagram for Party Guest Ages</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The data below shows the ages of people at a party.</p>
<p>\({@c3@}, {@b1@},{@a2@}, {@b3@},{@c1@}, {@c2@},{@a3@}, {@b2@},{@a1@}\)</p>
<p>Draw an ordered stem and leaf diagram for the data.</p>

<p>
<table border="solid">
  <tr>
    <th width="50%", style="text-align:center">Stem</th>
    <th width="50%", style="text-align:center">Leaf</th>
  </tr>
  <tr>
    <th width="50%", style="text-align:center">\(3\)</th>
    <th width="50%", style="text-align:center">[[input:ans1]]  [[input:ans2]]  [[input:ans3]]  [[validation:ans1]][[validation:ans2]] [[validation:ans3]] [[feedback:prt1]] [[feedback:prt2]][[feedback:prt3]]</th>
  </tr>
  <tr>
    <th width="50%", style="text-align:center">\(4\)</th>
    <th width="50%", style="text-align:center">[[input:ans4]]  [[input:ans5]]  [[input:ans6]] [[validation:ans4]] [[validation:ans5]]  [[validation:ans6]] [[feedback:prt4]][[feedback:prt5]][[feedback:prt6]]</th>
  </tr>
  <tr>
    <th width="50%", style="text-align:center">\(5\)</th>
    <th width="50%", style="text-align:center">[[input:ans7]]  [[input:ans8]]  [[input:ans9]]  [[validation:ans7]] [[validation:ans8]] [[validation:ans9]] [[feedback:prt7]][[feedback:prt8]] [[feedback:prt9]]</th>
  </tr>
</table>
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Arranging in ascending order: \({@sorted_list4@}\)</p>
<p>
<table border="solid">
  <tr>
    <th width="12.5%" style="text-align:center">Stem</th>
    <th width="12.5%" style="text-align:center">Leaf</th>
  </tr>
  <tr>
    <td width="12.5%" style="text-align:center">3</td>
    <td width="12.5%" style="text-align:center" colspan="2">{@ta1@}  {@ta2@} {@ta3@}</td>
  </tr>
  <tr>
    <td width="12.5%" style="text-align:center">4</td>
    <td width="12.5%" style="text-align:center" colspan="2">{@ta4@}   {@ta5@}  {@ta6@}</td>
  </tr>
  <tr>
    <td width="12.5%" style="text-align:center">5</td>
    <td width="12.5%" style="text-align:center" colspan="2">{@ta7@}  {@ta8@}  {@ta9@}</td>
  </tr>
</table>
</p>
<p>Key 3|{@ta1@} means {@sorted_list1[1]@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>a1:rand_with_step(30,39,1);
a2:rand_with_prohib(30,39,[a1]);
a3:rand_with_prohib(30,39,[a1,a2]);
b1:rand_with_step(40,49,1);
b2:rand_with_prohib(40,49,[b1]);
b3:rand_with_prohib(40,49,[b1,b2]);
c1:rand_with_step(50,59,1);
c2:rand_with_prohib(50,59,[c1]);
c3:rand_with_prohib(50,59,[c1,c2]);
my_list1: [a1, a2,a3];
sorted_list1: sort(my_list1);
ta1:mod(sorted_list1[1], 10);
ta2:mod(sorted_list1[2], 10);
ta3:mod(sorted_list1[3], 10);
my_list2: [b1, b2,b3];
sorted_list2: sort(my_list2);
ta4:mod(sorted_list2[1], 10);
ta5:mod(sorted_list2[2], 10);
ta6:mod(sorted_list2[3], 10);
my_list3: [c1,c2,c3];
sorted_list3: sort(my_list3);
ta7:mod(sorted_list3[1], 10);
ta8:mod(sorted_list3[2], 10);
ta9:mod(sorted_list3[3], 10);
my_list4: [a1, a2,a3,b1,b2,b3,c1,c2,c3];
sorted_list4: sort(my_list4);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Arranging in ascending order: \({@sorted_list4@}\)</p>
<p>
<table border="solid">
  <tr>
    <th width="12.5%", style="text-align:center">Stem</th>
    <th width="12.5%", style="text-align:center">Leaf</th>
  </tr>
  <tr>
    <th width="12.5%", style="text-align:center">\(3\)</th>
    <th width="12.5%", style="text-align:center">\({@ta1@}\, \, {@ta2@}\, \, {@ta3@} \)</th>
  </tr>
  <tr>
    <th width="12.5%", style="text-align:center">\(4\)</th>
    <th width="12.5%", style="text-align:center">\({@ta4@}\, \, {@ta5@}\, \, {@ta6@} \)</th>
  </tr>
  <tr>
    <th width="12.5%", style="text-align:center">\(5\)</th>
    <th width="12.5%", style="text-align:center">\({@ta7@}\, \, {@ta8@}\, \, {@ta9@} \)</th>
  </tr>
</table>
</p>
<p>Key \(3|{@ta1@}\) means \(3{@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description"> This question requires the student to organize and display the ages of people at a party using a stem and leaf diagram in ascending order for easy data analysis.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>ta9</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2088145598</deployedseed>
    <deployedseed>1697479738</deployedseed>
    <deployedseed>1567764222</deployedseed>
    <deployedseed>1737686448</deployedseed>
    <deployedseed>1408492743</deployedseed>
    <deployedseed>679263991</deployedseed>
    <deployedseed>623860955</deployedseed>
    <deployedseed>824893415</deployedseed>
    <deployedseed>2026986331</deployedseed>
    <deployedseed>1259068613</deployedseed>
    <deployedseed>1592885276</deployedseed>
    <deployedseed>1717472892</deployedseed>
    <deployedseed>277466458</deployedseed>
    <deployedseed>1506136160</deployedseed>
    <deployedseed>778175262</deployedseed>
    <deployedseed>2144267584</deployedseed>
    <deployedseed>716575669</deployedseed>
    <deployedseed>880941781</deployedseed>
    <deployedseed>356467126</deployedseed>
    <deployedseed>1113742406</deployedseed>
    <deployedseed>164203961</deployedseed>
    <deployedseed>1624630019</deployedseed>
    <deployedseed>736650007</deployedseed>
    <deployedseed>1357242886</deployedseed>
    <deployedseed>755516339</deployedseed>
    <deployedseed>703842374</deployedseed>
    <deployedseed>28192433</deployedseed>
    <deployedseed>1919273268</deployedseed>
    <deployedseed>1936172657</deployedseed>
    <deployedseed>1570462080</deployedseed>
    <deployedseed>1196338254</deployedseed>
    <deployedseed>369845917</deployedseed>
    <deployedseed>194730961</deployedseed>
    <deployedseed>1501769723</deployedseed>
    <deployedseed>993799311</deployedseed>
    <deployedseed>1598874300</deployedseed>
    <deployedseed>1105598589</deployedseed>
    <deployedseed>1511507310</deployedseed>
    <deployedseed>431801533</deployedseed>
    <deployedseed>1971921232</deployedseed>
    <deployedseed>1887213511</deployedseed>
    <deployedseed>1479852109</deployedseed>
    <deployedseed>1411823018</deployedseed>
    <deployedseed>302840463</deployedseed>
    <deployedseed>1773098787</deployedseed>
    <deployedseed>1751476512</deployedseed>
    <deployedseed>1300627972</deployedseed>
    <deployedseed>1098958696</deployedseed>
    <deployedseed>70713904</deployedseed>
    <deployedseed>620425941</deployedseed>
    <deployedseed>913866760</deployedseed>
    <deployedseed>604522796</deployedseed>
    <deployedseed>507455197</deployedseed>
    <deployedseed>643771758</deployedseed>
    <deployedseed>1718845440</deployedseed>
    <deployedseed>770736601</deployedseed>
    <deployedseed>710088429</deployedseed>
    <deployedseed>1641832578</deployedseed>
    <deployedseed>963015151</deployedseed>
    <deployedseed>56451505</deployedseed>
    <deployedseed>2114015771</deployedseed>
    <deployedseed>1383093782</deployedseed>
    <deployedseed>1385744004</deployedseed>
    <deployedseed>525390669</deployedseed>
    <deployedseed>1464751634</deployedseed>
    <deployedseed>1161165753</deployedseed>
    <deployedseed>129476148</deployedseed>
    <deployedseed>883026633</deployedseed>
    <deployedseed>1317914748</deployedseed>
    <deployedseed>1312616998</deployedseed>
    <deployedseed>210644868</deployedseed>
    <deployedseed>347676362</deployedseed>
    <deployedseed>1128799022</deployedseed>
    <deployedseed>485523831</deployedseed>
    <deployedseed>1490104873</deployedseed>
    <deployedseed>1469788624</deployedseed>
    <deployedseed>505820029</deployedseed>
    <deployedseed>1708191718</deployedseed>
    <deployedseed>135226666</deployedseed>
    <deployedseed>1908863698</deployedseed>
    <deployedseed>1541713184</deployedseed>
    <deployedseed>1356414629</deployedseed>
    <deployedseed>343361266</deployedseed>
    <deployedseed>1146396681</deployedseed>
    <deployedseed>300456287</deployedseed>
    <deployedseed>1472558942</deployedseed>
    <deployedseed>21153210</deployedseed>
    <deployedseed>1232822831</deployedseed>
    <deployedseed>2085475715</deployedseed>
    <deployedseed>1684556869</deployedseed>
    <deployedseed>860270399</deployedseed>
    <deployedseed>2118105458</deployedseed>
    <deployedseed>1010483034</deployedseed>
    <deployedseed>1608850871</deployedseed>
    <deployedseed>179640404</deployedseed>
    <deployedseed>313708406</deployedseed>
    <deployedseed>2015725112</deployedseed>
    <deployedseed>1762249894</deployedseed>
    <deployedseed>2040022910</deployedseed>
    <deployedseed>1716622988</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>ta9</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
