<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Data Presentation -Stem and Leaf Diagram-q3</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p></p>
<p>The data shows the age of patients in a certain clinic on a particular day. Construct a stem and leaf plot for the data.</p>
<p>
<table border="solid">
  <tr>
  

 <th width="4%", style="text-align:center">{#s_1#}{#l_1#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_3#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_4#}</th>
    <th width="4%", style="text-align:center">{#s_5#}{#l_5#}</th>
<th width="4%", style="text-align:center">{#s_6#}{#l_6#}</th>
   <th width="4%", style="text-align:center">{#s_7#}{#l_7#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_5#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_1#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_1#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_2#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
     <th width="4%", style="text-align:center">{#s_6#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_2#}</th>
  <th width="4%", style="text-align:center">{#s_7#}{#l_1#}</th>

  </tr>
  <tr>
 <th width="4%", style="text-align:center">{#s_2#}{#l_7#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_4#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_7#}{#l_2#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_1#}</th>
     <th width="4%", style="text-align:center">{#s_2#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_3#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_6#}</th>
<th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_7#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_4#}{#l_3#}</th>
  <th width="4%", style="text-align:center">{#s_4#}{#l_1#}</th>

  </tr>

</table>
</p>


<p>The above data should be organized in a stem and leaf diagram below:</p>
<p>
<!DOCTYPE html>
<html>

<head>
    <title>Stem and Leaf Diagram</title>
    <style>
        .stem-and-leaf {
            display: inline-block;
            font-family: monospace;
            font-size: 50px;
            line-height: 5;
        }
    </style>
</head>

<body>
    <h4>Stem and Leaf Diagram</h4>
    <div class="stem-and-leaf">
        <table border="1">
            <tr>
                <th>Stem</th>
                <th>Leaf</th>
            </tr>
            <tr>
                <td>{#s_1#}</td>
                <td border="1" colspan="2">{#l_1#} {#l_3#}[[input:ans1]]  [[validation:ans1]][[feedback:prt1]]</td>
            </tr>
            <tr>
                <td>{#s_2#}</td>
                <td border="1" colspan="2">[[input:ans2]] {#l_2#} [[input:ans3]]  {#l_4#}  {#l_5#}  {#l_6#}  {#l_7#}[[validation:ans2]][[feedback:prt2]][[validation:ans3]][[feedback:prt3]]</td>
            </tr>
            <tr>
                <td>{#s_3#}</td>
                <td border="1">{#l_2#}  {#l_3#}  [[input:ans4]]  {#l_5#}  {#l_6#}  {#l_7#}  [[validation:ans4]][[feedback:prt4]]</td>
            </tr>
            <tr>
                <td>{#s_4#}</td>
                <td border="1">{#l_1#}  {#l_2#}  {#l_3#}  {#l_4#}  [[input:ans5]]  {#l_7#}  [[validation:ans5]][[feedback:prt5]]</td>
            </tr>
            <tr>
                <td>{#s_5#}</td>
                <td border="1">{#l_2#}  {#l_3#}  {#l_4#}  {#l_5#} [[input:ans6]] [[validation:ans6]][[feedback:prt6]]</td>
            </tr>
            <tr>
                <td>{#s_6#}</td>
                <td border="1">{#l_2#}  {#l_3#}  {#l_6#}  [[input:ans7]] [[validation:ans7]][[feedback:prt7]]</td>
            </tr>
            <tr>
                <td>{#s_7#}</td>
                <td border="1">{#l_1#}  {#l_2#}  [[input:ans8]] [[validation:ans8]][[feedback:prt8]]</td>
            </tr>
         
        </table>
    </div>
</body>

</html>
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p></p>
<p>Given the data below we can construct a stem and leaf diagram.</p>
<p>
<table border="solid">
  <tr>
  

 <th width="4%", style="text-align:center">{#s_1#}{#l_1#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_3#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_4#}</th>
    <th width="4%", style="text-align:center">{#s_5#}{#l_5#}</th>
<th width="4%", style="text-align:center">{#s_6#}{#l_6#}</th>
   <th width="4%", style="text-align:center">{#s_7#}{#l_7#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_5#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_1#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_1#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_2#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
     <th width="4%", style="text-align:center">{#s_6#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_2#}</th>
  <th width="4%", style="text-align:center">{#s_7#}{#l_1#}</th>

  </tr>
  <tr>
 <th width="4%", style="text-align:center">{#s_2#}{#l_7#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_4#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_7#}{#l_2#}</th>
   <th width="4%", style="text-align:center">{#s_2#}{#l_1#}</th>
     <th width="4%", style="text-align:center">{#s_2#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_3#}{#l_3#}</th>
<th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_5#}{#l_7#}</th>
    <th width="4%", style="text-align:center">{#s_6#}{#l_6#}</th>
<th width="4%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="4%", style="text-align:center">{#s_3#}{#l_5#}</th>
     <th width="4%", style="text-align:center">{#s_3#}{#l_7#}</th>
     <th width="4%", style="text-align:center">{#s_4#}{#l_2#}</th>
    <th width="4%", style="text-align:center">{#s_4#}{#l_3#}</th>
  <th width="4%", style="text-align:center">{#s_4#}{#l_1#}</th>

  </tr>

</table>
</p>


<p>
<!DOCTYPE html>
<html>

<head>
    <title>Stem and Leaf Diagram</title>
    <style>
        .stem-and-leaf {
            display: inline-block;
            font-family: monospace;
            font-size: 50px;
            line-height: 5;
        }
    </style>
</head>

<body>
    <h4>Stem and Leaf Diagram</h4>
    <div class="stem-and-leaf">
        <table border="1">
            <tr>
                <th>Stem</th>
                <th>Leaf</th>
            </tr>
            <tr>
                <td>{#s_1#}</td>
                <td border="1">{#l_1#}  {#l_3#}  {#l_4#}</td>
            </tr>
            <tr>
                <td>{#s_2#}</td>
                <td border="1">{#l_1#}  {#l_2#}  {#l_3#}  {#l_4#}  {#l_5#}  {#l_6#}  {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_3#}</td>
                <td border="1">{#l_2#}  {#l_3#}  {#l_4#}  {#l_5#}  {#l_6#}  {#l_7#}  </td>
            </tr>
            <tr>
                <td>{#s_4#}</td>
                <td border="1">{#l_1#}  {#l_2#}  {#l_3#}  {#l_4#}  {#l_5#}  {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_5#}</td>
                <td border="1">{#l_2#}  {#l_3#}  {#l_4#}  {#l_5#}  {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_6#}</td>
                <td border="1">{#l_2#}  {#l_3#}  {#l_6#}  {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_7#}</td>
                <td border="1">{#l_1#}  {#l_2#}  {#l_7#}</td>
            </tr>
         
        </table>
    </div>
</body>

</html>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*STEM AND LEAF DIAGRAMS*/
s_1:rand(3);
s_2:s_1+1;
s_3:s_1+2;
s_4:s_1+3;
s_5:s_1+4;
s_6:s_1+5;
s_7:s_1+6;

/*leafs*/
l_1:rand(3)+1;
l_2:l_1+1;
l_3:l_1+2;
l_4:rand(3)+1;
l_5:l_1+3;
l_6:l_1+5;
l_7:l_1+6;
/*Model answers*/
ta1:l_4;
ta2:l_1;
ta3:l_3;
ta4:l_4;
ta5:l_5;
ta6:l_7 ;
ta7: l_7;
ta8: l_7;

</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Given the data below we can construct a stem and leaf diagram.</p>
<p>
<table border="solid">
  <tr>
  

 <th width="15%", style="text-align:center">{#s_1#}{#l_1#}</th>
   <th width="15%", style="text-align:center">{#s_2#}{#l_2#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_3#}</th>
     <th width="15%", style="text-align:center">{#s_4#}{#l_4#}</th>
    <th width="15%", style="text-align:center">{#s_5#}{#l_5#}</th>
<th width="15%", style="text-align:center">{#s_6#}{#l_6#}</th>
   <th width="15%", style="text-align:center">{#s_7#}{#l_7#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="15%", style="text-align:center">{#s_5#}{#l_2#}</th>
    <th width="15%", style="text-align:center">{#s_5#}{#l_3#}</th>
<th width="15%", style="text-align:center">{#s_1#}{#l_4#}</th>
   <th width="15%", style="text-align:center">{#s_2#}{#l_6#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="15%", style="text-align:center">{#s_4#}{#l_7#}</th>
    <th width="15%", style="text-align:center">{#s_1#}{#l_3#}</th>
<th width="15%", style="text-align:center">{#s_2#}{#l_4#}</th>
   <th width="15%", style="text-align:center">{#s_3#}{#l_2#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_4#}</th>
     <th width="15%", style="text-align:center">{#s_6#}{#l_7#}</th>
    <th width="15%", style="text-align:center">{#s_6#}{#l_2#}</th>
  <th width="15%", style="text-align:center">{#s_7#}{#l_1#}</th>

  </tr>
  <tr>
 <th width="15%", style="text-align:center">{#s_2#}{#l_7#}</th>
   <th width="15%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="15%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="15%", style="text-align:center">{#s_5#}{#l_4#}</th>
    <th width="15%", style="text-align:center">{#s_6#}{#l_3#}</th>
<th width="15%", style="text-align:center">{#s_7#}{#l_2#}</th>
   <th width="15%", style="text-align:center">{#s_2#}{#l_1#}</th>
     <th width="15%", style="text-align:center">{#s_2#}{#l_5#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_2#}</th>
    <th width="15%", style="text-align:center">{#s_3#}{#l_3#}</th>
<th width="15%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="15%", style="text-align:center">{#s_3#}{#l_6#}</th>
     <th width="15%", style="text-align:center">{#s_4#}{#l_5#}</th>
     <th width="15%", style="text-align:center">{#s_5#}{#l_7#}</th>
    <th width="15%", style="text-align:center">{#s_6#}{#l_6#}</th>
<th width="15%", style="text-align:center">{#s_3#}{#l_4#}</th>
   <th width="15%", style="text-align:center">{#s_3#}{#l_5#}</th>
     <th width="15%", style="text-align:center">{#s_3#}{#l_7#}</th>
     <th width="15%", style="text-align:center">{#s_4#}{#l_2#}</th>
    <th width="15%", style="text-align:center">{#s_4#}{#l_3#}</th>
  <th width="15%", style="text-align:center">{#s_4#}{#l_1#}</th>

  </tr>

</table>
</p>

<p>
<!DOCTYPE html>
<html>

<head>
    <title>Stem and Leaf Diagram</title>
    <style>
        .stem-and-leaf {
            display: inline-block;
            font-family: monospace;
            font-size: 50px;
            line-height: 5;
        }
    </style>
</head>

<body>
    <h4>Stem and Leaf Diagram</h4>
    <div class="stem-and-leaf">
        <table border="1">
            <tr>
                <th>Stem</th>
                <th>Leaf</th>
            </tr>
            <tr>
                <td>{#s_1#}</td>
                <td border="1">{#l_1#} {#l_3#}  {#l_4#}</td>
            </tr>
            <tr>
                <td>{#s_2#}</td>
                <td border="1">{#l_1#} {#l_2#}{#l_3#}{#l_4#}{#l_5#} {#l_6#}{#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_3#}</td>
                <td border="1">{#l_2#}{#l_3#} {#l_4#} {#l_5#} {#l_6#} {#l_7#}  </td>
            </tr>
            <tr>
                <td>{#s_4#}</td>
                <td border="1">{#l_1#}{#l_2#}{#l_3#}{#l_4#}{#l_5#} {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_5#}</td>
                <td border="1">{#l_2#} {#l_3#} {#l_4#} {#l_5#} {#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_6#}</td>
                <td border="1">{#l_2#}{#l_3#} {#l_6#}{#l_7#}</td>
            </tr>
            <tr>
                <td>{#s_7#}</td>
                <td border="1">{#l_1#}{#l_2#}{#l_7#}</td>
            </tr>
         
        </table>
    </div>
</body>

</html>
</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo<p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">The question focuses on using a given data to construct a stem and leaf diagram. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans2,ans3};
tans:{ta2,ta3};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if interchanged the correct values</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans2,ans3};
tans:{ta2,ta3};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if interchanged the correct values</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans5,ans6,ans7};
tans:{ta5,ta6,ta7};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if interchaged the order but still correct </description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans5,ans6,ans7};
tans:{ta5,ta6,ta7};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if interchaged the order but still correct </description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans5,ans6,ans7};
tans:{ta5,ta6,ta7};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if interchaged the order but still correct </description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1746873875</deployedseed>
    <deployedseed>1199792085</deployedseed>
    <deployedseed>25744329</deployedseed>
    <deployedseed>1588384591</deployedseed>
    <deployedseed>2039378864</deployedseed>
    <deployedseed>1490766366</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
