<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining Disease Probability from Positive Test Results in a Population</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A test for a certain rare disease is assumed to be correct \({@dis_test@}\)% of the time: if a person has the disease, the test results are positive with probability \({@pos_res@}\), and if the person does not have the disease, the test results are negative with probability \({@neg_res@}\). A random person drawn from a certain population has probability \({@hav_dis@}\) of having the disease. Given that the person just tested positive, what is the probability of having the disease?</p>
<p><b>Note:</b> Give your answer in exact form or correct to \(4\) decimal places where necessary.</p>
<p>\(P_r(A|B)=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>If \(A\) is the event that the person has the disease, and \(B\) is the event that the test results are positive, the desired probability \(P_r(A|B)\), is</p>
<p>\(\displaystyle P_r(A|B)=\frac{P_r(A)P_r(B|A)}{P_r(A)P_r(B|A)+P_r(A')P_r(B|A')}\)</p>
<p>\(\displaystyle P_r(A|B)=\frac{{@hav_dis@} \times {@pos_res@} }{({@hav_dis@} \times {@pos_res@})+ ({@(1-hav_dis)@} \times {@(1-pos_res)@}) }\)</p>
<p>\(P_r(A|B)={@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>dis_test:rand_with_step(51,98,1);
pos_res:(dis_test/100);
neg_res:(dis_test/100);
hav_dis:(rand_with_step(1,9,1)/1000);
ta:(hav_dis*pos_res)/((hav_dis*pos_res)+((1-hav_dis)*(1-pos_res)));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle P_r(A|B)=\frac{P_r(A)P_r(B|A)}{P_r(A)P_r(B|A)+P_r(A')P_r(B|A')}\)</p>
<p>\(\displaystyle P_r(A|B)=\frac{{@hav_dis@} \times {@pos_res@} }{{@hav_dis@} \times {@pos_res@}+ {@decimalplaces((1-hav_dis),3)@} \times {@decimalplaces((1-pos_res),2)@} }\)</p>
<p>\(P_r(A|B)={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/1Q3PyYwna8hI1fKLf72WjWteDB2rXa9n7/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question requires the student to calculate the probability of an individual having a rare disease given a positive test result, considering the test's accuracy and the prevalence of the disease in the population.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/
/*Node 2: Checks for the correct decimal answer*/
/*Node 3: Checks if the student gave too many decimal places*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct decimal answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>4</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks if the student gave too many decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.00005</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You got the correct answer but gave too many decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore/>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>376718021</deployedseed>
    <deployedseed>1662591957</deployedseed>
    <deployedseed>202226559</deployedseed>
    <deployedseed>915647749</deployedseed>
    <deployedseed>306962238</deployedseed>
    <deployedseed>1905848817</deployedseed>
    <deployedseed>275041860</deployedseed>
    <deployedseed>1376100857</deployedseed>
    <deployedseed>83666952</deployedseed>
    <deployedseed>472236544</deployedseed>
    <deployedseed>302037182</deployedseed>
    <deployedseed>919070154</deployedseed>
    <deployedseed>1259459377</deployedseed>
    <deployedseed>1965551654</deployedseed>
    <deployedseed>1661351278</deployedseed>
    <deployedseed>1607130407</deployedseed>
    <deployedseed>1591893711</deployedseed>
    <deployedseed>201154682</deployedseed>
    <deployedseed>1983651027</deployedseed>
    <deployedseed>1516950866</deployedseed>
    <deployedseed>2055057277</deployedseed>
    <deployedseed>78992487</deployedseed>
    <deployedseed>70726819</deployedseed>
    <deployedseed>1888031502</deployedseed>
    <deployedseed>2056195940</deployedseed>
    <deployedseed>240625788</deployedseed>
    <deployedseed>1222249886</deployedseed>
    <deployedseed>720485596</deployedseed>
    <deployedseed>1301540661</deployedseed>
    <deployedseed>2049084741</deployedseed>
    <deployedseed>2126605434</deployedseed>
    <deployedseed>792001359</deployedseed>
    <deployedseed>1978214835</deployedseed>
    <deployedseed>783180165</deployedseed>
    <deployedseed>399771741</deployedseed>
    <deployedseed>1913967858</deployedseed>
    <deployedseed>1662198083</deployedseed>
    <deployedseed>888759523</deployedseed>
    <deployedseed>1367932341</deployedseed>
    <deployedseed>1778396120</deployedseed>
    <deployedseed>181839488</deployedseed>
    <deployedseed>1538933115</deployedseed>
    <deployedseed>775121127</deployedseed>
    <deployedseed>518865364</deployedseed>
    <deployedseed>1173597132</deployedseed>
    <deployedseed>2135895416</deployedseed>
    <deployedseed>1074909706</deployedseed>
    <deployedseed>2009566794</deployedseed>
    <deployedseed>1978141400</deployedseed>
    <deployedseed>685343321</deployedseed>
    <deployedseed>160300419</deployedseed>
    <deployedseed>1894896598</deployedseed>
    <deployedseed>667671835</deployedseed>
    <deployedseed>591629116</deployedseed>
    <deployedseed>1621546527</deployedseed>
    <deployedseed>230402686</deployedseed>
    <deployedseed>17090149</deployedseed>
    <deployedseed>1980802266</deployedseed>
    <deployedseed>48527726</deployedseed>
    <deployedseed>1607320045</deployedseed>
    <deployedseed>1317988029</deployedseed>
    <deployedseed>267422180</deployedseed>
    <deployedseed>587838900</deployedseed>
    <deployedseed>509471368</deployedseed>
    <deployedseed>509768248</deployedseed>
    <deployedseed>2059391553</deployedseed>
    <deployedseed>634392071</deployedseed>
    <deployedseed>567354188</deployedseed>
    <deployedseed>181237132</deployedseed>
    <deployedseed>63917577</deployedseed>
    <deployedseed>1467661245</deployedseed>
    <deployedseed>825569949</deployedseed>
    <deployedseed>1656241488</deployedseed>
    <deployedseed>1201546199</deployedseed>
    <deployedseed>194626148</deployedseed>
    <deployedseed>699603985</deployedseed>
    <deployedseed>1946969836</deployedseed>
    <deployedseed>1775407219</deployedseed>
    <deployedseed>1538002772</deployedseed>
    <deployedseed>279881954</deployedseed>
    <deployedseed>97522636</deployedseed>
    <deployedseed>728421164</deployedseed>
    <deployedseed>1973342449</deployedseed>
    <deployedseed>1427080605</deployedseed>
    <deployedseed>667774468</deployedseed>
    <deployedseed>144936854</deployedseed>
    <deployedseed>1685593037</deployedseed>
    <deployedseed>908100330</deployedseed>
    <deployedseed>759687155</deployedseed>
    <deployedseed>499533456</deployedseed>
    <deployedseed>812361330</deployedseed>
    <deployedseed>1240352921</deployedseed>
    <deployedseed>275320292</deployedseed>
    <deployedseed>1772331526</deployedseed>
    <deployedseed>401105460</deployedseed>
    <deployedseed>890430802</deployedseed>
    <deployedseed>1875654333</deployedseed>
    <deployedseed>1777097738</deployedseed>
    <deployedseed>385101848</deployedseed>
    <deployedseed>2029707738</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
