<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining general fertility and total fertility rate</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>From the figures given below calculate the general fertility rate and the total fertility rate.</p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age Group</th>
    <th width="15%", style="text-align:center">No. of women</th>
     <th width="15%", style="text-align:center"> Age specific fertility rate  (per 1000)</th
  </tr>
  <tr>
    <td style="text-align:center">\({@age_group1@}-{@age_group2@}\)</td>
    <td style="text-align:center">\({@no_women1@}\)</td>
    <td style="text-align:center">\({@fertility_rate1@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group2@}-{@age_group3@}\)</td>
    <td style="text-align:center">\({@no_women2@}\)</td>
    <td style="text-align:center">\({@fertility_rate2@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group3@}-{@age_group4@}\)</td>
    <td style="text-align:center">\({@no_women3@}\)</td>
    <td style="text-align:center">\({@fertility_rate3@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group4@}-{@age_group5@}\)</td>
    <td style="text-align:center">\({@no_women4@}\)</td>
    <td style="text-align:center">\({@fertility_rate4@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group5@}-{@age_group6@}\)</td>
    <td style="text-align:center">\({@no_women5@}\)</td>
    <td style="text-align:center">\({@fertility_rate5@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group6@}-{@age_group7@}\)</td>
    <td style="text-align:center">\({@no_women6@}\)</td>
    <td style="text-align:center">\({@fertility_rate6@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group7@}-{@age_group8@}\)</td>
    <td style="text-align:center">\({@no_women7@}\)</td>
    <td style="text-align:center">\({@fertility_rate7@}\)</td>
  </tr>
</table>
</p>
<br>
<p>(Note: Leave your answers upto \(2\) dp where necessary)</p>
<p>Find, </p>
<p>a). General fertility rate \(\textbf{(GRR)}= \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b). Total fertility Rate \(\textbf{(TFR)}= \sum ASFR \times i \) where \(\text{ASFR}\) stands for Age-Specific Fertility Rate and \(i\) is the number of years in each age bracket.
<p>\(\textbf{TFR}= \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age Group</th>
    <th width="15%", style="text-align:center">No. of women</th>
     <th width="15%", style="text-align:center"> Age specific fertility rate  (per 1000)</th>
    <th width="15%", style="text-align:center"> Number of children born </th>
  </tr>
  <tr>
    <td style="text-align:center">\({@age_group1@}-{@age_group2@}\)</td>
    <td style="text-align:center">\({@no_women1@}\)</td>
    <td style="text-align:center">\({@fertility_rate1@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@no_women1@}\times {@fertility_rate1@}}{1000}={@decimalplaces(no_kids_born1,1)@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group2@}-{@age_group3@}\)</td>
    <td style="text-align:center">\({@no_women2@}\)</td>
    <td style="text-align:center">\({@fertility_rate2@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women2@}\times {@fertility_rate2@}}{1000}={@decimalplaces(no_kids_born2,1)@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group3@}-{@age_group4@}\)</td>
    <td style="text-align:center">\({@no_women3@}\)</td>
    <td style="text-align:center">\({@fertility_rate3@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women3@}\times {@fertility_rate3@}}{1000}={@decimalplaces(no_kids_born3,1)@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group4@}-{@age_group5@}\)</td>
    <td style="text-align:center">\({@no_women4@}\)</td>
    <td style="text-align:center">\({@fertility_rate4@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women4@}\times {@fertility_rate4@}}{1000}={@decimalplaces(no_kids_born4,1)@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group5@}-{@age_group6@}\)</td>
    <td style="text-align:center">\({@no_women5@}\)</td>
    <td style="text-align:center">\({@fertility_rate5@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women5@}\times {@fertility_rate5@}}{1000}={@decimalplaces(no_kids_born5,1)@}\)</td>
  </tr>
 <tr>
    <td style="text-align:center">\({@age_group6@}-{@age_group7@}\)</td>
    <td style="text-align:center">\({@no_women6@}\)</td>
    <td style="text-align:center">\({@fertility_rate6@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women6@}\times {@fertility_rate6@}}{1000}={@decimalplaces(no_kids_born6,1)@}\)</td>
 <tr>
    <td style="text-align:center">\({@age_group7@}-{@age_group8@}\)</td>
    <td style="text-align:center">\({@no_women7@}\)</td>
    <td style="text-align:center">\({@fertility_rate7@}\)</td>
   <td style="text-align:center">\(\displaystyle \frac{{@no_women7@}\times {@fertility_rate7@}}{1000}={@decimalplaces(no_kids_born7,1)@}\)</td>
  </tr>
<tr>
      <td style="text-align:center">\(\sum\)</td>
    <td style="text-align:center">total number of women \(={@sum_women@}\)</td>
    <td style="text-align:center"> total births per 1000 \(={@sum_fertility@}\)</td>
   <td style="text-align:center">sum of children \(={@sum_kids@}\)</td>
  </tr>
</table>
</p>
<br>
<p>a). General fertility rate </p>
<p>\(\textbf{GRR}= \displaystyle \frac{ \text{total number of children born}}{\text{total number of women}} \times 1000=\displaystyle \frac {{@decimalplaces(sum_kids,1)@}}{{@sum_women@}} \times 1000={@grr@}\)</p>
<br>
<p>b). Total fertility Rate \(\textbf{TFR}= \sum ASFR \times i \) where \(\text{ASFR}\) stands for Age-Specific Fertility Rate and \(i\) is the number of years in each age bracket.
<p>\(\textbf{TFR}={@sum_fertility@} \times {@age@}={@total_fertilityrate@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*general and fertility rates*/
age:5;
age_group1:rand_with_step(14,16,5);
age_group2:age_group1+age;
age_group3:age_group2+age;
age_group4:age_group3+age;
age_group5:age_group4+age;
age_group6:age_group5+age;
age_group7:age_group6+age;
age_group8:age_group7+age;

/*number of women*/
diff:rand(5)+5;
no_women:rand_with_step(80,150,10);
no_women1:rand_with_step(80,150,10);
no_women2:rand_with_step(80,150,10);
no_women3:rand_with_step(80,150,10);
no_women4:rand_with_step(80,130,10);
no_women5:rand_with_step(80,130,10);
no_women6:rand_with_step(80,120,10);
no_women7:rand_with_step(80,100,10);

/*fertility rates*/
fert_diff:rand(10)+10;
fertility_rate:rand_with_step(160,210,10);
fertility_rate1:rand_with_step(170,220,10);
fertility_rate2:rand_with_step(180,230,10);
fertility_rate3:rand_with_step(190,240,10);
fertility_rate4:rand_with_step(190,250,10);
fertility_rate5:rand_with_step(190,240,10);
fertility_rate6:rand_with_step(180,230,10);
fertility_rate7:rand_with_step(170,220,10);


/*number of children born*/
no_kids_born1:decimalplaces((no_women1*fertility_rate1)/1000,1);
no_kids_born2:decimalplaces((no_women2*fertility_rate2)/1000,1);
no_kids_born3:decimalplaces((no_women3*fertility_rate3)/1000,1);
no_kids_born4:decimalplaces((no_women4*fertility_rate4)/1000,1);
no_kids_born5:decimalplaces((no_women5*fertility_rate5)/1000,1);
no_kids_born6:decimalplaces((no_women6*fertility_rate6)/1000,1);
no_kids_born7:decimalplaces((no_women7*fertility_rate7)/1000,1);

/*worked solution variables*/
sum_women:no_women1+no_women2+no_women3+no_women4+no_women5+no_women6+no_women7;
sum_fertility:fertility_rate1+fertility_rate2+fertility_rate3+fertility_rate4+fertility_rate5+fertility_rate6+fertility_rate7;
sum_kids:no_kids_born1+no_kids_born2+no_kids_born3+no_kids_born4+no_kids_born5+no_kids_born6+no_kids_born7;
/*model answers*/
grr:decimalplaces((sum_kids/sum_women)*1000,2);/*general fertility rate*/
total_fertilityrate:sum_fertility*age;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>General fertility rate \(\textbf{GRR}= {@grr@}\)</p>
<p>Total fertility Rate \(\textbf{TFR}= \sum ASFR \times i \) where \(\text{ASFR}\) stands for Age-Specific Fertility Rate.
<p>\(\textbf{TFR}={@total_fertilityrate@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">Given a table with age group,number of women and  Age specific fertility rate  (per 1000), the question requires you to find General fertility rate and Total fertility Rate.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>grr</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>total_fertilityrate</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer on GRR </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>grr</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for n0. of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>grr</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because you did not leave to the correct number of dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for total_fertilityrate</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>total_fertilityrate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1620449862</deployedseed>
    <deployedseed>1775967909</deployedseed>
    <deployedseed>1363120725</deployedseed>
    <deployedseed>1248135504</deployedseed>
    <deployedseed>1270718576</deployedseed>
    <deployedseed>382499623</deployedseed>
    <deployedseed>991441889</deployedseed>
    <deployedseed>973018585</deployedseed>
    <deployedseed>120201689</deployedseed>
    <deployedseed>335796161</deployedseed>
    <deployedseed>1253948950</deployedseed>
    <deployedseed>415930698</deployedseed>
    <deployedseed>499037337</deployedseed>
    <deployedseed>2126963698</deployedseed>
    <deployedseed>1776322144</deployedseed>
    <deployedseed>1267141851</deployedseed>
    <deployedseed>1801797291</deployedseed>
    <deployedseed>1048774902</deployedseed>
    <deployedseed>1652258526</deployedseed>
    <deployedseed>250065166</deployedseed>
    <deployedseed>1376433220</deployedseed>
    <deployedseed>1449692201</deployedseed>
    <deployedseed>314208793</deployedseed>
    <deployedseed>1366948100</deployedseed>
    <deployedseed>964386449</deployedseed>
    <deployedseed>2030809906</deployedseed>
    <deployedseed>1658446659</deployedseed>
    <deployedseed>1473632925</deployedseed>
    <deployedseed>948147864</deployedseed>
    <deployedseed>1531379508</deployedseed>
    <deployedseed>1833004933</deployedseed>
    <deployedseed>460974929</deployedseed>
    <deployedseed>1394720159</deployedseed>
    <deployedseed>1023784190</deployedseed>
    <deployedseed>1738660545</deployedseed>
    <deployedseed>1947482790</deployedseed>
    <deployedseed>1131438476</deployedseed>
    <deployedseed>1721060006</deployedseed>
    <deployedseed>1380096308</deployedseed>
    <deployedseed>1204210367</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>grr</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>total_fertilityrate</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
