<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining factors that lowers birth rates</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>There is lower birth rate now because of </p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <i>Contraceptives </i>since it causes a delay to the return of fertility rate  for up to a year.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The correct answer is <i>More females want a career</i></p><p>Women who are more “Work centered” (i.e., work more hours and place a high value on career) are less likely to report a subsequent birth. In addition to individual preferences, the opportunity costs of having children have been identified as potential barriers for childbearing.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>The correct answer is <i>Improved health care</i></p>
<p> Improvements in child survival typically precede sustained fertility decline. Thus, interventions to improve the health of children will ultimately be followed by a decline in the birth rate.
</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "Contraceptives"],
[2, true, "More females want a career"],
[3, true, "Improved health care"]],1);


incorrect_answers:rand_selection([
[4, false, "Deaths"],
[5, false, "Miscarriages"],
[6, false, "Increase in Divorces"],
[7, false, "Baby burst"]],3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Identify a method that can be used to obtain Vital statistics</p>
<p>{@ta@}</p>
[[ if test="member(1, ta_index)" ]]
<p>Contraceptives.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>More females want a career </p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Improved health care</p>
[[/ if ]]
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/16s3aM5GC13Obwqm5uj8m-EzSj6_sR18uBoSCm5b0XIk/edit#heading=h.94wv8eijezx2</p>
<p id="reviewer"></p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question requires the student to determine factors that lowers birth rates </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct option*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct option</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>77</deployedseed>
    <deployedseed>883775737</deployedseed>
    <deployedseed>1678809143</deployedseed>
    <deployedseed>494542476</deployedseed>
    <deployedseed>302491989</deployedseed>
    <deployedseed>66152546</deployedseed>
    <deployedseed>496367606</deployedseed>
    <deployedseed>531607771</deployedseed>
    <deployedseed>852940610</deployedseed>
    <deployedseed>395610705</deployedseed>
    <deployedseed>1273515090</deployedseed>
    <deployedseed>893988380</deployedseed>
    <deployedseed>589126148</deployedseed>
    <deployedseed>830356288</deployedseed>
    <deployedseed>1447461032</deployedseed>
    <deployedseed>1979200276</deployedseed>
    <deployedseed>211621577</deployedseed>
    <deployedseed>1624415667</deployedseed>
    <deployedseed>1188440933</deployedseed>
    <deployedseed>833374768</deployedseed>
    <deployedseed>457676487</deployedseed>
    <deployedseed>520560727</deployedseed>
    <deployedseed>877116554</deployedseed>
    <deployedseed>1538904761</deployedseed>
    <deployedseed>1579337625</deployedseed>
    <deployedseed>258721661</deployedseed>
    <deployedseed>1908521336</deployedseed>
    <deployedseed>1647064931</deployedseed>
    <deployedseed>208187902</deployedseed>
    <deployedseed>1861906201</deployedseed>
    <deployedseed>1921820105</deployedseed>
    <deployedseed>1962799400</deployedseed>
    <deployedseed>1993120426</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
