<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Dividing a Class with Twins: Group Formation Challenge</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>In how many ways can a class of \({@class_tot@}\) be divided into two groups of \({@group_1@}\) and \({@group_2@}\) respectively if the class contains one set of twins who must not be separated?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>There are two options, depending on which group the twins are in:</p>
<p>If the twins are in the group of \({@group_2@}\), then there are \(^{@class_tot-2@} C_{@group_1@}\) ways of choosing the group of \({@group_1@}\) – the remaining \({@(class_tot-2)-group_1@}\) are put together with the twins.</p>
<p>If they are in the group of \({@group_1@}\), then there are \(^{@class_tot-2@} C_{@group_2@}\) ways of choosing the group of \({@group_2@}\) – the remaining \({@(class_tot-2)-group_2@}\) are put together with the twins.</p>
<p>Thus the number of possibilities is  \(^{@class_tot-2@} C_{@group_1@} + {}^{@class_tot-2@} C_{@group_2@} = {@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>class_tot:rand_with_step(15,25,1);
group_1:round(class_tot/2)+rand_with_step(1,3,1);
group_2:class_tot-group_1;
ta:(binomial(class_tot-2, group_1)+binomial(class_tot-2, group_2));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>In \({@ta@}\) ways can a class of  \({@class_tot@}\) be divided into two groups of \({@group_1@}\) and \({@group_2@}\) respectively if the class contains one set of twins who must not be separated</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/13_-Z-yfLg9hFHWa4QWnfM4zSE2RxP_ZuYxQSPKPMals/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question requires the student to determine the number of ways to divide a class of 20 students into two groups of 8 and 12, respectively, while ensuring that a set of twins remains together.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>126131650</deployedseed>
    <deployedseed>150248793</deployedseed>
    <deployedseed>1823260760</deployedseed>
    <deployedseed>869998937</deployedseed>
    <deployedseed>701226337</deployedseed>
    <deployedseed>1527452894</deployedseed>
    <deployedseed>100732241</deployedseed>
    <deployedseed>458959492</deployedseed>
    <deployedseed>1689520055</deployedseed>
    <deployedseed>1312535127</deployedseed>
    <deployedseed>1781784510</deployedseed>
    <deployedseed>970095974</deployedseed>
    <deployedseed>595214734</deployedseed>
    <deployedseed>824123141</deployedseed>
    <deployedseed>1482465897</deployedseed>
    <deployedseed>36563564</deployedseed>
    <deployedseed>1542089815</deployedseed>
    <deployedseed>749022492</deployedseed>
    <deployedseed>997228985</deployedseed>
    <deployedseed>981091468</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
