<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Estimating Measure by reading single Decrement Tables(Life Tables)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Use the BRAZIL \(1950\) LIFETABLE given below to answer the question that follows.
</p>
 <a href="https://drive.google.com/uc?export=download&id=1NfHh5bPWXGnfTfxSl68Ie5wy2umo3HQ8">Download Brazillian Demography Life tables</a>
<br>
<br>
<p>Using the life tables for Brazil, estimate the following measures.</p>
<ol class="HELM_parts">
<li> In \(1946\), there were \({@population@}\) female births, estimate the number of girls aged last birthday on \(31-12-{@1946+years@}\).</li>
<p>Estimate number of girls aged last birthday on \(31-12-{@1946+years@} = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<br>
<li> A couple aged (both) \(25\) years have male twins. What is the probability that one parent or both will die before either of the sons in the next \({@event_a_b@}\) years.</li>
<p>(Note: leave you answer to \(3 \text{dp}\).)</p>
 </ol>
<p><b>i.</b>  Probability that the male parent will die before either of the son \(=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p><b>ii.</b>  Probability that the female parent will die before either of the son \(=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p><b>iii.</b>  Probability that the both parents will die before either of the son \(=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>
<p><b>iv.</b>  Probability that the sons will die before the \({@event_a_b@}\) years. \(=\) [[input:ans5]] [[validation:ans5]][[feedback:prt5]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<ol class="HELM_parts">
<li>  Let the number of population over the years \({@1946+years@} -1946={@years@}\) be denoted as \(l_{@years@}\) and the original population be \(l_0\)</li>
<p>Thus,Estimate number of girls aged last birthday on \(31-12-{@1946+years@} \) is;</p>
<p>
\(\begin{align}
\text{Estimate number} &= \displaystyle \frac{l_{@years@}}{l_0} \times {@population@}\\
&=\displaystyle \frac{{@l_x@}} {{@l_0@}} \times {@population@}\\
&={@sol1@}\\
\end{align}\)
</p>
<br>
<li></li>
 </ol>
<p><b>i.</b>  Probability that the male parent will die before either of the son.</p>
<p>Let \(A\) be the event that the man die before the next {@event_a_b@} years and </p>
<p>\(l^m_{@25+event_a_b@}\) be population of men in \({@event_a_b@}\) and \(l_{25}\) be the population of men at \(25\) years.</p>
<p>\( \begin{align}
P(A)&=1-\displaystyle \frac{l^m_{@25+event_a_b@}}{l_{25}}\\
&=1-\displaystyle \frac{{@l_m_x@}}{{@l_m_age@}}\\
&=1-{@decimalplaces(soln2,3)@}\\
&={@ta3@}\\
\end{align}\)</p>
<br>
<p><b>ii.</b>  Probability that the female parent will die before either of the son. </p>
<p>Let \(B\) be the event that the female die before the next {@event_a_b@} years and </p>
<p>\(l^f_{@25+event_a_b@}\) be population of females in \({@event_a_b@}\) and \(l_{25}\) be the population of female at \(25\) years.</p>
<p>\( \begin{align}
P(A)&=1-\displaystyle \frac{l^f_{@25+event_a_b@}}{l_{25}}\\
&=1-\displaystyle \frac{{@l_f_x@}}{{@l_f_age@}}\\
&=1-{@decimalplaces(soln1,3)@}\\
&={@ta2@}\\
\end{align}\)</p>
<br>
<p><b>iii.</b>  Probability that the both parents will die before either of the son.</p>
<p>Probability that either or both parents die before the next {@event_a_b@} years.</p>
<p>\(P(A \cup B)=P(A)+P(B)-P(A)P(B)\)</p>
<p>\(\begin{align}
P(A \cup B)&={@ta3@}+{@ta2@}-({@ta2@} \times{@ta3@})\\
&={@ta4@}\\
\end{align}\)</p>
<br>
<p><b>iv.</b> Probability that the sons will die before the \({@event_a_b@}\) years.</p>
<p>Let \(C_1\) and \(C_2\) be the events that the sons will die before the \({@event_a_b@}\) years.</p>
<p>\(P(C_1)=1-\displaystyle \frac{l_{@event_a_b@}}{l_0}\) which is \(=P(C_2)=1-\displaystyle \frac{l_{@event_a_b@}}{l_0}\)</p>
<p>\(\begin{align}
P(C_1)&=P(C_2)\\
&=1-\displaystyle \frac{{@l_x_sons@}}{100000}\\
&={@s1@}\\
&=2({@s1@})-({@s1@}^2)\\
&={@ta5@}
\end{align}\)</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*LIFE TABLES*/
/*part a*/
years:rand([5,10,15,20,25]);
l_0:100000;/*original population*/
l_x:if years=5 then 77007 else if years=10 then 74889 else if years=15 then 73260 else if years=20 then 71116 else if years=25 then 68469; /*population on the table*/
population:rand(50000)+300000;/*popluation over the years*/
/*solutions*/
sol1:floor(float((l_x/l_0)*population));


/*part b*/
/*event b*/
event_a_b:rand([10,15,20,25,30]);
age:25;
l_f:age+event_a_b;/*age in the next 10yrs*/
l_f_x:if event_a_b=10 then 62472 else if event_a_b=15 then 59159 else if event_a_b=20 then 55674 else if event_a_b=25 then 51922 else if event_a_b=30 then 47279;/*l_x in next yrs*/
l_f_age:68469; /*l_x in exact age*/
/*solution*/
soln1:1-(l_f_x/l_f_age);
probability_b:float(soln1);
/*male events*/
age:25;
l_m:age+event_a_b;/*age in the next 10yrs*/
l_m_x:if event_a_b=10 then 60279 else if event_a_b=15 then 56705 else if event_a_b=20 then 52558 else if event_a_b=25 then 47858 else if event_a_b=30 then 42223;/*l_x in next 10yrs*/
l_m_age:66352; /*l_x in exact age*/
/*solution*/
soln2:1-(l_m_x/l_m_age);
probability_a:float(soln2);
/* both parents probability*/
p_aub:probability_a+probability_b-(probability_a*probability_b);

/*sons probabilities*/
sons_age:event_a_b;
l_x_sons:if sons_age=10 then 72523 else if sons_age=15 then 71094 else if sons_age=20 then 69096 else if sons_age=25 then 66352 else if sons_age=30 then 63457;
/*solution*/
s1:decimalplaces(1-(l_x_sons/100000),3);
p_c1_c2:2*(s1)-(s1^2);
s2:decimalplaces(p_c1_c2,3);

/*model answers*/
ta1:sol1;
ta2:decimalplaces(probability_b,3);
ta3:decimalplaces(probability_a,3);
ta4:decimalplaces(p_aub,3);
ta5:s2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Estimate number of girls aged last birthday on \(31-12-{@1946+years@} ={@ta1@} \) </p>
<br>
<p><i>i.</i>  Probability that the male parent will die before either of the son \(={@ta2@}\)</p>
<p><i>ii.</i>  Probability that the female parent will die before either of the son \(={@ta3@}\)</p>
<p><i>iii.</i>  Probability that the both parents will die before either of the son \(={@ta4@}\)</p>
<p><i>iv.</i>  Probability that the sons will die before the \({@event_a_b@}\) years. \(={@ta5@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1DxOxfp2iCmNd-ZoURRoYJJ19c4gk_Ftz89A9sOcbeFo/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">Given life tables for Brazil Demography analysis, a question requires you ton read the tables and determine the probabilities that a couple at \(25\)yrs, will each die after a given random number of years and the probabilities that the sons die in the given number of years.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct  answer ta1 on female population</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for probability that male parent dies</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct, you left your answer to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>  checks if the answer is a fraction</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>soln1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer should not be a fraction but float numbers upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks if the decimal places are more than 3</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.1</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks if the decimal places are more than 3</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks the probability that female parent dies</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct, you left your answer to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if the answer has more than 3 dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>soln2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer should not be a fraction but float numbers upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks the probability that both parents die</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct, you left your answer to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.1</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks the probability that both sons die</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct, you left your answer to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks if answer has less or more than 3 dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>0.1</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You answer is not correct because it's not upto to \(3\)dp</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>22753903</deployedseed>
    <deployedseed>1247248823</deployedseed>
    <deployedseed>1350701102</deployedseed>
    <deployedseed>808501414</deployedseed>
    <deployedseed>1282195367</deployedseed>
    <deployedseed>1522797549</deployedseed>
    <deployedseed>847812891</deployedseed>
    <deployedseed>1499155148</deployedseed>
    <deployedseed>957757707</deployedseed>
    <deployedseed>2105344875</deployedseed>
    <deployedseed>904506500</deployedseed>
    <deployedseed>1032550657</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>soln1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>soln2</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value/>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value/>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore/>
        <expectedpenalty/>
        <expectedanswernote>NULL</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
