<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find median in grouped data</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The following is the height of \({@total_f@}\) students in a class. </p>
<p>Fill in the cumulative frequency column of the table:</p>

<table border="solid">
  <tr>
    <th width="25%", style="text-align:center">Height (in cm)</th>
    <th width="25%", style="text-align:center">Height (in cm) (class boundaries)</th>
    <th width="25%", style="text-align:center">Number of students (f)</th>
    <th width="30%", style="text-align:center">Cumulative frequency (cf)</th>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range1@}-{@upper_range1@}\)</td>
    <td style="text-align:center">\({@lower_bound1@}-{@upper_bound1@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf1@}\)</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range2@}-{@upper_range2@}\)</td>
    <td style="text-align:center">\({@lower_bound2@}-{@upper_bound2@}\)</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</td>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range3@}-{@upper_range3@}\)</td>
    <td style="text-align:center">\({@lower_bound3@}-{@upper_bound3@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range4@}-{@upper_range4@}\)</td>
    <td style="text-align:center">\({@lower_bound4@}-{@upper_bound4@}\)</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range5@}-{@upper_range5@}\)</td>
    <td style="text-align:center">\({@lower_bound5@}-{@upper_bound5@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td>
  </tr>
</table>

<br>

<p>Find the median height and give your answer to \(2\) decimal places where necessary.</p>
<p>The median height is [[input:ans5]] [[validation:ans5]] \(\text{cm}\)</p>
<p>[[feedback:prt5]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Since the height ranges provided are discontinuous, we introduce a correcting factor of \( \frac{{@lower_range2@}-{@upper_range1@}}{2}={@error_bound@}\) to determine the class boundaries.</p>

<p>From this, you can prepare the class boundary column and the cumulative frequency column as follows:</p>

<table border="solid">
  <tr>
    <th width="25%", style="text-align:center">Height (in cm)</th>
    <th width="25%", style="text-align:center">Height (in cm) (class boundaries)</th>
    <th width="25%", style="text-align:center">Number of students (f)</th>
    <th width="30%", style="text-align:center">Cumulative frequency (cf)</th>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range1@}-{@upper_range1@}\)</td>
    <td style="text-align:center">\({@lower_bound1@}-{@upper_bound1@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf1@}\)</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range2@}-{@upper_range2@}\)</td>
    <td style="text-align:center">\({@lower_bound2@}-{@upper_bound2@}\)</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">\({@cf1@}+{@f2@}={@cf2@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range3@}-{@upper_range3@}\)</td>
    <td style="text-align:center">\({@lower_bound3@}-{@upper_bound3@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">\({@cf2@}+{@f3@}={@cf3@}\)</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range4@}-{@upper_range4@}\)</td>
    <td style="text-align:center">\({@lower_bound4@}-{@upper_bound4@}\)</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">\({@cf3@}+{@f4@}={@cf4@}\)</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range5@}-{@upper_range5@}\)</td>
    <td style="text-align:center">\({@lower_bound5@}-{@upper_bound5@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">\({@cf4@}+{@f5@}={@cf5@}\)</td>
  </tr>
  <tr>
   <td style="text-align:center"></td>
    <td style="text-align:center">Total</td>
    <td style="text-align:center">\( \sum f= {@total_f@}\)</td>
    <td style="text-align:center"></td>
  </tr>
</table>

<br>

<p>The median class is that class containing the item in position \( \frac{{@total_f@}}{2} \), which is the \({@median_cf@}^{\text{th}}\) item.</p>
<p>Therefore, the median class is \({@lower_bound@}-{@upper_bound@}\).</p>

<br>

<p>To estimate the median height from the data, we use linear interpolation in the median class. For this, we assume that the data is evenly spaced across the class and determine the height of the median position, in this case the \({@median_cf@}^{\text{th}}\) item.</p>

<p>The width of the median class is, \({@upper_bound@}-{@lower_bound@}={@width@}\)
<br>The lower boundary of the class is, {@lower_bound@}</p>

<p>The number of items in the class is, {@f_median@}
<br>The number of items in classes before the median class is, {@cf_lower@}</p>

<br>

<p>With this, we have extracted all the required information to perform the linear interpolation and estimate the median height of the data, \(m\). We compare the ratios between the position of the median value in the data points with the position of the median in the median class:</p>

<p style="margin-left:30px">\(\displaystyle \frac{m - {@lower_bound@}}{{@width@}} = \frac{{@median_cf@} - {@cf_lower@}}{{@f_median@}}\)</p>
<p style="margin-left:30px">\(\begin{align*}
\Rightarrow m 
&= {@lower_bound@} + \frac{{@median_cf@} - {@cf_lower@}}{{@f_median@}} \cdot {@width@} \\
&= {@lower_bound@}+{@md1@} \\
&= {@ta@} \hspace{10px} \text{(2pd)}
\end{align*}\)</p>

<br>

<p>So, the median height is \({@ta@} \text{cm}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*randomly generating column 1 height intervals*/
aa:rand(200)+10;
lower_range1: aa+ 5;  /*Adding 5 on each variable to ensure the subsequent class width is the same*/
upper_range1: lower_range1 + 5;

lower_range2: upper_range1 + 1;
upper_range2: lower_range2 + 5;

lower_range3: upper_range2 + 1;
upper_range3: lower_range3 + 5;

lower_range4: upper_range3 + 1;
upper_range4: lower_range4 + 5;

lower_range5: upper_range4 + 1;
upper_range5: lower_range5 + 5;

error_bound:float((lower_range2-upper_range1)/2);

lower_bound1: lower_range1-error_bound;  /*Adding 0.5 and subtracting 0.5 as the correcting factor for class boundaries*/
upper_bound1: upper_range1+error_bound;

lower_bound2:  lower_range2-error_bound;
upper_bound2:upper_range2+error_bound;

lower_bound3:  lower_range3-error_bound;
upper_bound3:upper_range3+error_bound;

lower_bound4:  lower_range4-error_bound;
upper_bound4:upper_range4+error_bound;

lower_bound5:  lower_range5-error_bound;
upper_bound5:upper_range5+error_bound;


frequency_list:makelist(rand_with_step(2,20,2),5);/*create a list with 5 elements randomly picked between 2-20 for even N*/
f1:first(frequency_list);          /*the first item in the list*/
f2:second(frequency_list);
f3:third(frequency_list)+2;  /*Adding by 10 to make sure this is the modal class freq*/
f4:fourth(frequency_list);
f5:fifth(frequency_list);      /*the fifth item in the list*/

total_f:f1+f2+f3+f4+f5;  /*number of  students in our class*/

cf1:f1; /*cumulative frequencies*/
cf2:cf1+f2; 
cf3:cf2+f3; 
cf4:cf3+f4; 
cf5:cf4+f5;

/*Calculating the median of Grouped Data*/    
frequencies: [f1, f2, f3, f4, f5];
cf: [cf1, cf2, cf3, cf4, cf5];
n:total_f; /*Total Frequency*/

lower_boundaries: [lower_bound1, lower_bound2, lower_bound3, lower_bound4, lower_bound5];
upper_boundaries: [upper_bound1, upper_bound2, upper_bound3, upper_bound4, upper_bound5];

class_widths: upper_boundaries - lower_boundaries;
midpoints: lower_boundaries + class_widths / 2;

median_cf: n / 2;

i: 1;
while cf[i] < median_cf do i: i + 1;

median_cf;
lower_bound:lower_boundaries[i];/*L is the lower limit of the median class.*/
upper_bound: upper_boundaries[i];
cf_lower: cf[i-1]; /*cf is the cumulative frequency of the class before the median class.*/
f_median: frequencies[i]; /*f is the frequency of the median class.*/
width: width: upper_range2 - upper_range1; /*w is the interval width (the difference between the upper and lower limits of each class).*/

median_height: lower_boundaries[i] + (median_cf - cf[i-1]) * class_widths[i] / frequencies[i];
md1: float( (median_cf - cf[i-1]) * class_widths[i] / frequencies[i]);

/*closest median height integer*/
median_height1:decimalplaces(median_height,2);
ta1:round(median_height1);
ta_1:is(abs(round(median_height1)-median_height1)=0.0);
ta: if ta_1=true then ta1 else median_height1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br>Number of student \({@total_f@}\)
<br>\(B_L={@lower_bound@}, \, \displaystyle  \left( \frac{n}{2}\right)={@median_cf@}, \, f_c={@f_median@}, \, i={@width@}, \, cf_b={@cf_lower@} \)
<br>The median height is \({@ta@}\)cm.]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">NA</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Dan Kelly</p>
<p id="description">Given a table with grouped data and frequencies find the cumulative frequencies and find the median.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>cf2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>cf3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>cf4</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>cf5</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>cf2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>cf3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>cf4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>cf5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/*node 1 checks for correct answer when an integer or 2 decimal places*/

/*Node 2 checks for correct float answer*/
median_height2: if ta=false then  float(median_height);

/*Node 3 checks for 3 decimal place*/

/*Node 4 tests sa close to the teacher's answer, abs(sa-ta)<opt */

/*node 5 mistake:Using the upper bound of the median class*/
median_mistake: upper_boundaries[i] + (median_cf - cf[i-1]) * class_widths[i] / frequencies[i];
mistake_2:decimalplaces(median_mistake,2);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer in exact form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>median_height2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for correct answer in with wrong number of decimal places</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans5</sans>
        <tans>median_height</tans>
        <testoptions>3</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but not given to the right number of decimal places.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt5-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for correct answer in rounding mistake</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>median_height</tans>
        <testoptions>0.02</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is close to the correct answer but you are likely to have made a rounding mistake.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt5-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for correct answer in only in he upper limit of the median class.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your computations are correct but you used the upper limit of the median class.</p>
<p>Instead you should use the lower limit of the median class when computing the median.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2097273858</deployedseed>
    <deployedseed>204974300</deployedseed>
    <deployedseed>338731865</deployedseed>
    <deployedseed>1370681493</deployedseed>
    <deployedseed>691949249</deployedseed>
    <deployedseed>715050382</deployedseed>
    <deployedseed>1673481494</deployedseed>
    <deployedseed>430980473</deployedseed>
    <deployedseed>56471626</deployedseed>
    <deployedseed>66273648</deployedseed>
    <deployedseed>943614167</deployedseed>
    <deployedseed>146834349</deployedseed>
    <deployedseed>1547146956</deployedseed>
    <deployedseed>590032037</deployedseed>
    <deployedseed>560178124</deployedseed>
    <deployedseed>462838300</deployedseed>
    <deployedseed>1366745135</deployedseed>
    <deployedseed>1756198286</deployedseed>
    <deployedseed>252408909</deployedseed>
    <deployedseed>1094234934</deployedseed>
    <deployedseed>2031161360</deployedseed>
    <deployedseed>1007727263</deployedseed>
    <deployedseed>2077473078</deployedseed>
    <deployedseed>1140860535</deployedseed>
    <deployedseed>1295397471</deployedseed>
    <deployedseed>711940560</deployedseed>
    <deployedseed>2042792177</deployedseed>
    <deployedseed>186717030</deployedseed>
    <deployedseed>2083457834</deployedseed>
    <deployedseed>867943688</deployedseed>
    <deployedseed>839111281</deployedseed>
    <deployedseed>2130612292</deployedseed>
    <deployedseed>241673348</deployedseed>
    <deployedseed>1416551321</deployedseed>
    <deployedseed>686651145</deployedseed>
    <deployedseed>728878428</deployedseed>
    <deployedseed>64692486</deployedseed>
    <deployedseed>1514671231</deployedseed>
    <deployedseed>333541939</deployedseed>
    <deployedseed>962003931</deployedseed>
    <deployedseed>1944886916</deployedseed>
    <deployedseed>618566905</deployedseed>
    <deployedseed>1583824524</deployedseed>
    <deployedseed>1157475208</deployedseed>
    <deployedseed>1871750205</deployedseed>
    <deployedseed>1686534481</deployedseed>
    <deployedseed>787658268</deployedseed>
    <deployedseed>1802035244</deployedseed>
    <deployedseed>690805062</deployedseed>
    <deployedseed>1949989056</deployedseed>
    <deployedseed>776232719</deployedseed>
    <deployedseed>75168850</deployedseed>
    <deployedseed>1968049319</deployedseed>
    <deployedseed>95605798</deployedseed>
    <deployedseed>1628469727</deployedseed>
    <deployedseed>673076889</deployedseed>
    <deployedseed>2130230581</deployedseed>
    <deployedseed>1903752002</deployedseed>
    <deployedseed>1227357364</deployedseed>
    <deployedseed>2015351821</deployedseed>
    <deployedseed>1873007698</deployedseed>
    <deployedseed>720020104</deployedseed>
    <deployedseed>668287920</deployedseed>
    <deployedseed>275299656</deployedseed>
    <deployedseed>86692966</deployedseed>
    <deployedseed>1670180472</deployedseed>
    <deployedseed>448187670</deployedseed>
    <deployedseed>1965246701</deployedseed>
    <deployedseed>1455158870</deployedseed>
    <deployedseed>148616570</deployedseed>
    <deployedseed>1030287868</deployedseed>
    <deployedseed>1372998320</deployedseed>
    <deployedseed>1745946915</deployedseed>
    <deployedseed>1483367552</deployedseed>
    <deployedseed>362751089</deployedseed>
    <deployedseed>1461475956</deployedseed>
    <deployedseed>1152949723</deployedseed>
    <deployedseed>1229046538</deployedseed>
    <deployedseed>1533576746</deployedseed>
    <deployedseed>208977911</deployedseed>
    <deployedseed>987324154</deployedseed>
    <deployedseed>1022367637</deployedseed>
    <deployedseed>1194159063</deployedseed>
    <deployedseed>1172476738</deployedseed>
    <deployedseed>1590817153</deployedseed>
    <deployedseed>807400532</deployedseed>
    <deployedseed>1990041982</deployedseed>
    <deployedseed>881988</deployedseed>
    <deployedseed>1938603424</deployedseed>
    <deployedseed>1249187963</deployedseed>
    <deployedseed>988679604</deployedseed>
    <deployedseed>711318597</deployedseed>
    <deployedseed>1482729569</deployedseed>
    <deployedseed>881926538</deployedseed>
    <deployedseed>129933455</deployedseed>
    <deployedseed>680095690</deployedseed>
    <deployedseed>2020479761</deployedseed>
    <deployedseed>547572363</deployedseed>
    <deployedseed>847913684</deployedseed>
    <deployedseed>1888884366</deployedseed>
    <deployedseed>346612103</deployedseed>
    <deployedseed>152949030</deployedseed>
    <deployedseed>1528395696</deployedseed>
    <deployedseed>1735657755</deployedseed>
    <deployedseed>1336496661</deployedseed>
    <deployedseed>417827031</deployedseed>
    <deployedseed>1198019622</deployedseed>
    <deployedseed>217359018</deployedseed>
    <deployedseed>688626762</deployedseed>
    <deployedseed>1499528989</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>cf2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>cf3</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>cf4</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>cf5</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
