<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding a conditional probability of the product of 2 dice throws</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Two fair dice are thrown and the product of the numbers on the dice is recorded. Given that one die lands on {@one_die@}, find the probability that the product on the dice is</p>
<p>(a) Exactly {@exact_product@} = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>(b) More than {@more_than_product@} = [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Remember that the probability of an event \(B\) occurring given that event \(A\) has occurred is denoted as \(P(B|A)\) and is calculated as follows: \(P(B|A)=\frac{P(B|A)}{P(A)}\). This is known as conditional probability. </p>
<p>Also, keep in mind that the probability space for this question can be represented in a table as shown.</p>
<p dir="ltr"><b id="docs-internal-guid-cdbe3a72-7fff-1ea0-e9e2-04def6d53ae0"></b></p>
<div dir="ltr" align="left">
    <table>
        <colgroup>
            <col width="69">
            <col width="67">
            <col width="74">
            <col width="65">
            <col width="65">
            <col width="70">
            <col width="67">
        </colgroup>
        <tbody>
            <tr>
                <td>
                    <p dir="ltr">Dice 1<br></p>
                    <p dir="ltr"></p>
                    <hr>
                    <p></p>
                    <p dir="ltr">Dice 2<br></p>
                </td>
                <td>
                    <p dir="ltr">1</p>
                </td>
                <td>
                    <p dir="ltr">2</p>
                </td>
                <td>
                    <p dir="ltr">3</p>
                </td>
                <td>
                    <p dir="ltr">4</p>
                </td>
                <td>
                    <p dir="ltr">5</p>
                </td>
                <td>
                    <p dir="ltr">6</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">1</p>
                </td>
                <td>
                    <p dir="ltr">1x1=1</p>
                </td>
                <td>
                    <p dir="ltr">1x2=2</p>
                </td>
                <td>
                    <p dir="ltr">1x3=3</p>
                </td>
                <td>
                    <p dir="ltr">1x4=4</p>
                </td>
                <td>
                    <p dir="ltr">1x5=5</p>
                </td>
                <td>
                    <p dir="ltr">1x6=6</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">2</p>
                </td>
                <td>
                    <p dir="ltr">2x1=2</p>
                </td>
                <td>
                    <p dir="ltr">2x2=4</p>
                </td>
                <td>
                    <p dir="ltr">2x3=6</p>
                </td>
                <td>
                    <p dir="ltr">2x4=8</p>
                </td>
                <td>
                    <p dir="ltr">2x5=10</p>
                </td>
                <td>
                    <p dir="ltr">2x6=12</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">3</p>
                </td>
                <td>
                    <p dir="ltr">3x1=3</p>
                </td>
                <td>
                    <p dir="ltr">3x2=6</p>
                </td>
                <td>
                    <p dir="ltr">3x3=9</p>
                </td>
                <td>
                    <p dir="ltr">4x3=12</p>
                </td>
                <td>
                    <p dir="ltr">3x5=15</p>
                </td>
                <td>
                    <p dir="ltr">3x6=18</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">4</p>
                </td>
                <td>
                    <p dir="ltr">4x1=4</p>
                </td>
                <td>
                    <p dir="ltr">4x2=8</p>
                </td>
                <td>
                    <p dir="ltr">4x3=12</p>
                </td>
                <td>
                    <p dir="ltr">4x4=16</p>
                </td>
                <td>
                    <p dir="ltr">4x5=20</p>
                </td>
                <td>
                    <p dir="ltr">4x6=24</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">5</p>
                </td>
                <td>
                    <p dir="ltr">5x1=5</p>
                </td>
                <td>
                    <p dir="ltr">5x2=10</p>
                </td>
                <td>
                    <p dir="ltr">5x3=15</p>
                </td>
                <td>
                    <p dir="ltr">5x4=20</p>
                </td>
                <td>
                    <p dir="ltr">5x5=25</p>
                </td>
                <td>
                    <p dir="ltr">5x6=30</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr">6</p>
                </td>
                <td>
                    <p dir="ltr">6x1=6</p>
                </td>
                <td>
                    <p dir="ltr">6x2=12</p>
                </td>
                <td>
                    <p dir="ltr">6x3=18</p>
                </td>
                <td>
                    <p dir="ltr">6x4=24</p>
                </td>
                <td>
                    <p dir="ltr">6x5=30</p>
                </td>
                <td>
                    <p dir="ltr">6x6=36</p>
                </td>
            </tr>
        </tbody>
    </table>
</div><br><p>From the table above, you notice that there are 36 possible outcomes when the two dice are rolled. <br></p><p>Given that one dice lands on a {@one_die@}, <br></p>
<p>(a) The probability that the product of the two dices is Exactly {@exact_product@} will be {@ta_1@}<br>
</p><p></p><p>(b) The probability that the product of the two dices is more than {@more_than_product@} will be&nbsp; {@ta_2@}<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Question 1*/
/*random generation of outcomes*/
one_die:rand(4)+2;
exact_product:one_die*(rand(4)+2);
more_than_product:rand_with_step(2*one_die+1,4*one_die-1,1);

/*Model answer*/
ta_1:2/11;
sample_21:makelist(i*one_die,i,1,6);
sample_22:delete(one_die^2,sample_21);
count_more_than:0;
for i:1 thru 6 do (if sample_21[i] >more_than_product then count_more_than:count_more_than+1);
for i:1 thru 5 do (if sample_22[i] >more_than_product then count_more_than:count_more_than+1);
ta_2:count_more_than/11;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>die lands on {@one_die@},  {@exact_product@} and  {@more_than_product@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checking correct answer*/

/*Node2: Check if answer is taken as the addition of the two dice to give the probability of exactly 6*/
mistake_2:12;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span id="docs-internal-guid-e90ab2f8-7fff-3faf-0941-21fcbcd32684"><span id="docs-internal-guid-654e81b7-7fff-720d-bc9e-e715a4f0c713"><span id="docs-internal-guid-61ed7f48-7fff-be71-7cd3-87259bca0854"><span id="docs-internal-guid-d975528c-7fff-b1dc-f9ae-c6f58a8a052b"></span></span></span></span></p><p dir="ltr" style="">You have given the sum of 6 on two dice and not the product of 6 and 2 on the die over the probability of two on one die.</p><br><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node1: Checking correct answer*/

/*Node2:Checking the total sum of outcomes that one die lands on 2 as more than 5*/
mistake_2:36;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span id="docs-internal-guid-06e7b14c-7fff-355b-aa03-cbb2731794d9">The answer given is addition of one values on one die but not the other die.</span><br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
