<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying from a list of non-probabilistic sampling techniques the exception</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The following are types of non-probabilistic sampling techniques except</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]

<p>Stratified sampling is a sampling technique used in statistics and research to ensure that the sample drawn from a population is representative and includes subgroups or strata in the same proportion as they exist in the population. It is a systematic approach to sampling that can improve the accuracy and reliability of survey results by ensuring that each subgroup within the population is adequately represented.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>Systematic sampling is a probability sampling technique used in statistics and research to select a sample from a larger population in a systematic and structured way. It is an efficient method for obtaining a representative sample while maintaining simplicity in the sampling process.</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]

<p>Multistage sampling, also known as multi-stage cluster sampling, is a complex probability sampling technique used in statistics and research to sample from large and diverse populations. This method involves breaking down the population into multiple stages or levels, with sampling occurring at each stage. Multistage sampling is particularly useful when it is impractical or too costly to conduct a simple random sample of the entire population. </p>

[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*STA 141 week2 quizes*/
correct_answer:rand_selection([
[1, true, "Stratified sampling"],
[2, true, "Systematic sampling"],
[3, true, "Multistage sampling"]
], 1);
incorrect_answers:rand_selection([
[4, false, "Voluntary response sampling"],
[5, false, "Quota sampling"],
[6, false, "Purposive sampling"],
[7, false, "Convenience sampling"],
[8, false, "Judgmental sampling"]
], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The following are types of non-probabilistic sampling techniques except</p>
[[ if test="member(1, ta_index)" ]]

<p>Stratified sampling</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>Systematic sampling</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]

<p>Multistage sampling</p>

[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on Identify what  types of non-probabilistic sampling techniques.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1402190781</deployedseed>
    <deployedseed>1167768145</deployedseed>
    <deployedseed>1330004817</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
