<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying Improper Set Representations</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following is <b>NOT</b> a proper or conventional presentation of a set?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <i> Set a={A,B,C,D}</i></p>
<p>This is not a conventional name for a set. By convention, names of sets should be capital letters.</p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The correct answer is <i> Set B=1,2,3,4</i></p>
<p> This is not a proper way of representing a set, elements in a set are enclosed in curly brackets.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>The correct answer is <i> Set D=(1,2,3,4)</i></p>
<p> This is not a proper way of representing a set, elements in a set are enclosed in curly brackets.</p>
[[/ if ]]

[[ if test="member(4, ta_index)" ]]
<p>The correct answer is <i> Set E=[4,5,6,7]</i></p>
<p>This is not a proper way of representing a set, elements in a set are enclosed in curly brackets.</p>
[[/ if ]]

[[ if test="member(5, ta_index)" ]]
<p>The correct answer is <i> Set A={a b c d}</i></p>
<p>This is not a proper way of representing a set, elements in a set are separated by commas.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Identifying sets*/
/*model answer*/
ta_1:random_permutation([["Set a={A,B,C,D}",true],["Set B={1,2,3,4,5}",false],["Set C={a,b,c,d}",false]]);
    
correct_answer:rand_selection([
[1, true, "Set a={A,B,C,D}"],
[2, true, "Set B=1,2,3,4"],
[3, true, "Set D=(1,2,3,4)"],
[4, true, "Set E=[4,5,6,7]"],
[5, true, "Set A={a b c d}"]],1);
incorrect_answers:rand_selection([
[6, false, "Set B={1,2,3,4,5}"],
[7, false, "Set C={a,b,c,d}"],
[8, false, "Set D={}"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following is <b>NOT</b> a proper presentation of sets</p>
{@ta@}
[[ if test="member(1, ta_index)" ]]
<p> Set a={A,B,C,D} </p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p> Set B=1,2,3,4 </p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p> Set d=(1,2,3,4) </p>
[[/ if ]]

[[ if test="member(4, ta_index)" ]]
<p> Set E=[4,5,6,7]</p>
[[/ if ]]

[[ if test="member(5, ta_index)" ]]
<p>Set A={a b c d}</p>
[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1Qs3otO6IBjGuFaVHi95A-OYUFIaKtsHVsHD4lWFe9mc/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">The question focuses on identifying proper ways of representing sets </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct option</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>887581462</deployedseed>
    <deployedseed>1717226483</deployedseed>
    <deployedseed>1882696228</deployedseed>
    <deployedseed>1147799668</deployedseed>
    <deployedseed>1326597233</deployedseed>
    <deployedseed>1617031880</deployedseed>
    <deployedseed>1284733327</deployedseed>
    <deployedseed>1734002432</deployedseed>
    <deployedseed>1410238195</deployedseed>
    <deployedseed>1621483781</deployedseed>
    <deployedseed>1778827114</deployedseed>
    <deployedseed>2070874960</deployedseed>
    <deployedseed>1686932128</deployedseed>
    <deployedseed>1241228846</deployedseed>
    <deployedseed>1003251529</deployedseed>
    <deployedseed>1767905140</deployedseed>
    <deployedseed>1023097339</deployedseed>
    <deployedseed>285214685</deployedseed>
    <deployedseed>2075195739</deployedseed>
    <deployedseed>768851236</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
