<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying advantages of using Mean</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which statement DOES NOT represent a disadvantage of using the mean?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]

<p>The desired algebraic property, refers to a specific property that is not possessed by the mean, and it's known as "linearity." The mean does not satisfy the property of linearity when applied to the sum or difference of random variables or datasets. This means that the mean of the sum (or difference) of two datasets is not generally equal to the sum (or difference) of the means of those datasets.

To illustrate this, let's consider two random variables \(X\) and \(Y\), and their respective means, denoted as \(E(X)\) and \(E(Y)\). The expected value (mean) of the sum of these two random variables is not necessarily equal to the sum of their means:

\(E(X + Y) ≠ E(X) + E(Y)\)</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>It cannot be computed for grouped data in the cases of open classes
because computing the mean (average) in grouped data with open-ended class intervals can be challenging due to the lack of specific data points within those intervals. This is because the mean relies on knowing the individual data points, and when you have open-ended intervals, you don't have access to those specific values</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]

<p> "it is least affected by fluctuations of the sample" is not accurate when referring to the mean. In fact, the mean is quite sensitive to fluctuations in a sample. If you want a measure that is less influenced by fluctuations or outliers, it's better to consider the median or the mode.</p>

[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*STA 141 week2 quizes*/
correct_answer:rand_selection([
[1, true, "It does not possess the desired algebraic property"],
[2, true, "It is uniquely defined"],
[3, true, "It is least affected by fluctuations of the sample"]
], 1);
incorrect_answers:rand_selection([
[4, false, "It is affected by extreme values"],
[5, false, "It cannot be computed in grouped data with open-ended class intervals"],
[6, false, "It does not deal with qualitative data that cannot be expressed numerically"],
[7, false, "It cannot be computed for grouped data in the cases of open classes
"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Which statement DOES NOT represent a disadvantage of using the mean?</p>

[[ if test="member(1, ta_index)" ]]

<p>The desired algebraic property, refers to a specific property that is not possessed by the mean, and it's known as "linearity." The mean does not satisfy the property of linearity when applied to the sum or difference of random variables or datasets. This means that the mean of the sum (or difference) of two datasets is not generally equal to the sum (or difference) of the means of those datasets.

To illustrate this, let's consider two random variables \(X\) and \(Y\), and their respective means, denoted as \(E(X)\) and \(E(Y)\). The expected value (mean) of the sum of these two random variables is not necessarily equal to the sum of their means:

\(E(X + Y) ≠ E(X) + E(Y)\)</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]

<p>Computing the mean (average) in grouped data with open-ended class intervals can be challenging due to the lack of specific data points within those intervals. This is because the mean relies on knowing the individual data points, and when you have open-ended intervals, you don't have access to those specific values</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]

<p> "it is least affected by fluctuations of the sample" is not accurate when referring to the mean. In fact, the mean is quite sensitive to fluctuations in a sample. If you want a measure that is less influenced by fluctuations or outliers, it's better to consider the median or the mode.</p>

[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer"></p>
<p id="description">The question focuses on Identify the advantages of using mean where the gives a list of responses and the student is required to identify that which is not a disadvantage of using mean.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2065242451</deployedseed>
    <deployedseed>541377929</deployedseed>
    <deployedseed>1237152706</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
