<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interpreting a two-way frequency table</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Weather forecasts are often inaccurate. The following frequency table describes the occasions during a year in which rain was predicted for the following day and the actual occurrence of rain on that day.</p>
<p>On what proportion of days was the forecast {@forecast@} and the actual weather was {@actual@}?</p>

<p>
<style type="text/css">
.tg  {border-collapse:collapse;border-spacing:0;}
.tg td{font-family:Arial, sans-serif;font-size:14px;padding:10px 5px;border-style:solid;border-width:1px;overflow:hidden;word-break:normal;border-color:black;}
.tg th{font-family:Arial, sans-serif;font-size:14px;font-weight:normal;padding:10px 5px;border-style:solid;border-width:1px;overflow:hidden;word-break:normal;border-color:black;}
.tg .tg-baqh{text-align:center;vertical-align:top}
.tg .tg-c3ow{border-color:inherit;text-align:center;vertical-align:top}
.tg .tg-7btt{font-weight:bold;border-color:inherit;text-align:center;vertical-align:top}
.tg .tg-0pky{border-color:inherit;text-align:left;vertical-align:top}
.tg .tg-0lax{text-align:left;vertical-align:top}
</style>
<table class="tg">
  <tr>
    <th class="tg-c3ow"></th>
    <th class="tg-7btt" colspan="2">Actual</th>
    <th class="tg-0pky"></th>
  </tr>
  <tr>
    <td class="tg-7btt">Forecast</td>
    <td class="tg-c3ow">Rain</td>
    <td class="tg-c3ow">No Rain</td>
    <td class="tg-0pky"></td>
  </tr>
  <tr>
    <td class="tg-0pky">Rain</td>
    <td class="tg-c3ow">{@rain_rain@}<br></td>
    <td class="tg-c3ow">{@rain_norain@}</td>
    <td class="tg-0pky">{@f_rain@}</td>
  </tr>
  <tr>
    <td class="tg-0pky">No Rain</td>
    <td class="tg-c3ow">{@norain_rain@}</td>
    <td class="tg-c3ow">{@norain_norain@}</td>
    <td class="tg-0pky">{@f_norain@}</td>
  </tr>
  <tr>
    <td class="tg-0lax"></td>
    <td class="tg-baqh">{@a_rain@}</td>
    <td class="tg-baqh">{@a_norain@}</td>
    <td class="tg-baqh">{@total@}</td>
  </tr>
</table>
</p>

<p>
Proportion:
<table>
<caption></caption>
<thead>
<tr>
<th scope="col" style="text-align: center; " width="100">[[input:ans1]]</th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: center;border-top: 1px solid black;height: 60px;">[[input:ans2]]</td>
</tr>
</tbody>
</table>
</p>
<p>[[validation:ans1]][[validation:ans2]]</p>
[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>The question asks for the joint proportion for (forecast = {@forecast@}, actual = {@actual@}).
The proportion is therefore the count for the cell ({@forecast@}, {@actual@}) divided by the total count for the table, {@total@}.</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[dp(x, n) := float(round((x+10^(-8))*10^n)/10^n);

forecast: rand(["Rain","No Rain"]);
actual: rand(["Rain","No Rain"]);

rain_rain: 90+rand(50);
norain_norain: 150+rand(50);
rr: 365-rain_rain-norain_norain;
rain_norain: floor(rr*rand([0.2,0.25,0.3,0.35,0.4,0.45,0.55,0.6,0.65]));
norain_rain: rr-rain_norain;
total: rain_rain+norain_norain+rain_norain+norain_rain;
f_rain: rain_rain+rain_norain;
f_norain: norain_rain+norain_norain;
a_rain: rain_rain+norain_rain;
a_norain: rain_norain+norain_norain;

ta1: if forecast="Rain" and actual="Rain" then rain_rain elseif forecast="Rain" and actual="No Rain" then rain_norain elseif forecast="No Rain" and actual="Rain" then norain_rain elseif forecast="No Rain" and actual="No Rain" then norain_norain;
ta2: total;
ta: ta1/ta2;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>s_ans: ans1/ans2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>s_ans</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[Numerator is correct! You correctly identified the value inside the table corresponding to a forecast of {@forecast@} and the actual weather of {@actual@}.
<br>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[The numerator is not correct. The numerator should be the value inside the table corresponding to a forecast of {@forecast@} and the actual weather of {@actual@}. Find the value at the intersection of these.
<br>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[Denominator is correct. "On what proportion of days" means the proportion is out of all days.]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Denominator is not correct. The question says "On what proportion of days" which means the proportion is out of all days, so the denominator should be {@total@}.]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
