<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interpret the text and use the information to calculate conditional probabilities</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Of all the households in Kisumu County, {@prob_plasma@}% have a plasma TV and {@prob_laptop@}% have a Laptop computer. There are {@prob_both@}% of households that have both a plasma TV and a Laptop. Find the probability that a household chosen at random has either a plasma TV or a Laptop computer but not both. Give your answer in percentage.</p>
<p>[[input:ans1]]% [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To find the probability that a household chosen at random has either a plasma TV or a Laptop computer but not both. We need to find first the probability a household has only a laptop computer i.e P(plasmaTv)-P(both plasmaTV and laptop computer) then find the probability a household has only a plasma TV i.e P(Laptop computer)-P(both plasmaTV and laptop computer) and get the sum.</p>
<p>\( P (\text{Plasma TV or Laptop but not both}) = (P(\text{Plasma TV}) - P(\text{Both})) + (P(\text{Laptop computer}) - P(\text{Both}))\)</p>
<p>\(= ({@prob_plasma@}-{@prob_both@})+({@prob_laptop@}-{@prob_both@})\)</p>
<p>= \({@prob_plasma-prob_both@}+ {@prob_laptop-prob_both@}\)</p>
<p>= {@ta1@}%</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*random generating probabilities*/
prob_both:rand_with_step(10,30,5);
prob_plasma:prob_both+rand_with_step(10,50,5);
prob_laptop:prob_both+rand_with_step(10,100-prob_plasma-5,5);

/* model answer*/
ta1:(prob_plasma-prob_both) + (prob_laptop-prob_both);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>%</forbidwords>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checking correct answer*/

/*Node 2: Check if answer is not given in decimal form*/
n2_mistake:ta1/100;

/*Node 3: brief description of error*/
n3_mistake:prob_plasma + prob_laptop - prob_both;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>n2_mistake</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>&nbsp;&lt;p&gt;Your answer is correct but you have not given it in percentage as requested&lt;/p&gt;</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>n3_mistake</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span id="docs-internal-guid-82395a41-7fff-7f42-e054-aa77e89d1b04"></span></p><p dir="ltr" style="">&nbsp;&lt;p&gt;You computed the probability of having a Laptop, Plasma TV or both, while the questions asked for probability of having a Laptop, Plasma TV but not both. Instead take the total sum of probability of plasma TV and laptop&lt;/p&gt;</p><br>.<p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
