<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>LIFE TABLES (Single Decrement Tables)_q4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following is <b>NOT</b> a notation used in life table.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p><b>\(nZx,\)</b>: there is no such notation in the use of a life table.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p><b>\(_{n}I_{x}\):</b> there is no such notation in the use of a life table, only \(I_{x}\) with "\(n\)" subscript exists.
</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p><b>\(b_{x}\)</b>:  there is no such notation in the use of a life table.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a:2;
correct_answer:rand_selection([
[1, true, "\\(nZx\\)"],
[2, true, "\\(_{n}I_{x}\\)"],
[3, true, "\\(b_{x}\\)"]
],a);

b:3;
incorrect_answers:rand_selection([
[5, false, "\\(I_{x}\\)"],
[6, false, "\\(_{n}d_{x}\\)"],
[7, false, "\\(nq_{x}\\)"],
[8, false, "\\(L_{x}\\)"]
], b);

ta:random_permutation(append(correct_answer,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[[[ if test="member(1, ta_index)" ]]
<p><b>\(nZx,\)</b>: there is no such notation in the use of a life table.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p><b>\(_{n}I_{x}\):</b> there is no such notation in the use of a life table, only \(I_{x}\) with "\(n\)" subscript exists.
</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p><b>\(b_{x}\)</b>:  there is no such notation in the use of a life table.</p>
[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo<p>
<p id="concept">https://docs.google.com/document/d/1ioECBsnRwApAxz2ZzavQKR2YDOXMKT3ezRA5r5d8Cr4/edit</p>
<p id="reviewer">Mary</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">The question requires the student to identify the notations not used in life tables from a list of notations given.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>422821229</deployedseed>
    <deployedseed>2096880197</deployedseed>
    <deployedseed>536602932</deployedseed>
    <deployedseed>643177017</deployedseed>
    <deployedseed>703442279</deployedseed>
    <deployedseed>1285864149</deployedseed>
    <deployedseed>888346550</deployedseed>
    <deployedseed>1379542024</deployedseed>
    <deployedseed>145753024</deployedseed>
    <deployedseed>475744777</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
