<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>LIFE TABLES (Single Decrement Tables)_q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following is an assumption made in constructing a life table?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <i>Population is closed (net migration is zero).</i> </p>
<p>In the context of a life table, a closed population means that there is no migration into or out of the population being studied. Individuals are born, live, and die within the population without any external movement.</p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>Population is stationary (there is equal levels of mortality and fertility rates).</i></p>
<p>In a stationary population, the birth rate (fertility) is equal to the death rate (mortality). This means that on average, each individual that dies is replaced by a new birth, resulting in a stable population size. So, over time, the number of births balances out the number of deaths, maintaining a constant population size.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Deaths and births are evenly spread over the year.</i>
</p>
<p>While births and deaths may actually fluctuate throughout the year due to various factors such as seasonal variations, epidemics, or cultural practices, assuming an even distribution provides a reasonable approximation for the average mortality and fertility rates over time.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "Population is closed (net migration is zero).
"],
[2, true, "Population is stationary (there is equal levels of mortality and fertility rates).
"],
[3, true, "Deaths and births are evenly spread over the year.
"]],1);
incorrect_answers:rand_selection([
[4, false, "Population is stable."],
[5, false, "Population is aging."],
[6, false, "Population is constant."],
[7, false, "Deaths and births are constant over the year"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@ta@}</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept"></p>
<p id="reviewer">Mary</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question requires the student to identify the underlying assumption pivotal in the construction process of a life table.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>77</deployedseed>
    <deployedseed>751145021</deployedseed>
    <deployedseed>2138445174</deployedseed>
    <deployedseed>176406133</deployedseed>
    <deployedseed>1265362524</deployedseed>
    <deployedseed>1676429480</deployedseed>
    <deployedseed>514312335</deployedseed>
    <deployedseed>1183696066</deployedseed>
    <deployedseed>2098692860</deployedseed>
    <deployedseed>543882866</deployedseed>
    <deployedseed>1683713566</deployedseed>
    <deployedseed>405800180</deployedseed>
    <deployedseed>289368978</deployedseed>
    <deployedseed>1093743956</deployedseed>
    <deployedseed>1600529688</deployedseed>
    <deployedseed>788836605</deployedseed>
    <deployedseed>464832602</deployedseed>
    <deployedseed>2004280262</deployedseed>
    <deployedseed>991543283</deployedseed>
    <deployedseed>183921785</deployedseed>
    <deployedseed>1347407595</deployedseed>
    <deployedseed>746897481</deployedseed>
    <deployedseed>147880452</deployedseed>
    <deployedseed>735531074</deployedseed>
    <deployedseed>1434808750</deployedseed>
    <deployedseed>999585494</deployedseed>
    <deployedseed>1446097053</deployedseed>
    <deployedseed>398620129</deployedseed>
    <deployedseed>781879106</deployedseed>
    <deployedseed>873801758</deployedseed>
    <deployedseed>1361519054</deployedseed>
    <deployedseed>555373133</deployedseed>
    <deployedseed>1703615465</deployedseed>
    <deployedseed>1643686282</deployedseed>
    <deployedseed>1204424494</deployedseed>
    <deployedseed>966930163</deployedseed>
    <deployedseed>1858781896</deployedseed>
    <deployedseed>525390502</deployedseed>
    <deployedseed>1604779223</deployedseed>
    <deployedseed>1933224900</deployedseed>
    <deployedseed>1055007290</deployedseed>
    <deployedseed>1388373144</deployedseed>
    <deployedseed>2116785590</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
