<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>LINEAR REGRESSION-Q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the following data summaries:</p>
<ul>
<li>\( n={@n@}\)</li>
<li> \(\sum{x}={@sum_x@}\)</li>
<li> \(\sum{y}={@sum_y@}\)</li>
<li> \(\sum{xy}={@sum_xy@}\)</li>
<li> \(\sum{x^2}={@sum_x2@}\)</li>
<li>\(\sum{y^2}={@sum_y2@}\)</li>
</ul>
<p> Find the equation of the regression line:</p>
<p>a) \(\textbf{Y}\) on \(\textbf{X}\), given by.. \(\textbf{Y}=a+b\textbf{X}\)   </p>
<p>\( \textbf{Y}=\) [[input:ans1]] [[validation:ans1]]. [[feedback:prt1]]</p>
<p>b) \(\textbf{X}\) on \(\textbf{Y}\), given by \(\textbf{X}= \alpha+\beta \textbf{Y}\)   </p>
<p>\(\textbf{X}=\) [[input:ans2]] [[validation:ans2]]. [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
    <p>For the data with the following summaries:</p>
  <ul>
        <li>\(n = {@n@}\)</li>
        <li> \( \sum{x}= {@sum_x@}\)</li>
        <li>\( \sum{y} = {@sum_y@}\)</li>
        <li>\( \sum{xy} = {@sum_xy@}\)</li>
        <li>\( \sum{x^2} = {@sum_x2@}\)</li>
        <li> \(\sum{y^2} = {@sum_y2@}\)</li>
    </ul>
    <p>To find the equation of the regression line for \(Y\) on \(X\) and \(X\) on \(Y\), we can use the formulas for the slope \((b\) or \(\beta)\) and the intercept \((a\) or \(\alpha)\)  of the regression lines. The formulas for simple linear regression are as follows:</p>
<br>
    <p>a) For \(\textbf{Y}\) on \(\textbf{X}\), given by \(\textbf{Y}=a+b\textbf{X}\)  </p>
    <p>\( \displaystyle b = \frac{n\sum{xy} - (\sum{x})(\sum{y})}{n\sum{x^2} - (\sum{x})^2}\)</p>
    <p>\( \displaystyle b ={@b@} \)</p>
    <p>\(b \approx {@decimalplaces(b,2)@} \)</p>

    <p>Now, calculate \(a\):</p>
    <p>\( \displaystyle a = \frac{\sum{y} - b\sum{x}}{n}\)</p>
     <p>\( \displaystyle a ={@a@} \)</p>
    <p>\(a \approx {@decimalplaces(a,2)@} \)</p>

    <p>So, the equation for \(\textbf{Y}\) on \(\textbf{X}\) is approximately:</p>
    <p>\( \displaystyle \textbf{Y} ={@ta1@} \)</p>
    <p>\(\textbf{Y}  \approx {@alt_ta1@}\)</p>

<br>
    <p>b) For \(\textbf{X}\) on \(\textbf{Y}\), given by \(\textbf{X}=\alpha + \beta \textbf{Y}\)  </p>
    <p>\(  \displaystyle \beta = \frac{n\sum{xy} -( \sum{x})(\sum{y})}{n\sum{y^2} - (\sum{y})^2}\)</p>
    <p>\( \displaystyle \beta  ={@beta@} \)</p>
    <p>\( \beta  \approx {@decimalplaces(beta,2)@} \)</p>

    <p>Now, calculate α:</p>
    <p>\(  \displaystyle \alpha = \frac{\sum{x} - \beta\sum{y}}{n}\)</p>
     <p>\(\displaystyle \alpha  ={@alpha@} \)</p>
    <p>\( \alpha \approx {@decimalplaces(alpha,2)@} \)</p>

    <p>So, the equation for \(\textbf{X}\) on \(\textbf{Y}\) is approximately:</p>
    <p>\( \displaystyle \textbf{X} ={@ta2@} \)</p>
    <p>\(\textbf{X}  \approx {@alt_ta2@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the data*/ 
XX:makelist(rand(20),i,1,7); /*Number of absences*/
YY:makelist(rand(100),i,1,7); /*Final grade */
xy:XX * YY;
x_sq:XX^2;
y_sq:YY^2;
/* Calculate the number of data points (n) */ 
n: length(XX);  
/* Calculate the sums of X and Y */ 
sumX: sum(XX[i], i, 1, n); 
sumY: sum(YY[i], i, 1, n);  
/* Calculate the sum of the products of X and Y */ 
sumXY: sum(XX[i] * YY[i], i, 1, n);  
/* Calculate the sum of the squares of X and Y */ 
sumX2: sum(XX[i]^2, i, 1, n); 
sumY2: sum(YY[i]^2, i, 1, n);  

/* Given data */ 
n: n; 
sum_x: sumX; 
sum_y: sumY; 
sum_xy: sumXY; 
sum_x2: sumX2;  
/* Calculate the slope (b) */ 
b: (n * sum_xy - sum_x * sum_y) / (n * sum_x2 - sum_x^2);  
/* Calculate the intercept (a) */ 
a: (sum_y - b * sum_x) / n; 
/* Define the regression equation Y on X */ 
Y_on_X: a + b * X;


/* Given data */ 
 sum_y2: sumY2;  
/* Calculate the slope (beta) */ 
beta: (n * sum_xy - sum_x * sum_y) / (n * sum_y2 - sum_y^2);  
/* Calculate the intercept (alpha) */ 
alpha: (sum_x - beta * sum_y) / n;  
/* Define the regression equation X on Y */ 
X_on_Y: alpha + beta * Y; 

/*Model answers*/
ta1:Y_on_X;
alt_ta1: decimalplaces(a,2) + decimalplaces(b,2)*X;
ta2:X_on_Y;
alt_ta2: decimalplaces(alpha,2) + decimalplaces(beta,2)*Y;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( n={@n@}, \sum{x}={@sum_x@}, \sum{y}={@sum_y@}, \sum{xy}={@sum_xy@}, \sum{x^2}={@sum_x2@}, \sum{y^2}={@sum_y2@}\)</p>
<p> Find the equation of the regression line:</p>
<p>a) \(\textbf{Y}\) on \(\textbf{X}\), given by.. \(\textbf{Y}=a+b\textbf{X}\)   </p>
<p>\( \textbf{Y}={@ta1@}\)</p>
<p>b) \(\textbf{X}\) on \(\textbf{Y}\), given by \(\textbf{X}= \alpha+\beta \textbf{Y}\)   </p>
<p>\(\textbf{X}={@ta2@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire.</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description"> Given values of sum of x,y and the mean of x and y, value of n and the sum of their squares, the question requires the students to find the equation of the regression</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer 2dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>alt_ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer 2dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>alt_ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>583700453</deployedseed>
    <deployedseed>892939632</deployedseed>
    <deployedseed>1108750488</deployedseed>
    <deployedseed>222837215</deployedseed>
    <deployedseed>1109707024</deployedseed>
    <deployedseed>2070622469</deployedseed>
    <deployedseed>638485177</deployedseed>
    <deployedseed>1835732972</deployedseed>
    <deployedseed>1434643747</deployedseed>
    <deployedseed>583991390</deployedseed>
    <deployedseed>419987551</deployedseed>
    <deployedseed>1266862337</deployedseed>
    <deployedseed>1850206992</deployedseed>
    <deployedseed>492025424</deployedseed>
    <deployedseed>1788723596</deployedseed>
    <deployedseed>977774707</deployedseed>
    <deployedseed>1435737926</deployedseed>
    <deployedseed>1306460660</deployedseed>
    <deployedseed>1369826140</deployedseed>
    <deployedseed>717785444</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>alt_ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>alt_ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
