<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>LINEAR REGRESSION-Q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>To explore the potential relationship between a father's weight (in kg) and his son's weight (in kg) and derive the equation of the regression line, a physician collected the data below.</p>
<br>
<p>
<table border="solid",width="80%","text-align:center">
  <tr>
<th>\( \text {Father’s weight (x)}\)</th>
<th>\({@x_1@}\)</th>
<th>\({@x_2@}\)</th>
<th>\({@x_3@}\)</th>
<th>\({@x_4@}\)</th>
<th>\({@x_5@}\)</th>
<th>\({@x_6@}\)</th>
<th>\({@x_7@}\)</th>
<th>\({@x_8@}\)</th>
</tr>
<tr>
<th> \(\text{Son’s weight (y)}\)</th>
  <th>\({@y_1@}\)</th>
  <th>\({@y_2@}\)</th>
  <th>\({@y_3@}\)</th>
  <th>\({@y_4@}\)</th>
  <th>\({@y_5@}\)</th>
  <th>\({@y_6@}\)</th>
  <th>\({@y_7@}\)</th>
  <th>\({@y_8@}\)</th>
  </tr>
</table>
</p>


<p>Using the data in the table above fill in the blanks below.</p>
<p>
<table border="solid", width="80%", style="text-align:center">
  <tr>
    <th width="80%">\(x-M_x \)</th>
    <th width="80%">\(y-M_y\)</th>
    <th width="80%">\((x-M_x)^2\)</th>
    <th width="80%">\((x-M_x) \times (y-M_y)\)</th> 
  </tr>
  <tr>
    <td>\({@decimalplaces(x_1-m_x,2)@}\)</td>
    <td>[[input:ans1]] [[validation:ans1]][[feedback:prt1]] </td>
    <td>{@(decimalplaces(x_1-m_x,2))^2@}</td>
    <td>{@decimalplaces(x_1-m_x,2)*(decimalplaces(y_2-m_y,2))@}</td>
  
  </tr>
   <tr>
    <td>[[input:ans2]] [[validation:ans2]][[feedback:prt2]] </td>
    <td>{@decimalplaces((y_2-m_y),2)@}</td>
    <td>{@(decimalplaces(x_2-m_y,2))^2@}</td>
    <td>{@decimalplaces(x_2-m_y,2)*(decimalplaces((y_2-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces(x_3-m_x,2)@}</td>
    <td>{@decimalplaces((y_3-m_y),2)@}</td>
    <td>[[input:ans3]] [[validation:ans3]][[feedback:prt3]] </td>
    <td>{@decimalplaces(x_3-m_x,2)*(decimalplaces((y_3-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_4-m_x),2)@}</td>
    <td>{@decimalplaces((y_4-m_y),2)@}</td>
    <td>[[input:ans4]] [[validation:ans4]][[feedback:prt4]] </td>
    <td>{@decimalplaces(decimalplaces((x_4-m_x),2)*(decimalplaces((y_4-m_y),2)),3)@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_5-m_x),2)@}</td>
    <td>{@decimalplaces((y_5-m_y),2)@}</td>
    <td>{@(decimalplaces(x_5-m_x,2))^2@}</td>
    <td>[[input:ans5]] [[validation:ans5]][[feedback:prt5]]</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_6-m_x),2)@}</td>
    <td>[[input:ans6]] [[validation:ans6]][[feedback:prt6]]</td>
    <td>{@(decimalplaces(x_6-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_6-m_x),2)*(decimalplaces((y_6-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_7-m_x),2)@}</td>
    <td>{@decimalplaces((y_7-m_y),2)@}</td>
    <td>{@(decimalplaces(x_7-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_7-m_x),2)*(decimalplaces((y_7-m_y),2))@}</td>
     
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_8-m_x),2)@}</td>
    <td>{@decimalplaces((y_8-m_y),2)@}</td>
    <td>{@(decimalplaces(x_8-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_8-m_x),2)*(decimalplaces((y_8-m_y),2))@}</td>
   
  </tr>
  <tr>
    <td></td>
    <td></td>
    <td>\(\sum (x-M_x)^2 \) [[input:ans7]] [[validation:ans7]][[feedback:prt7]]</td>
    <td>\(\sum (x-M_x) \times (y-M_y) \) [[input:ans8]] [[validation:ans8]][[feedback:prt8]] </td>
  </tr>
</table>



</p>
<p>The Regression Equation \(\displaystyle ŷ = bX + a\)</p>
<p>\(\textbf{a}=\)  [[input:ans9]] [[validation:ans9]][[feedback:prt9]]  </p>
<p>\(\textbf{b}=\)  [[input:ans10]] [[validation:ans10]][[feedback:prt10]]</p>
<p>\(\textbf{ŷ}=\)  [[input:ans11]] [[validation:ans11]][[feedback:prt11]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3></h3>
<p>The table should be filled as follows;</p>
<hr>
<p>
<table border="solid", width="100%", style="text-align:center">
  <tr>
    <th  width="60%" >\(x-M_x\)</th>
    <th  width="60%" >\(y-M_y\)</th>
    <th  width="60%" >\((x-M_x)^2\)</th>
    <th  width="60%" >\((x-M_x) \times (y-M_y)\)</th> 
  </tr>
  <tr>
    <td>{@decimalplaces(x_1-m_x,2)@}</td>
    <td>{@decimalplaces(y_2-m_y,2)@}</td>
    <td>{@(decimalplaces(x_1-m_x,2))^2@}</td>
    <td>{@decimalplaces(x_1-m_x,2)*(decimalplaces(y_2-m_y,2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces(x_2-m_y,2)@}</td>
    <td>{@decimalplaces((y_2-m_y),2)@}</td>
    <td>{@(decimalplaces(x_2-m_y,2))^2@}</td>
    <td>{@decimalplaces(x_2-m_y,2)*(decimalplaces((y_2-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces(x_3-m_x,2)@}</td>
    <td>{@decimalplaces((y_3-m_y),2)@}</td>
    <td>{@(decimalplaces(x_3-m_x,2))^2@}</td>
    <td>{@decimalplaces(x_3-m_x,2)*(decimalplaces((y_3-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_4-m_x),2)@}</td>
    <td>{@decimalplaces((y_4-m_y),2)@}</td>
    <td>{@(decimalplaces(x_4-m_x,2))^2@}</td>
    <td>{@decimalplaces(decimalplaces((x_4-m_x),2)*(decimalplaces((y_4-m_y),2)),3)@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_5-m_x),2)@}</td>
    <td>{@decimalplaces((y_5-m_y),2)@}</td>
    <td>{@(decimalplaces(x_5-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_5-m_x),2)*(decimalplaces((y_5-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_6-m_x),2)@}</td>
    <td>{@decimalplaces((y_6-m_y),2)@}</td>
    <td>{@(decimalplaces(x_6-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_6-m_x),2)*(decimalplaces((y_6-m_y),2))@}</td>
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_7-m_x),2)@}</td>
    <td>({@decimalplaces((y_7-m_y),2)@}</td>
    <td>{@(decimalplaces(x_7-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_7-m_x),2)*(decimalplaces((y_7-m_y),2))@}</td>
     
  
  </tr>
   <tr>
    <td>{@decimalplaces((x_8-m_x),2)@}</td>
    <td>{@decimalplaces((y_8-m_y),2)@}</td>
    <td>{@(decimalplaces(x_8-m_x,2))^2@}</td>
    <td>{@decimalplaces((x_8-m_x),2)*(decimalplaces((y_8-m_y),2))@}</td>
   
  </tr>
  <tr>
    <td></td>
    <td></td>
    <td>\(\sum={@ss_x@}\) </td>
    <td>\(\sum={@sp@}\) </td>
  </tr>
</table>



</p>


<hr>
<p> Recall:</p>
<p>The mean of the father's weights  \( M_x = {@decimalplaces(m_x,2)@}\)</p>
<p>The mean of the son's weights  \( M_y = {@decimalplaces(m_y,2)@}\)</p>
<p></p>
<p>Therefore:</p>
<p>\(b=\displaystyle\frac{(x-M_x) \times (y-M_y)}{(x-M_x)^2} = \frac{{@sp@}}{{@ss_x@}} = {@decimalplaces(value_b,2)@}\)</p>

<p>\( a = M_y - bM_x ={@decimalplaces(value_a,2)@}\)</p>

<p>The Regression Equation is written as \(\displaystyle ŷ = bX + a\)</p>
<p>\(\displaystyle {ŷ}={@value_y@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*line of regerssion*/
/*values of x*/
x_1:rand(5)+60;
x_2:x_1+3;
x_3:x_1+6;
x_4:x_1+2;
x_5:x_1+4;
x_6:x_1+6;
x_7:x_1+5;
x_8:x_1+9;
/*values of y*/
y_1:rand(6)+60;
y_2:y_1+1;
y_3:y_1+2;
y_4:y_1+3;
y_5:y_1-4;
y_6:y_1+5;
y_7:y_1+6;
y_8:y_1-7;
/*solutions x*/
m_x:(x_1+x_2+x_3+x_4+x_5+x_6+x_7+x_8)/(8);
m_y:(y_1+y_2+y_3+y_4+y_5+y_6+y_7+y_8)/8;
/*x-m_x*/
x_m_x:x_1-m_x;x_2-m_x;x_3-m_x;x_4-m_x;x_5-m_x;x_6-m_x;x_7-m_x;x_8-m_x;
y_m_y:y_1-m_y;y_2-m_y;y_3-m_y;y_4-m_y;y_5-m_y;y_6-m_y;y_7-m_y;y_8-m_y;
ss_x_1:((x_1-m_x)^2+(x_2-m_x)^2+(x_3-m_x)^2+(x_4-m_x)^2+(x_5-m_x)^2+(x_6-m_x)^2);
ss_x_2:((x_7-m_x)^2+(x_8-m_x)^2);
ss_x:(decimalplaces(x_1-m_x,2))^2+(decimalplaces(x_2-m_x,2))^2+(decimalplaces(x_3-m_x,2))^2+(decimalplaces(x_4-m_x,2))^2+(decimalplaces(x_5-m_x,2))^2+(decimalplaces(x_6-m_x,2))^2+(decimalplaces(x_7-m_x,2))^2+(decimalplaces(x_8-m_x,2))^2;
/*sp*/
sol_p:decimalplaces(x_1-m_x,2)*(decimalplaces(y_2-m_y,2))+decimalplaces(x_2-m_y,2)*(decimalplaces((y_2-m_y),2))+decimalplaces(x_3-m_x,2)*(decimalplaces((y_3-m_y),2))+decimalplaces((x_4-m_x),2)*(decimalplaces((y_4-m_y),2))+decimalplaces((x_5-m_x),2)*(decimalplaces((y_5-m_y),2))+decimalplaces((x_6-m_x),2)*(decimalplaces((y_6-m_y),2))+decimalplaces((x_7-m_x),2)*(decimalplaces((y_7-m_y),2))+decimalplaces((x_8-m_x),2)*(decimalplaces((y_8-m_y),2));
sp:decimalplaces(sol_p,3);

/*other variables*/
value_b:decimalplaces(sp/ss_x,2);
value_a:decimalplaces(m_y-(value_b*m_x),2);
value_y:(value_b)*x+value_a;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The Regression Equation \( ŷ = bX + a\)</p>
<p>\(\textbf{a}={@value_a@}\) </p>
<p>\(\textbf{b}={@value_b@}\)  </p>
<p>\(\textbf{ŷ}={@value_y@}\)  </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer"></p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">The question focuses on finding the regression equation of the weight of a father and son given the weights of the father and that of the son </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>decimalplaces(y_2-m_y,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>value_b</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>value_y</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>decimalplaces(x_2-m_y,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>(decimalplaces(x_3-m_x,2))^2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>(decimalplaces(x_4-m_x,2))^2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>decimalplaces((x_5-m_x),2)*(decimalplaces((y_5-m_y),2))</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>decimalplaces((y_6-m_y),2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ss_x</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>sp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>value_a</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>decimalplaces(y_2-m_y,2)</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>y_2-m_y</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>value_b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>value_y</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer to  2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>decimalplaces(x_2-m_y,2)</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for corrrect answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>x_2-m_y</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>(decimalplaces(x_3-m_x,2))^2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>(x_3-m_x)^2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>(decimalplaces(x_4-m_x,2))^2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for 2dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>(x_4-m_x)^2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for 2dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>decimalplaces((x_5-m_x),2)*(decimalplaces((y_5-m_y),2))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description> checks for 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans6</sans>
        <tans>decimalplaces((y_6-m_y),2)</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ss_x</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>sp</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>value_a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>894756463</deployedseed>
    <deployedseed>2034989337</deployedseed>
    <deployedseed>256039462</deployedseed>
    <deployedseed>1947630959</deployedseed>
    <deployedseed>1478786928</deployedseed>
    <deployedseed>697938406</deployedseed>
    <deployedseed>907125095</deployedseed>
    <deployedseed>802881401</deployedseed>
    <deployedseed>902523515</deployedseed>
    <deployedseed>22976202</deployedseed>
    <deployedseed>919382811</deployedseed>
    <deployedseed>756892699</deployedseed>
    <deployedseed>1228160656</deployedseed>
    <deployedseed>501034314</deployedseed>
    <deployedseed>818061387</deployedseed>
    <deployedseed>570938560</deployedseed>
    <deployedseed>1871361629</deployedseed>
    <deployedseed>1563386113</deployedseed>
    <deployedseed>1932320464</deployedseed>
    <deployedseed>7758861</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>decimalplaces(y_2-m_y,2)</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>value_b</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>value_y</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>decimalplaces(x_2-m_y,2)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>(decimalplaces(x_3-m_x,2))^2</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>(decimalplaces(x_4-m_x,2))^2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>decimalplaces((x_5-m_x),2)*(decimalplaces((y_5-m_y),2))</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>decimalplaces((y_6-m_y),2)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ss_x</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>sp</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>value_a</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt4-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
