<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>multilevel data, Kenya child-health programme</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A Kenyan child-health programme is being evaluated using survey data collected from children in several counties. Survey clusters received either <strong>Programme A</strong> or <strong>Programme B</strong>. We want to investigate which programme is associated with a higher proportion of fully vaccinated children.</p>

<p>Download the dataset for your question below. Each row represents one child.</p>

<p><a href="[[textdownload name="kenya_child_health.csv"]]{@stack_csv_formatter(kenya_data,kenya_columns)@}[[/textdownload]]">Download the CSV dataset</a></p>

<p>Use a statistical software package to analyse the data and answer the following questions. Give all percentages to <strong>one decimal place</strong>.</p>

<h5>a. Programme A</h5>

<p>What percentage of children receiving <strong>Programme A</strong> were fully vaccinated?</p>

<p>[[input:ans1]] % [[validation:ans1]]</p>
[[feedback:prt1]]

<h5>b. Programme B</h5>

<p>What percentage of children receiving <strong>Programme B</strong> were fully vaccinated?</p>

<p>[[input:ans2]] % [[validation:ans2]]</p>
[[feedback:prt2]]

<h5>c. Overall comparison</h5>

<p>Based on these overall vaccination percentages, which programme appears to have the higher vaccination rate?</p>

<p>[[input:ans3]] [[validation:ans3]]</p>
[[feedback:prt3]]

<h5>d. Vaccination rates within wealth quintiles</h5>

<p>Now calculate the percentage of fully vaccinated children for each programme within each wealth quintile. Give all percentages to <strong>one decimal place</strong>.</p>

<table class="table">
    <thead>
        <tr>
            <th>Wealth quintile</th>
            <th>Programme A (%)</th>
            <th>Programme B (%)</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>1</td>
            <td>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td>
            <td>[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</td>
        </tr>
        <tr>
            <td>2</td>
            <td>[[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</td>
            <td>[[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</td>
        </tr>
        <tr>
            <td>3</td>
            <td>[[input:ans8]] [[validation:ans8]] [[feedback:prt8]]</td>
            <td>[[input:ans9]] [[validation:ans9]] [[feedback:prt9]]</td>
        </tr>
        <tr>
            <td>4</td>
            <td>[[input:ans10]] [[validation:ans10]] [[feedback:prt10]]</td>
            <td>[[input:ans11]] [[validation:ans11]] [[feedback:prt11]]</td>
        </tr>
        <tr>
            <td>5</td>
            <td>[[input:ans12]] [[validation:ans12]] [[feedback:prt12]]</td>
            <td>[[input:ans13]] [[validation:ans13]] [[feedback:prt13]]</td>
        </tr>
    </tbody>
</table>

<h5>e. Comparison within wealth quintiles</h5>

<p>Which programme has the higher vaccination rate within <strong>every wealth quintile</strong>?</p>

<p>[[input:ans14]] [[validation:ans14]]</p>
[[feedback:prt14]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The calculations in this question can be carried out using a statistical software package. The key quantities are the proportions of fully vaccinated children, first for each programme overall and then separately within each wealth quintile.</p>

<h5>a. Programme A</h5>

<p>The percentage of children receiving <strong>Programme A</strong> who were fully vaccinated is {@float(round(10*ta1)/10)@}%.</p>

<h5>b. Programme B</h5>

<p>The percentage of children receiving <strong>Programme B</strong> who were fully vaccinated is {@float(round(10*ta2)/10)@}%.</p>

<h5>c. Overall comparison</h5>

<p>Based on the overall vaccination percentages, <strong>Programme A</strong> appears to have the higher vaccination rate.</p>

<h5>d. Vaccination rates within wealth quintiles</h5>

<p>When the vaccination percentages are calculated separately within each wealth quintile, the results are:</p>

<table class="table">
    <thead>
        <tr>
            <th>Wealth quintile</th>
            <th>Programme A (%)</th>
            <th>Programme B (%)</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>1</td>
            <td>{@float(round(10*ta4)/10)@}</td>
            <td>{@float(round(10*ta5)/10)@}</td>
        </tr>
        <tr>
            <td>2</td>
            <td>{@float(round(10*ta6)/10)@}</td>
            <td>{@float(round(10*ta7)/10)@}</td>
        </tr>
        <tr>
            <td>3</td>
            <td>{@float(round(10*ta8)/10)@}</td>
            <td>{@float(round(10*ta9)/10)@}</td>
        </tr>
        <tr>
            <td>4</td>
            <td>{@float(round(10*ta10)/10)@}</td>
            <td>{@float(round(10*ta11)/10)@}</td>
        </tr>
        <tr>
            <td>5</td>
            <td>{@float(round(10*ta12)/10)@}</td>
            <td>{@float(round(10*ta13)/10)@}</td>
        </tr>
    </tbody>
</table>

<h5>e. Comparison within wealth quintiles</h5>

<p><strong>Programme B</strong> has the higher vaccination rate within every wealth quintile. This is the opposite of the comparison obtained from the overall vaccination percentages.</p>

<p>The apparent reversal occurs because the two programmes have different distributions of children across the wealth quintiles. Programme A contains a larger proportion of children from the higher wealth quintiles, where vaccination rates are higher, whereas Programme B contains a larger proportion from the lower wealth quintiles. Consequently, comparing the programmes without taking wealth quintile into account gives a different result from comparing them within each wealth quintile.</p>

<p>This is an example of <strong>Simpson's paradox</strong>: an association observed in aggregated data can reverse when the data are examined within relevant groups. It illustrates why the structure of the data and potentially important grouping variables should be considered when interpreting an overall comparison.</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* ============================================================
   KENYA VACCINATION DATA
   Direct-construction version
   ============================================================ */


/* ------------------------------------------------------------
   1. Parameters
   ------------------------------------------------------------ */

counties:["Kakamega","Kisumu","Machakos","Turkana"];

A_households:[8,11,19,38,75];
B_households:[75,38,19,11,8];

A_prob:[55/100,65/100,75/100,85/100,90/100];
B_prob:[70/100,80/100,90/100,95/100,98/100];


/* ------------------------------------------------------------
   2. Cluster counties
   ------------------------------------------------------------ */

/*
   Clusters 1--4 belong to Programme A.
   Clusters 5--8 belong to Programme B.
*/

cluster_counties:
    makelist(
        counties[rand(4)+1],
        i,1,8
    );


/* ------------------------------------------------------------
   3. Initialise dataset
   ------------------------------------------------------------ */

/*
   Final child-level row structure:

   [cluster_id,
    county,
    programme,
    household_id,
    wealth_quintile,
    child_id,
    age_months,
    sex,
    fully_vaccinated]
*/

kenya_data:[];

household_id:0;
child_id:0;


/* ------------------------------------------------------------
   4. Initialise summary variables
   ------------------------------------------------------------ */

/*
   These scalar variables store the number of children and
   number vaccinated in each programme x wealth-quintile cell.

   We calculate them during data generation rather than
   searching kenya_data afterwards.
*/

A_n1:0;
A_n2:0;
A_n3:0;
A_n4:0;
A_n5:0;

A_vacc1:0;
A_vacc2:0;
A_vacc3:0;
A_vacc4:0;
A_vacc5:0;

B_n1:0;
B_n2:0;
B_n3:0;
B_n4:0;
B_n5:0;

B_vacc1:0;
B_vacc2:0;
B_vacc3:0;
B_vacc4:0;
B_vacc5:0;


/* ------------------------------------------------------------
   5. Programme A
   ------------------------------------------------------------ */

for q:1 thru 5 do block(
    [cell_households,
     n_cell,
     n_vaccinated,
     vaccination,
     hh_size,
     cluster_id,
     county,
     sex,
     age_months,
     k],

    /*
       Generate household sizes.

       Each household independently contains 4, 5 or 6 children.
    */

    cell_households:
        makelist(
            rand(3)+4,
            h,1,A_households[q]
        ),

    /*
       Number of children in this programme x wealth cell.
    */

    n_cell:
        sum(
            cell_households[h],
            h,1,A_households[q]
        ),

    /*
       Exact number vaccinated in the cell.
    */

    n_vaccinated:
        round(n_cell*A_prob[q]),

    /*
       Store the summary counts as scalar variables.
    */

    if q=1 then (
        A_n1:n_cell,
        A_vacc1:n_vaccinated
    ),

    if q=2 then (
        A_n2:n_cell,
        A_vacc2:n_vaccinated
    ),

    if q=3 then (
        A_n3:n_cell,
        A_vacc3:n_vaccinated
    ),

    if q=4 then (
        A_n4:n_cell,
        A_vacc4:n_vaccinated
    ),

    if q=5 then (
        A_n5:n_cell,
        A_vacc5:n_vaccinated
    ),

    /*
       Construct and randomly permute the vaccination vector.
    */

    vaccination:
        random_permutation(
            append(
                makelist(
                    1,
                    i,1,n_vaccinated
                ),
                makelist(
                    0,
                    i,1,n_cell-n_vaccinated
                )
            )
        ),

    k:1,

    /*
       Generate households and child-level records directly.
    */

    for h:1 thru A_households[q] do (

        household_id:household_id+1,

        hh_size:cell_households[h],

        cluster_id:rand(4)+1,

        county:cluster_counties[cluster_id],

        for c:1 thru hh_size do (

            child_id:child_id+1,

            age_months:rand(12)+12,

            sex:
                if rand(2)=0
                then "Female"
                else "Male",

            kenya_data:
                cons(
                    [
                        cluster_id,
                        county,
                        "A",
                        household_id,
                        q,
                        child_id,
                        age_months,
                        sex,
                        vaccination[k]
                    ],
                    kenya_data
                ),

            k:k+1
        )
    )
);


/* ------------------------------------------------------------
   6. Programme B
   ------------------------------------------------------------ */

for q:1 thru 5 do block(
    [cell_households,
     n_cell,
     n_vaccinated,
     vaccination,
     hh_size,
     cluster_id,
     county,
     sex,
     age_months,
     k],

    /*
       Generate household sizes.
    */

    cell_households:
        makelist(
            rand(3)+4,
            h,1,B_households[q]
        ),

    /*
       Number of children in this programme x wealth cell.
    */

    n_cell:
        sum(
            cell_households[h],
            h,1,B_households[q]
        ),

    /*
       Exact number vaccinated.
    */

    n_vaccinated:
        round(n_cell*B_prob[q]),

    /*
       Store summary counts as scalar variables.
    */

    if q=1 then (
        B_n1:n_cell,
        B_vacc1:n_vaccinated
    ),

    if q=2 then (
        B_n2:n_cell,
        B_vacc2:n_vaccinated
    ),

    if q=3 then (
        B_n3:n_cell,
        B_vacc3:n_vaccinated
    ),

    if q=4 then (
        B_n4:n_cell,
        B_vacc4:n_vaccinated
    ),

    if q=5 then (
        B_n5:n_cell,
        B_vacc5:n_vaccinated
    ),

    /*
       Construct and randomise vaccination outcomes.
    */

    vaccination:
        random_permutation(
            append(
                makelist(
                    1,
                    i,1,n_vaccinated
                ),
                makelist(
                    0,
                    i,1,n_cell-n_vaccinated
                )
            )
        ),

    k:1,

    /*
       Generate final child-level records.
    */

    for h:1 thru B_households[q] do (

        household_id:household_id+1,

        hh_size:cell_households[h],

        cluster_id:rand(4)+5,

        county:cluster_counties[cluster_id],

        for c:1 thru hh_size do (

            child_id:child_id+1,

            age_months:rand(12)+12,

            sex:
                if rand(2)=0
                then "Female"
                else "Male",

            kenya_data:
                cons(
                    [
                        cluster_id,
                        county,
                        "B",
                        household_id,
                        q,
                        child_id,
                        age_months,
                        sex,
                        vaccination[k]
                    ],
                    kenya_data
                ),

            k:k+1
        )
    )
);


/* ------------------------------------------------------------
   7. Randomise row order
   ------------------------------------------------------------ */

kenya_data:
    random_permutation(kenya_data);


/* ------------------------------------------------------------
   8. Column names and metadata
   ------------------------------------------------------------ */

kenya_columns:[
    "cluster_id",
    "county",
    "programme",
    "household_id",
    "wealth_quintile",
    "child_id",
    "age_months",
    "sex",
    "fully_vaccinated"
];

n_households:household_id;
n_children:child_id;


/* ------------------------------------------------------------
   9. Overall summary counts
   ------------------------------------------------------------ */

A_total_n:
    A_n1+A_n2+A_n3+A_n4+A_n5;

A_total_vacc:
    A_vacc1+A_vacc2+A_vacc3+A_vacc4+A_vacc5;

B_total_n:
    B_n1+B_n2+B_n3+B_n4+B_n5;

B_total_vacc:
    B_vacc1+B_vacc2+B_vacc3+B_vacc4+B_vacc5;


/* ------------------------------------------------------------
   10. Teacher answers
   ------------------------------------------------------------ */

/*
   Overall vaccination percentages.
*/

ta1:
    100*A_total_vacc/A_total_n;

ta2:
    100*B_total_vacc/B_total_n;


/*
   Overall apparent winner.

   Multiple-choice/radio teacher answer.
*/

ta3:[
    [1,true,"Programme A"],
    [2,false,"Programme B"]
];


/*
   Wealth quintile 1.
*/

ta4:
    100*A_vacc1/A_n1;

ta5:
    100*B_vacc1/B_n1;


/*
   Wealth quintile 2.
*/

ta6:
    100*A_vacc2/A_n2;

ta7:
    100*B_vacc2/B_n2;


/*
   Wealth quintile 3.
*/

ta8:
    100*A_vacc3/A_n3;

ta9:
    100*B_vacc3/B_n3;


/*
   Wealth quintile 4.
*/

ta10:
    100*A_vacc4/A_n4;

ta11:
    100*B_vacc4/B_n4;


/*
   Wealth quintile 5.
*/

ta12:
    100*A_vacc5/A_n5;

ta13:
    100*B_vacc5/B_n5;


/*
   Programme with the higher vaccination percentage
   within every wealth quintile.

   Multiple-choice/radio teacher answer.
*/

ta14:[
    [1,false,"Programme A"],
    [2,true,"Programme B"]
];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@n_children@} children; {@n_households@} households; Programme A higher overall; Programme B higher within every wealth quintile.</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>float(round(10*ta1)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>numerical</type>
      <tans>float(round(10*ta10)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>numerical</type>
      <tans>float(round(10*ta11)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>numerical</type>
      <tans>float(round(10*ta12)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>numerical</type>
      <tans>float(round(10*ta13)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14</name>
      <type>radio</type>
      <tans>ta14</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>float(round(10*ta2)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>radio</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>float(round(10*ta4)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>float(round(10*ta5)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>numerical</type>
      <tans>float(round(10*ta6)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>float(round(10*ta7)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>numerical</type>
      <tans>float(round(10*ta8)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>numerical</type>
      <tans>float(round(10*ta9)/10)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans10</sans>
        <tans>ta10</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans11</sans>
        <tans>ta11</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans12</sans>
        <tans>ta12</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt13</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans13</sans>
        <tans>ta13</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt13-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt13-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt14</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans14</sans>
        <tans>first(mcq_correct(ta14))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt14-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt14-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>first(mcq_correct(ta3))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans9</sans>
        <tans>ta9</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1722241041</deployedseed>
    <deployedseed>830728031</deployedseed>
    <deployedseed>1114633813</deployedseed>
    <deployedseed>772262201</deployedseed>
    <deployedseed>1341233461</deployedseed>
    <deployedseed>349862584</deployedseed>
    <deployedseed>1220861889</deployedseed>
    <deployedseed>1904297111</deployedseed>
    <deployedseed>1830000406</deployedseed>
    <deployedseed>235440781</deployedseed>
    <deployedseed>940950539</deployedseed>
    <deployedseed>518612876</deployedseed>
    <deployedseed>1310941416</deployedseed>
    <deployedseed>1525766317</deployedseed>
    <deployedseed>53313987</deployedseed>
    <deployedseed>948941607</deployedseed>
    <deployedseed>635402953</deployedseed>
    <deployedseed>1682962528</deployedseed>
    <deployedseed>571059838</deployedseed>
    <deployedseed>1291438536</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>float(round(10*ta1)/10)</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>float(round(10*ta10)/10)</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>float(round(10*ta11)/10)</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>float(round(10*ta12)/10)</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>float(round(10*ta13)/10)</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>first(mcq_correct(ta14))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>float(round(10*ta2)/10)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>float(round(10*ta4)/10)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>float(round(10*ta5)/10)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>float(round(10*ta6)/10)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>float(round(10*ta7)/10)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>float(round(10*ta8)/10)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>float(round(10*ta9)/10)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt10-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt11-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt12-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt13</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt13-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt14</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt14-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt2-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt3-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt4-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt5-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt6-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt7-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt8-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt9-1-T ]</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>float(round(10*ta1)/10)</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>float(round(10*ta10)/10)</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>float(round(10*ta11)/10)</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>float(round(10*ta12)/10)</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>float(round(10*ta13)/10)</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>first(mcq_correct(ta14))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>float(round(10*ta2)/10)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>-1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>float(round(10*ta5)/10)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>float(round(10*ta6)/10)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>float(round(10*ta7)/10)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>float(round(10*ta8)/10)</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>float(round(10*ta9)/10)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt13</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt13-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt14</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt14-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
