<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>POPULATION PROJECTION_q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the following female population of Trindad in \(1960\) and \(1975\), study it and answer the questions that follow. </p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age</th>
    <th width="15%", style="text-align:center">\(1960\)</th>
    <th width="15%", style="text-align:center">\(1975\)</th>
  </tr>
  <tr>
  <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\({@ageset_1a@}\)</td>
    <td style="text-align:center">\({@ageset_1b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\({@ageset_2a@}\)</td>
    <td style="text-align:center">\({@ageset_2b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\({@ageset_3a@}\)</td>
    <td style="text-align:center">\({@ageset_3b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(45-59\)</td>
    <td style="text-align:center">\({@ageset_4a@}\)</td>
    <td style="text-align:center">\({@ageset_4b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">60+</td>
    <td style="text-align:center">\({@ageset_5a@}\)</td>
    <td style="text-align:center">\({@ageset_5b@}\)</td>
  </tr>
</table>
</p>

<p>Note: Leave your answer correct to \(3\) dp</p>
<p>Determine the growth rates over the given years denoted as \(\text{p}\)</p>
<p>\(\begin{bmatrix}
f_1 &f_2  & f_3 & 0 & 0\\ 
p_1 & 0 & 0 & 0 &0 \\ 
0 & p_2 & 0 &0  &0 \\ 
 0& 0 & p_3 & 0 & 0\\ 
0 & 0 &0  & p_4 & p_5
\end{bmatrix}\begin{bmatrix}
v(1,t)\\ 
v(2,t)\\ 
v(3,t)\\ 
v(4,t)\\ 
v(5,t)
\end{bmatrix}\)</p>
<p>\(\text{p}_1=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\(\text{p}_2=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>\(\text{p}_3=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p>\(\text{p}_4=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>
<p>\(\text{p}_5=\) [[input:ans5]] [[validation:ans5]][[feedback:prt5]]</p>
<br>
<p>Using the growth rate above, project the population from \(1975-1990\)</p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age</th>
    <th width="15%", style="text-align:center">\(1975\)</th>
<th width="15%", style="text-align:center">\(1990\)</th>
  </tr>
  <tr>
  <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\({@ageset_1b@}\)</td>
    <td style="text-align:center">[[input:ans6]] [[validation:ans6]][[feedback:prt6]]</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\({@ageset_2b@}\)</td>
    <td style="text-align:center">[[input:ans7]] [[validation:ans7]][[feedback:prt7]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\({@ageset_3b@}\)</td>
    <td style="text-align:center">[[input:ans8]] [[validation:ans8]][[feedback:prt8]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(44-59\)</td>
    <td style="text-align:center">\({@ageset_4b@}\)</td>
    <td style="text-align:center">[[input:ans9]] [[validation:ans9]][[feedback:prt9]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(60+\)</td>
    <td style="text-align:center">\({@ageset_5b@}\)</td>
    <td style="text-align:center">[[input:ans10]] [[validation:ans10]][[feedback:prt10]]</td>
  </tr>
 </table>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>

<table border="solid" style="width: 100%;">
  <tr>
    <th width="25%" style="text-align:center">Age</th>
    <th width="25%" style="text-align:center">\(1960\)</th>
    <th width="25%" style="text-align:center">\(1975\)</th>
    <th width="25%" style="text-align:center">Projection \(1990\)</th>
  </tr>
  <tr>
    <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">V\((1,t)={@ageset_1a@}\)</td>
    <td style="text-align:center">V\((1,t+15)={@ageset_1b@}\)</td>
    <td style="text-align:center">V(1,t)→P1</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">V\((2,t)={@ageset_2a@}\)</td>
    <td style="text-align:center">V\((2,t+15)={@ageset_2b@}\)</td>
    <td style="text-align:center">V(2,t)→P2</td>
  </tr>
  <tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">V\((3,t)={@ageset_3a@}\)</td>
    <td style="text-align:center">V(3,t+15)={@ageset_3b@}</td>
    <td style="text-align:center">V(3,t)→P3</td>
  </tr>
  <tr>
    <td style="text-align:center">\(45-59\)</td>
    <td style="text-align:center">V\((4,t)={@ageset_4a@}\)</td>
    <td style="text-align:center">V\((4,t+15)={@ageset_4b@}\)</td>
    <td style="text-align:center">V(4,t)→P4</td>
  </tr>
  <tr>
    <td style="text-align:center">60+</td>
    <td style="text-align:center">V(5,t)={@ageset_5a@}</td>
    <td style="text-align:center">V(5,t+15)={@ageset_5b@}</td>
    <td style="text-align:center">V(5,t)→P5</td>
  </tr>
</table>


</p>
<p>To find the growth rates for \(\text{p}_1-\text{p}_5\) and population projection of \(1990\) we proceed as follows</p>
<p>\(\text{p}_1\displaystyle \frac{V(2,t+15)}{V(1,t)}\) where \(15\) is the the years for which you have data.</p>
<p>\(=\displaystyle \frac{{@ageset_2b@}}{{@ageset_1a@}}={@p_1_rate@}\)</p>
<p>To project the population from \(1975\) to \(1990\) we use the formula below</p>
<p>\(\text{P}_{1} 1990​=\text{Population in} 1975×(1+\text{Growth Rate})\)</p>
<p>\(\quad \quad={@ageset_1b@} \times (1+{@p_1_rate@})={@p_1@}\)</p>
<p>Applying the same technique to the other projections we have;</p>
<br>
<p>\(\text{p}_2=\displaystyle \frac{V(3,t+15)}{V(2,t)}\) </p>
<p>\(\quad \quad=\displaystyle \frac{{@ageset_3b@}}{{@ageset_2a@}}={@p_2_rate@}\)</p>
<p>\(\quad \quad={@ageset_2b@} \times (1+{@p_2_rate@})={@p_2@}\)</p>
<br>
<p>\(\text{p}_3=\displaystyle \frac{V(4,t+15)}{V(3,t)}\) </p>
<p>\(\quad \quad=\displaystyle \frac{{@ageset_4b@}}{{@ageset_3a@}}={@p_3_rate@}\)</p>
<p>\(\quad \quad={@ageset_3b@} \times (1+{@p_3_rate@})={@p_3@}\)</p>
<br>
<p>\(\text{p}_4=\displaystyle \frac{V(5,t+15)}{V(4,t)}\) </p>
<p>\(\quad \quad=\displaystyle \frac{{@ageset_5b@}}{{@ageset_4a@}}={@p_4_rate@}\)</p>
<p>\(\quad \quad={@ageset_4b@} \times (1+{@p_4_rate@})={@p_4@}\)</p>
<br>]]></text>
    </generalfeedback>
    <defaultgrade>0</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*sta 246*/
aa:makelist((rand_with_step(60000,65000,500)), i, 1, 2);/*population of 0-14 1960 and 1975*/
ageset_1a:aa[1];
ageset_1b:aa[2];
bb:makelist((rand_with_step(50000,60000,500)), i, 1, 2);/*population of 15-29 1960 and 1975*/
ageset_2a:bb[1];
ageset_2b:bb[2];
cc:makelist((rand_with_step(45000,49000,500)), i, 1, 2);/*population of 30-34 1960 and 1975*/
ageset_3a:cc[1];
ageset_3b:cc[2];
dd:makelist((rand_with_step(40000,44000,500)), i, 1, 2);/*population of 45-59 1960 and 1975*/
ageset_4a:dd[1];
ageset_4b:dd[2];
ee:makelist((rand_with_step(35000,40000,500)), i, 1, 2);/*population of 60+ 1960 and 1975*/
ageset_5a:ee[1];
ageset_5b:ee[2];

/*worked solution variables*/
p_1_rate:decimalplaces(ageset_2b/ageset_1a,3);
p_2_rate:decimalplaces(ageset_3b/ageset_2a,3);
p_3_rate:decimalplaces(ageset_4b/ageset_3a,3);
p_4_rate:decimalplaces(ageset_5b/ageset_4a,3);
p_5_rate:decimalplaces(ageset_5b/ageset_4a,3);

p_1:ceiling((1+p_1_rate)*ageset_1a);
p_2:ceiling((1+p_2_rate)*ageset_2a);
p_3:ceiling((1+p_3_rate)*ageset_3a);
p_4:ceiling((1+p_4_rate)*ageset_4a);
p_5:ceiling((1+p_5_rate)*ageset_5a);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Growth rates</p>
<p>\(p_1={@p_1_rate@}\) </p>
<p>\(p_2={@p_2_rate@}\)</p>
<p>\(p_3={@p_3_rate@}\) </p>
<p>\(p_4={@p_4_rate@}\) </p>
<p>\(p_5={@p_5_rate@}\) </p>
<p>Population in \(1990\)</p>
<p>\(0-14 ={@p_1@}\)</p>
<p>\(15-29 ={@p_2@}\)</p>
<p>\(30-44 ={@p_3@}\)</p>
<p>\(45-59 ={@p_4@}\)</p>
<p>\(60+  ={@p_5@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">Give population of a given region for \(1960-1975\) , the question requires you to find the projection the year \(1990\).</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>p_1_rate</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>p_5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>p_2_rate</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>p_3_rate</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>p_4_rate</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>p_5_rate</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>p_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>p_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>p_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>p_4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 0-14</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>p_1_rate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>p_1_rate</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection 60+ 1990</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>p_5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 15-29</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>p_2_rate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>p_2_rate</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 34-44</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>p_3_rate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>p_3_rate</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 44-59</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>p_4_rate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>p_4_rate</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth in 60+</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>p_5_rate</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>p_5_rate</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection 0-14 1990</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>p_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection 15-29 1990</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>p_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection 30-44 1990</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>p_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection 45-59 1990</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>p_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1856660881</deployedseed>
    <deployedseed>606775363</deployedseed>
    <deployedseed>112960354</deployedseed>
    <deployedseed>738508520</deployedseed>
    <deployedseed>2019250361</deployedseed>
    <deployedseed>2103840000</deployedseed>
    <deployedseed>1791388928</deployedseed>
    <deployedseed>1500579429</deployedseed>
    <deployedseed>224654143</deployedseed>
    <deployedseed>774774320</deployedseed>
    <deployedseed>1320164101</deployedseed>
    <deployedseed>134581158</deployedseed>
    <deployedseed>2107191495</deployedseed>
    <deployedseed>336652481</deployedseed>
    <deployedseed>480657810</deployedseed>
    <deployedseed>1054203176</deployedseed>
    <deployedseed>261459561</deployedseed>
    <deployedseed>1636526578</deployedseed>
    <deployedseed>150037449</deployedseed>
    <deployedseed>436222370</deployedseed>
    <deployedseed>1556106091</deployedseed>
    <deployedseed>385880763</deployedseed>
    <deployedseed>1440296353</deployedseed>
    <deployedseed>337359604</deployedseed>
    <deployedseed>117805701</deployedseed>
    <deployedseed>1895410606</deployedseed>
    <deployedseed>308007022</deployedseed>
    <deployedseed>1731864150</deployedseed>
    <deployedseed>2415687</deployedseed>
    <deployedseed>1092634381</deployedseed>
    <deployedseed>600935970</deployedseed>
    <deployedseed>940160491</deployedseed>
    <deployedseed>232204159</deployedseed>
    <deployedseed>1116970248</deployedseed>
    <deployedseed>1991117319</deployedseed>
    <deployedseed>60024228</deployedseed>
    <deployedseed>1573791686</deployedseed>
    <deployedseed>1230448000</deployedseed>
    <deployedseed>1740456178</deployedseed>
    <deployedseed>1169309290</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p_1_rate</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>p_5</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>p_2_rate</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>p_3_rate</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>p_4_rate</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>p_5_rate</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>p_1</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>p_2</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>p_3</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>p_4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
