<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>POPULATION PROJECTION_q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the following female population of Trinidad in \(1960\) and \(1975\), calculate the projection of the female population in \(1990\).</p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age</th>
    <th width="15%", style="text-align:center">\(1960\)</th>
    <th width="15%", style="text-align:center">\(1975\)</th>
  </tr>
  <tr>
  <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\({@ageset_1a@}\)</td>
    <td style="text-align:center">\({@ageset_1b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\({@ageset_2a@}\)</td>
    <td style="text-align:center">\({@ageset_2b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\({@ageset_3a@}\)</td>
    <td style="text-align:center">\({@ageset_3b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(45-59\)</td>
    <td style="text-align:center">\({@ageset_4a@}\)</td>
    <td style="text-align:center">\({@ageset_4b@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">60+</td>
    <td style="text-align:center">\({@ageset_5a@}\)</td>
    <td style="text-align:center">\({@ageset_5b@}\)</td>
  </tr>
</table>
</p>
<p>During the period \(1960-1975\), the female live births were distributed as follows:</p>
<p>\(15-29……….{@births_1@}\)</p>
<p>\(30-44……….{@births_2@}\)</p>
<!--<p>You may assume that there is mortality and fertility are constant, there is no net migration, and ages are distributed uniformly within each age group.</p>-->
<hr>
<p>Note: Leave your answer correct to \(3\) dp</p>
<p>Calculate the survivorship ratios; you may assume that \(P_4=P_5\).</p>
<p>\(\text{P}_1=\) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>
<p>\(\text{P}_2=\) [[input:ans5]] [[validation:ans5]][[feedback:prt5]]</p>
<p>\(\text{P}_3=\) [[input:ans6]] [[validation:ans6]][[feedback:prt6]]</p>
<p>\(\text{P}_4=\) [[input:ans7]] [[validation:ans7]][[feedback:prt7]]</p>
<p>\(\text{P}_5=\) [[input:ans8]] [[validation:ans8]][[feedback:prt8]]</p>
<br>
<p>Find the Probability of a baby born surviving to time \(t+15\) which is denoted as \(P_{b} =\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p>Find the fertility rate \(f_j^*\) for ages \(15-29, 30-44\)</p>  
<p>\(f_2^*=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\(f_3^*=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>Using the information above, project the population for \(1990\).</p>
<p>Note: Round your answer to the nearest integer</p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age</th>
    <th width="15%", style="text-align:center">\(1975\)</th>
<th width="15%", style="text-align:center">\(1990\)</th>
  </tr>
  <tr>
  <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\({@ageset_1b@}\)</td>
    <td style="text-align:center">[[input:ans9]] [[validation:ans9]][[feedback:prt9]]</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\({@ageset_2b@}\)</td>
    <td style="text-align:center">[[input:ans10]] [[validation:ans10]][[feedback:prt10]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\({@ageset_3b@}\)</td>
    <td style="text-align:center">[[input:ans11]] [[validation:ans11]][[feedback:prt11]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(44-59\)</td>
    <td style="text-align:center">\({@ageset_4b@}\)</td>
    <td style="text-align:center">[[input:ans12]] [[validation:ans12]][[feedback:prt12]]</td>
  </tr>
<tr>
    <td style="text-align:center">\(60+\)</td>
    <td style="text-align:center">\({@ageset_5b@}\)</td>
    <td style="text-align:center">[[input:ans13]] [[validation:ans13]][[feedback:prt13]]</td>
  </tr>
 </table>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Assumptions made are: Constant mortality, Constant fertility, No net migration, Uniform distribution by age within each age group of the female population</p>
<p>We have age intervals of \(15\) years. Thus, let</p>
<p>\(V(1,t)=\) population aged \(0-14\) years last birth day at time \(t\)</p>
<p>\(V(2,t)=\) population aged \(15-29\) years last birth day at time \(t\)</p>
<p>\(V(3,t)=\) population aged \(30-44\) years last birth day at time \(t\)</p>
<p>\(V(4,t)=\) population aged \(45-59\) years last birth day at time \(t\)</p>
<p>\(V(5,t)=\) population aged \(60+\) years last birth day at time \(t\)</p>
<p>\(P_i=\) the probability of surviving from age group \(i\) to the next age group for \(i=1\dots 5\)</p>
<p>\(P_b=\) probability of a baby born surviving to time \(t+15\)</p>
<p>\(f_i^*\) the fertility rate at age group \(i\)</p>
<hr>
<p>First, we arrange a schematic diagram for the population distribution. For \(t=1960\), we have:

<table border="solid" style="width: 100%;">
  <tr>
    <th width="25%" style="text-align:center">Age</th>
    <th width="25%" style="text-align:center">\(1960\)</th>
    <th width="25%" style="text-align:center">\(1975\)</th>
    <th width="25%" style="text-align:center">Projection \(1990\)</th>
  </tr>
  <tr>
    <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\(V(1,t)={@ageset_1a@}\)</td>
    <td style="text-align:center"> \(V(1,t+15)={@ageset_1b@}\) </td>
    <td style="text-align:center">\(V(1,t+30)\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\(V(2,t)={@ageset_2a@}\)</td>
    <td style="text-align:center">\(V(2,t+15)={@ageset_2b@}\)</td>
    <td style="text-align:center">\(V(2,t+30)\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\(V(3,t)={@ageset_3a@}\)</td>
    <td style="text-align:center">\(V(3,t+15)={@ageset_3b@}\)</td>
    <td style="text-align:center">\(V(3,t+30)\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(45-59\)</td>
    <td style="text-align:center">\(V(4,t)={@ageset_4a@}\)</td>
    <td style="text-align:center">\(V(4,t+15)={@ageset_4b@}\)</td>
    <td style="text-align:center">\(V(4,t+30)\)</td>
  </tr>
  <tr>
    <td style="text-align:center">60+</td>
    <td style="text-align:center">\(V(5,t)={@ageset_5a@}\)</td>
    <td style="text-align:center">\(V(5,t+15)={@ageset_5b@}\)</td>
    <td style="text-align:center">\(V(5,t+30)\)</td>
  </tr>
</table>
</p>
<hr>
<p>We start by calculating the survivorship ratios.</p>
<p>After \(15\) years, the surviving members of age group \(V(i,t)\) move to age group \(V(i+1,t+15)\) for \(i=1,2,3\). In mathematical terms: \(V(i+1,t+15) = V(i,t) P_i\). We can use this to compute the survivorship ratios:</p>
<p>\(\text{P}_1=\displaystyle \frac{V(2,t+15)}{V(1,t)}=\displaystyle \frac{{@ageset_2b@}}{{@ageset_1a@}}={@p1@}\) </p>
<p>\(\text{P}_2=\displaystyle \frac{V(3,t+15)}{V(2,t)}=\displaystyle \frac{{@ageset_3b@}}{{@ageset_2a@}}={@p2@}\) </p>
<p>\(\text{P}_3=\displaystyle \frac{V(4,t+15)}{V(3,t)}=\displaystyle \frac{{@ageset_4b@}}{{@ageset_3a@}}={@p3@}\) </p>
<p>For \(i=4\), i.e. the last age group, we have \(V(5,t+15) = V(4,t) P_4 + V(5,t) P_5\). We assume that \(P_5=P_4\), to get:</p>
<p>\(P_4 = P_5 = \displaystyle \frac{V(5,t+15)}{V(4,t) + V(5,t)} \displaystyle \frac{{@ageset_5b@}}{{@ageset_4a@}+{@ageset_5a@}}={@p4@}\) </p>
<hr>
<p>Average number of women in the fertility group \(15-29\) and \(30-44\):</p>
<p>\(w_2=  \displaystyle \frac{V(2,t)+V(2,t+15)}{2} = \displaystyle \frac{V(2,t)+V(1,t)P_1}{2}\)</p>
<p>\(w_3=  \displaystyle \frac{V(3,t)+V(3,t+15)}{2} =  \displaystyle\frac{V(3,t)+V(2,t)P_2}{2}\)</p>
<p>Computing the fertility rates:</p>
<p>\(f_2^*= \displaystyle \frac{ \displaystyle\frac{1}{15}\cdot \text{births from 15-29}}{w_3}=  \displaystyle \frac{ \displaystyle\frac{1}{15}\cdot{@births_1@}}{0.5({@ageset_2a@}+{@ageset_1a@}\cdot {@p1@})}={@f_2@}\)</p>
<p>and</p>
<p>\(f_3^*= \displaystyle \frac{\frac{1}{15}\cdot \text{births from 30-44}}{w_3}=  \displaystyle \frac{ \displaystyle\frac{1}{15}\cdot{@births_2@}}{0.5({@ageset_3a@}+{@ageset_2a@}\cdot {@p2@})}={@f_3@}\)</p>

<p>Within the \(15\) year period \(1960-1975\), we have \({@births_1@}+{@births_2@}\) live births. Of those, \(V(1,t+15)={@ageset_1b@}\) survive until \(1975\), thus we get \(P_b = \displaystyle \frac{{@ageset_1b@}}{{@births_1@}+{@births_2@}} = {@pb@}\)</p>

<p>No. of surviving births \(V(1,t+15)=w_2 \cdot 15 f_2^* \cdot P_b + w_3 \cdot 15 f_3^* \cdot P_b\)</p>
<p>\(\quad \quad =\displaystyle\frac{V(2,t)+V(1,t)P_1}{2} \cdot 15 \cdot f_2^* \cdot P_b + \displaystyle \frac{V(3,t)+V(2,t)P_2}{2} \cdot 15 \cdot f_3^* \cdot P_b\)</p>
<p>\(\quad \quad =\displaystyle V(2,t)\frac{15}{2}\cdot f_2^*\cdot P_b+V(2,t)P_2\cdot \displaystyle\frac{15}{2}\cdot f_3^*P_b+V(3,t)\frac{15}{2}\cdot f_3^*  \cdot P_b+ V(1,t)P_1 \cdot \displaystyle \frac{15}{2} \cdot f_2^* \cdot P_b \)</p>
<p>\(\quad \quad=\displaystyle V(1,t)P_1 \cdot \displaystyle \frac{15}{2}\cdot f_2^*\cdot P_b+V(2,t)(\displaystyle\frac{15}{2}\cdot P_b \cdot P_2 \cdot f_3^* +\frac{15}{2}\cdot P_b \cdot f_2^*)+V(3,t)\cdot \frac{15}{2} \cdot f_3^* \cdot P_b\)</p>
<p>Letting</p>
<p>\(f_1=\displaystyle P_1 \cdot \displaystyle\frac{15}{2} \cdot f_2^* \cdot P_b = {@f1@}\)</p>
<p>\(f_2=\displaystyle P_2 \cdot \displaystyle\frac{15}{2} \cdot f_3^* \cdot P_b+\displaystyle\frac{15}{2}\cdot f_2^* \cdot P_b = {@f2@}\)</p>
<p>\(f_3=\displaystyle \frac{15}{2} \cdot f_3^* \cdot P_b = {@f3@}\),</p>
<p>the above expression becomes \(V(1,t+15)=V(1,t)f_1+V(2,t)f_2+V(3,t)f_3\).</p>
<p>We get the following Leslie matrix:</p>
<p>\(\begin{bmatrix}
V(1,t+15)\\ 
V(2,t+15)\\ 
V(3,t+15)\\ 
V(4,t+15)\\ 
V(5,t+15)
\end{bmatrix}=\begin{bmatrix}
f_1 & f_2 & f_3 & 0 & 0\\ 
P_1 &0  & 0 & 0 & 0\\ 
 0&P_2  &0  &0  &0 \\ 
 0&0  & P_3 &0  &0 \\ 
 0&  0&  0& P_4 & P_5
\end{bmatrix}~\begin{bmatrix}
V(1,t)\\ 
V(2,t)\\ 
V(3,t)\\ 
V(4,t)\\ 
V(5,t)
\end{bmatrix}\)</p>
<hr>
<p>By evaluating this computation for \(t=1975\), </p>
<p>\(\begin{bmatrix}
V(1,1990)\\ 
V(2,1990)\\ 
V(3,1990)\\ 
V(4,1990)\\ 
V(5,1990)
\end{bmatrix}=\begin{bmatrix}
{@f1@} & {@f2@} & {@f3@} & 0 & 0\\ 
{@p1@} &0  & 0 & 0 & 0\\ 
 0& {@p2@}  &0  &0  &0 \\ 
 0&0  & {@p3@} &0  &0 \\ 
 0&  0&  0& {@p4@} & {@p5@}
\end{bmatrix}~\begin{bmatrix}
{@ageset_1b@}\\ 
{@ageset_2b@}\\ 
{@ageset_3b@}\\ 
{@ageset_4b@}\\ 
{@ageset_5b@}
\end{bmatrix}\),</p>
<p>we get the following projection:</p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age</th>
<th width="15%", style="text-align:center">\(1990\)</th>
  </tr>
  <tr>
  <td style="text-align:center">\(0-14\)</td>
    <td style="text-align:center">\({@proj_1@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">\(15-29\)</td>
    <td style="text-align:center">\({@proj_2@}\)</td>
  </tr>
<tr>
    <td style="text-align:center">\(30-44\)</td>
    <td style="text-align:center">\({@proj_3@}\)</td>
  </tr>
<tr>
    <td style="text-align:center">\(44-59\)</td>
   <td style="text-align:center">\({@proj_4@}\)</td>
  </tr>
<tr>
    <td style="text-align:center">\(60+\)</td>
   <td style="text-align:center">\({@proj_5@}\)</td>
  </tr>
 </table>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*sta 246*/
aa:makelist((rand_with_step(60000,65000,500)), i, 1, 2);/*population of 0-14 1960 and 1975*/
ageset_1a:aa[1];
ageset_1b:aa[2];
bb:makelist((rand_with_step(50000,60000,500)), i, 1, 2);/*population of 15-29 1960 and 1975*/
ageset_2a:bb[1];
ageset_2b:bb[2];
cc:makelist((rand_with_step(45000,49000,500)), i, 1, 2);/*population of 30-34 1960 and 1975*/
ageset_3a:cc[1];
ageset_3b:cc[2];
dd:makelist((rand_with_step(40000,44000,500)), i, 1, 2);/*population of 45-59 1960 and 1975*/
ageset_4a:dd[1];
ageset_4b:dd[2];
ee:makelist((rand_with_step(35000,40000,500)), i, 1, 2);/*population of 60+ 1960 and 1975*/
ageset_5a:ee[1];
ageset_5b:ee[2];

/* Have to make sure that the number of births exceeds ageset_1b */
births_1:rand_with_step(round(0.5*ageset_1b),round(0.6*ageset_1b),1);
births_2:rand_with_step(round(0.5*ageset_1b),round(0.6*ageset_1b),1);

/*worked solution variables*/
p1:decimalplaces(ageset_2b/ageset_1a,3);
p2:decimalplaces(ageset_3b/ageset_2a,3);
p3:decimalplaces(ageset_4b/ageset_3a,3);
p4:decimalplaces(ageset_5b/(ageset_4a+ageset_5a),3);
p5:decimalplaces(ageset_5b/(ageset_4a+ageset_5a),3);

f_2:decimalplaces(births_1/((ageset_2a+ageset_1a*p1)*15/2),3);/*fertility rate 2*/
f_3:decimalplaces(births_2/((ageset_3a+ageset_2a*p2)*15/2),3);/*fertility rate 3*/
pb:decimalplaces(ageset_1b/(births_1+births_2),3);
f1:decimalplaces(p1*(15/2)*f_2*pb,3);
f2:decimalplaces(p2*(15/2)*f_3*pb+(7.5*f_2*pb),3);
f3:decimalplaces((15/2)*f_3*pb,3);


/*projection*/
proj_1:round((f1*ageset_1b)+(f2*ageset_2b)+(f3*ageset_3b));
proj_2:round(ageset_1b*p1);
proj_3:round(ageset_2b*p2);
proj_4:round(ageset_3b*p3);
proj_5:round((p4*ageset_4b)+(p5*ageset_5b));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(P_1={@p1@}\) </p>
<p>\(P_2={@p2@}\) </p>
<p>\(p_3={@p3@}\) </p>
<p>\(p_4=P_5={@p4@}\) </p>
<p>Population in \(1990\)</p>
<p>\(0-14 ={@proj_1@}\)</p>
<p>\(15-29 ={@proj_2@}\)</p>
<p>\(30-44 ={@proj_3@}\)</p>
<p>\(45-59 ={@proj_4@}\)</p>
<p>\(60+  ={@proj_5@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">Give population of a given region for \(1960-1975\) , the question requires you to find the projection the year \(1990\).</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>f_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>proj_2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>proj_3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>proj_4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>algebraic</type>
      <tans>proj_5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>f_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>pb</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>p1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>p2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>p3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>p4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>p5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>proj_1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 0-14</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>f_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>f_2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection for 15-29</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>proj_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection for 30-44</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>proj_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection for 45-59</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>proj_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt13</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection for 60+</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>proj_5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt13-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt13-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for growth rate in 15-29</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>f_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>f_3</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for pb</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>pb</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>pb</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for P1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>p1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>p1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for P2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>p2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>p2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for P3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>p3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans6</sans>
        <tans>p3</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for P4</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>p4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans7</sans>
        <tans>p4</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for P5</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>p5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans8</sans>
        <tans>p5</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct projection for 0-14</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>proj_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1856660881</deployedseed>
    <deployedseed>606775363</deployedseed>
    <deployedseed>112960354</deployedseed>
    <deployedseed>738508520</deployedseed>
    <deployedseed>2019250361</deployedseed>
    <deployedseed>2103840000</deployedseed>
    <deployedseed>1791388928</deployedseed>
    <deployedseed>1500579429</deployedseed>
    <deployedseed>224654143</deployedseed>
    <deployedseed>774774320</deployedseed>
    <deployedseed>1320164101</deployedseed>
    <deployedseed>134581158</deployedseed>
    <deployedseed>2107191495</deployedseed>
    <deployedseed>336652481</deployedseed>
    <deployedseed>480657810</deployedseed>
    <deployedseed>1054203176</deployedseed>
    <deployedseed>261459561</deployedseed>
    <deployedseed>1636526578</deployedseed>
    <deployedseed>150037449</deployedseed>
    <deployedseed>436222370</deployedseed>
    <deployedseed>1556106091</deployedseed>
    <deployedseed>385880763</deployedseed>
    <deployedseed>1440296353</deployedseed>
    <deployedseed>337359604</deployedseed>
    <deployedseed>117805701</deployedseed>
    <deployedseed>1895410606</deployedseed>
    <deployedseed>308007022</deployedseed>
    <deployedseed>1731864150</deployedseed>
    <deployedseed>2415687</deployedseed>
    <deployedseed>1092634381</deployedseed>
    <deployedseed>600935970</deployedseed>
    <deployedseed>940160491</deployedseed>
    <deployedseed>232204159</deployedseed>
    <deployedseed>1116970248</deployedseed>
    <deployedseed>1991117319</deployedseed>
    <deployedseed>60024228</deployedseed>
    <deployedseed>1573791686</deployedseed>
    <deployedseed>1230448000</deployedseed>
    <deployedseed>1740456178</deployedseed>
    <deployedseed>1169309290</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>f_2</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>proj_2</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>proj_3</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>proj_4</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>proj_5</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>f_3</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>pb</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>p1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>p2</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>p3</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>p4</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>p5</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>proj_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt13</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt13-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
