<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Probability of Drawing a Red Ball from Bag B Given One Red Ball Observed: A Bayesian Approach</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A bag \(A\) contains \({@w_a@}\) white balls and \({@r_a@}\) red balls and a bag \(B\) contains \({@w_b@}\) white balls and \({@r_b@}\)
red balls. One of the bags is chosen at random, one ball is drawn randomly from it, and this ball is found to be red. Given this observation, find
the probability that the ball was drawn from bag \(B\).</p>
<p>The probability that it was drawn from bag \(B\) is [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p><b>Note:</b> Give your answer correct to four decimal places or as a fraction.</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Let \(E_1\) be the event that the ball is drawn from bag \(A\) and \(E_2\) be the event that the ball is drawn from bag \(B\).</p>
<p>Then, a priori, \(P(E_1)=P(E_2)=\frac{1}{2}\)</p>
<p>Let \(E\) be the event that the ball drawn is red.</p>
<p>The probability of drawing a red ball from bag \(A\) is \(P(E|E_1)={@r_a/(r_a+w_a)@}\)</p>
<p>The probability of drawing a red ball from bag \(B\) is \(P(E|E_2)={@r_b/(r_b+w_b)@}\)</p>
<p>Thus the probability that a ball was drawn from bag \(B\), given that it was red, is \( \displaystyle P \bigg(E_2|E \bigg )=\frac{(\frac{1}{2})({@r_b/(r_b+w_b)@})}{(\frac{1}{2})({@r_a/(r_a+w_a)@})+(\frac{1}{2})({@r_b/(r_b+w_b)@})}={@ta@}\)</p>
<p>Therefore the answer is \({@ta@}\) or \({@decimalplaces(ta,4)@}\) (\(4\) significant figures or \(4\) decimal places). </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>w_a:rand(5)+2;
r_a:rand_with_prohib(2,6,[w_a]);
w_b:rand_with_prohib(2,7,[w_a,r_a]);
r_b:rand_with_prohib(2,7,[w_a,r_a,w_b]);
ta:((1/2)*(r_b/(w_b+r_b)))/((1/2)*(r_a/(w_a+r_a))+(1/2)*(r_b/(w_b+r_b)));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(A-{@w_a@}\) white balls and \({@r_a@}\) red balls.</p>
 <p>\(B-{@w_b@}\) white balls and \({@r_b@}\)
red balls.</p> 
<p>The probability that ball is drawn and found to be red was from bag B is \({@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/1qelTr6Tnajf0Z2NHlqlSHEbPpp5ZRyWW/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question requires the student to calculate the probability of drawing a red ball from Bag B, given the information that one red ball was randomly drawn from either Bag A  or Bag B.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct fraction answer*/
/*Node 2: Checks for the correct decimal answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct fraction answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correctly rounded decimal answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check for correct decimal answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.00005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1782169399</deployedseed>
    <deployedseed>1170579815</deployedseed>
    <deployedseed>1310419784</deployedseed>
    <deployedseed>34016904</deployedseed>
    <deployedseed>88004317</deployedseed>
    <deployedseed>1165109447</deployedseed>
    <deployedseed>2145020487</deployedseed>
    <deployedseed>329194199</deployedseed>
    <deployedseed>1366552409</deployedseed>
    <deployedseed>738262366</deployedseed>
    <deployedseed>1021554995</deployedseed>
    <deployedseed>1796313449</deployedseed>
    <deployedseed>1863380814</deployedseed>
    <deployedseed>1418770657</deployedseed>
    <deployedseed>1015623862</deployedseed>
    <deployedseed>2130279216</deployedseed>
    <deployedseed>1082356239</deployedseed>
    <deployedseed>249869544</deployedseed>
    <deployedseed>1041355562</deployedseed>
    <deployedseed>1168508960</deployedseed>
    <deployedseed>653053794</deployedseed>
    <deployedseed>32858290</deployedseed>
    <deployedseed>1328110739</deployedseed>
    <deployedseed>1974793900</deployedseed>
    <deployedseed>151126514</deployedseed>
    <deployedseed>538948713</deployedseed>
    <deployedseed>776230112</deployedseed>
    <deployedseed>2014401446</deployedseed>
    <deployedseed>9273648</deployedseed>
    <deployedseed>1186799782</deployedseed>
    <deployedseed>137643158</deployedseed>
    <deployedseed>750393816</deployedseed>
    <deployedseed>784953681</deployedseed>
    <deployedseed>525636078</deployedseed>
    <deployedseed>33691054</deployedseed>
    <deployedseed>1675021215</deployedseed>
    <deployedseed>230404792</deployedseed>
    <deployedseed>866657192</deployedseed>
    <deployedseed>1607363230</deployedseed>
    <deployedseed>788026654</deployedseed>
    <deployedseed>672813116</deployedseed>
    <deployedseed>143614707</deployedseed>
    <deployedseed>1049234194</deployedseed>
    <deployedseed>1052521609</deployedseed>
    <deployedseed>192262070</deployedseed>
    <deployedseed>286956103</deployedseed>
    <deployedseed>1470048298</deployedseed>
    <deployedseed>29200303</deployedseed>
    <deployedseed>1522709410</deployedseed>
    <deployedseed>982711822</deployedseed>
    <deployedseed>250164110</deployedseed>
    <deployedseed>1705769311</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
