<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Q1_Concerns of a demographer</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which among the following is a concern for a demographer?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <i>The changing of structure of a population</i>.</p>
<p>The changing structure of a population refers to shifts and variations in the composition of a population over time. This can involve changes in the distribution of age groups, gender, ethnicity, socioeconomic status, or other demographic characteristics. </p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The correct answer is <i>The current size of a population</i>.</p>
<p>Population size is a key demographic indicator, and understanding its magnitude provides valuable insights into various aspects of a region's dynamics, such as resource allocation, age distribution, economic development among others:
</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>The correct answer is <i>Population composition or structure</i>.</p>
<p>Population composition or structure is the arrangement or distribution of individuals within a population based on various demographic characteristics. These characteristics include age, sex, education, occupation, ethnicity, marital status, and other relevant factors. Analyzing population composition helps demographers and researchers understand the diversity and dynamics within a population.
</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "The changing of structure of a population"],
[2, true, "The current size of a population"],
[3, true, "Population composition or structure"]],1);
incorrect_answers:rand_selection([
[4, false, "Carrying out population census"],
[5, false, "Producing Demographic health data"],
[6, false, "Infant deaths"],
[7, false, "Doing survey"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Which among the following is  a concern for a demographer</p>
<p>{@ta@}</p>
[[ if test="member(1, ta_index)" ]]
<p>The changing of structure of a population.</p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The current size of a population.</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Population composition or structure.</p>
[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1yKMASRNBrxv56EctlLkgvdHRMASolQhCBNM67lq8h9M/edit</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer"></p>
<p id="description">The question focuses on outlining the major concerns of a demographer</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2110384822</deployedseed>
    <deployedseed>617513444</deployedseed>
    <deployedseed>1473694765</deployedseed>
    <deployedseed>988625460</deployedseed>
    <deployedseed>736983476</deployedseed>
    <deployedseed>1627238698</deployedseed>
    <deployedseed>1029307470</deployedseed>
    <deployedseed>574819684</deployedseed>
    <deployedseed>1679080041</deployedseed>
    <deployedseed>1168730881</deployedseed>
    <deployedseed>884872604</deployedseed>
    <deployedseed>762654898</deployedseed>
    <deployedseed>235727052</deployedseed>
    <deployedseed>1587891385</deployedseed>
    <deployedseed>2122831240</deployedseed>
    <deployedseed>1605625566</deployedseed>
    <deployedseed>1842066752</deployedseed>
    <deployedseed>1150787794</deployedseed>
    <deployedseed>233314499</deployedseed>
    <deployedseed>1375188841</deployedseed>
    <deployedseed>1576679748</deployedseed>
    <deployedseed>744864973</deployedseed>
    <deployedseed>721602713</deployedseed>
    <deployedseed>831214936</deployedseed>
    <deployedseed>837731955</deployedseed>
    <deployedseed>347740008</deployedseed>
    <deployedseed>981702248</deployedseed>
    <deployedseed>1199138847</deployedseed>
    <deployedseed>12160704</deployedseed>
    <deployedseed>527836176</deployedseed>
    <deployedseed>1704004176</deployedseed>
    <deployedseed>663763556</deployedseed>
    <deployedseed>2024934615</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
