<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Q4_Identifying methods that can be used in vital statistics (copy)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Indentify a method that can be used to obtain Vital statistics</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <i>Population registers </i> since they are designed to provide some indication of the possibilities of  obtaining statistical information on individuals, households, families, and other units of analysis for potential use in demographic, biologic, medical, sociological and genetic studies.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The correct answer is <i>Population census</i> since it is a source for population estimates needed to calculate vital rates based on data derived from civil registration. It is also important in providing the base population for the estimates of statistics obtained from demographic surveys..</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p><i>Sample survey</i> can be used to obtain vital statistics as inferences about the entire population can be obtained from a subset, or sample, of the population members.
</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "Population registers"],
[2, true, "Population census"],
[3, true, "Sample survey"]],1);


incorrect_answers:rand_selection([
[4, false, "Fecudity"],
[5, false, "Migration"],
[6, false, "Questionnaires"],
[7, false, "Calendars"]],3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Identify a method that can be used to obtain Vital statistics</p>
[[ if test="member(1, ta_index)" ]]
<p>Population registers.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>Population census </p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Sample survey </p>
[[/ if ]]
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/17AQFfmO_SNw-3CJsWxBBDUa7GJyckWab5WtQDjoR71w/edit</p>
<p id="reviewer">Mary sayuni</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question requires the student to identify the methods that can be used to obtain  vital statistics</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct option*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct option</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>77</deployedseed>
    <deployedseed>883775737</deployedseed>
    <deployedseed>1678809143</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
