<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Set Operations Challenge</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the universal set as \(U={@u_set@}\) and \(A={@set_a@}\,, B={@set_b@}\) and  \(C={@set_c@}\)</p>
<p>Find \((B \cup C)'= \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Given the universal set as \(U={@u_set@}\) and \(A={@set_a@}\,, B={@set_b@}\)  and \(C={@set_c@}\)</p>
<p> \((B \cup C)'\) represents the complement of the union of sets \(B\) and \(C\). In other words, it is the set of elements that do not belong to the union of sets \(B\) and \(C\).</p>
<p>The  \((B \cup C)\) is union of all the elements that are in set \(B\) and \(C\) which is \({@b_u_c@}\)</p>
<p>Elements in the universal set \(U\) that are not in the union \((B \cup C)\) form  \((B \cup C)'\) </p>
<p>\((B \cup C)'={@ta1@}\)</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*sets*/
a1: rand(2)-5;
a2: a1+2;
a3:a2+1;
a4:0;
a5:rand(4)+1;
a6:a5+1;
a7:a6+1;
a8:a7+1;
a9:a8+1;
a10:a9+1;

u_set:{a1,a2,a3,a4,a5,a6,a7,a8,a9,a10};
set_a:{a1,a2,a3,a5};
set_b:{a2,a4,a6,a7};
set_c:{a2,a3,a4,a8};
/*model answer*/
b_u_c:{a2,a3,a4,a6,a7,a8};
b_u_c_complement:{a1,a5,a9,a10};
ta1:b_u_c_complement;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(U={@u_set@}\) and \(A={@set_a@}\,, B={@set_b@} \,,C={@set_c@}\)</p>
<p>Find \((B \cup C)'={@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1T9mbNupNv0Cs_PjnlrMwrN44HSDBD5KqV-slJLcnpqc/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer"></p>
<p id="description">The question focuses on identifying \((B \cup C)'\) given the universal set \(U\) and sets \(A,B\) and \(C\)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer on \((B \cup C)'\)</description>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1472787931</deployedseed>
    <deployedseed>469471871</deployedseed>
    <deployedseed>1629541959</deployedseed>
    <deployedseed>308844995</deployedseed>
    <deployedseed>2083693928</deployedseed>
    <deployedseed>2046463799</deployedseed>
    <deployedseed>472502338</deployedseed>
    <deployedseed>92654709</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
