<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Skewness (SK) and Kurtosis_Q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Suppose that a sample of \({@n@}\) batches of different whooping cough vaccines were tested on groups of children to assess their protective immunizing activity. The following data give immunity percentage in home exposure (IPHE) values: Fill the blanks in the table and answer the questions that follow.</p>
<p><b>Note: \(1\). Give your answer to \(2\) decimal places where necessary in all your calculations.</b></p>
<p>\( \qquad 2. X_i-\mu =d\)  </p>
<p>\( \qquad 3\) The central moment is given by the formula: \( \displaystyle M_n = \frac{1}{N} \sum_{i=1}^{N} (x_i - \bar{x})^n \cdot f_i\)</p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">IPHE values CI(\(\%\))</th>
    <th width="15%", style="text-align:center">IPHE boundaries CI(\(\%\))</th>
     <th width="15%", style="text-align:center">Midpoint of intervals \((X_i)\)</th>
     <th width="15%", style="text-align:center">Number of children \((f)\)</th>
    <th width="15%", style="text-align:center">Cumm frequency \((cf)\)</th>
     <th width="15%", style="text-align:center">\(fX_i\)</th>
     <th width="15%", style="text-align:center">\(f(X_i-\mu)^2\)</th>
     <th width="60%", style="text-align:center">\(fd^3\)</th>
     <th width="60%", style="text-align:center">\(fd^4\)</th>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range1@}-{@upper_range1@}\)</td>
    <td style="text-align:center">\({@float(lower_bound1)@}-{@float(upper_bound1)@}\)</td>
    <td style="text-align:center">\({@midpoints[1]@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf[1]@}\)</td>
    <td style="text-align:center"> [[input:ans1]][[validation:ans1]][[feedback:prt1]]</td>
    <td style="text-align:center">\({@mean_fdiffs[1]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[1]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[1]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range2@}-{@upper_range2@}\)</td>
    <td style="text-align:center">\({@float(lower_bound2)@}-{@float(upper_bound2)@}\)</td>
     <td style="text-align:center">[[input:ans2]][[validation:ans2]][[feedback:prt2]]</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">\({@cf[2]@}\)</td>
    <td style="text-align:center">\({@fxi[2]@}\)</td>
    <td style="text-align:center"> [[input:ans3]][[validation:ans3]][[feedback:prt3]]</td>
    <td style="text-align:center"> [[input:ans3b]][[validation:ans3b]][[feedback:prt3b]]</td>
    <td style="text-align:center"> [[input:ans4b]][[validation:ans4b]][[feedback:prt4b]]</td>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range3@}-{@upper_range3@}\)</td>
    <td style="text-align:center">\({@float(lower_bound3)@}-{@float(upper_bound3)@}\)</td>
    <td style="text-align:center">\({@midpoints[3]@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">[[input:ans4]][[validation:ans4]][[feedback:prt4]]</td>
    <td style="text-align:center">[[input:ans5]][[validation:ans5]][[feedback:prt5]]</td>
    <td style="text-align:center">[[input:ans6]][[validation:ans6]][[feedback:prt6]]</td>
    <td style="text-align:center"> {@mean_fdiffss3[3]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[3]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range4@}-{@upper_range4@}\)</td>
    <td style="text-align:center">\({@float(lower_bound4)@}-{@float(upper_bound4)@}\)</td>
    <td style="text-align:center">[[input:ans7]][[validation:ans7]][[feedback:prt7]]</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">\({@cf[4]@}\)</td>
    <td style="text-align:center">[[input:ans8]][[validation:ans8]][[feedback:prt8]]</td>
    <td style="text-align:center">\({@mean_fdiffs[4]@}\)</td>
    <td style="text-align:center">[[input:ans3d]][[validation:ans3d]][[feedback:prt3d]]</td>
    <td style="text-align:center"> {@mean_fdiffss4[4]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range5@}-{@upper_range5@}\)</td>
    <td style="text-align:center">\({@float(lower_bound5)@}-{@float(upper_bound5)@}\)</td>
    <td style="text-align:center">\({@midpoints[5]@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">[[input:ans9]][[validation:ans9]][[feedback:prt9]]</td>
    <td style="text-align:center">[[input:ans10]][[validation:ans10]][[feedback:prt10]]</td>
    <td style="text-align:center">\({@mean_fdiffs[5]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[5]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[5]@}</td>
  </tr>
    <tr>
   <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center">\( \sum f =\) [[input:ans11]][[validation:ans11]][[feedback:prt11]]</td>
    <td style="text-align:center"></td>
    <td style="text-align:center">\(\sum fX_i=\) [[input:ans12]][[validation:ans12]][[feedback:prt12]]</td>
    <td style="text-align:center">\(\sum{fd^2}=\) [[input:ans13]][[validation:ans13]][[feedback:prt13]]</td>
    <td style="text-align:center">\(\sum{fd^3}={@sum_m3@}\) </td>
    <td style="text-align:center">\(\sum{fd^4}={@sum_m4@}\) </td>
  </tr>
</table>
</p>
<p>Using the information in the table above, compute the following statistics:</p>
<br>\(\textbf{Part A}\)
<p>(a) Mean \( \displaystyle \bar{x}= \mu =\frac{\sum{fX}}{\sum{f}}=\) [[input:ans14]][[validation:ans14]][[feedback:prt14]]</p>

<p>(b) Median \(=\) [[input:ans15]][[validation:ans15]][[feedback:prt15]]</p>

<p>(c) Mode\(=\) [[input:ans16]][[validation:ans16]][[feedback:prt16]]</p>

<p>(d) Standard deviation \( \displaystyle  \sigma = \sqrt{\frac{\sum{f(x-\bar{x})^2}}{\sum f}}=\)  [[input:ans17]][[validation:ans17]][[feedback:prt17]]</p>

<p>(e) Variance \({ \sigma^2}=\) [[input:ans18]][[validation:ans18]][[feedback:prt18]]</p>
<br>\(\textbf{Part B}\)
<p>(a) Central \(2\)nd Moment (Variance): \(M_2=\) [[input:ansm2]][[validation:ansm2]][[feedback:prtm2]]</p>
<p>(b) Central \(3\)rd Moment: \(M_3=\) [[input:ansm3]][[validation:ansm3]][[feedback:prtm3]]</p>
<p>(c) Central \(4\)th Moment: \(M_4=\) [[input:ansm4]][[validation:ansm4]][[feedback:prtm4]]</p>
<p>(d) \( \displaystyle \text{Moment Coefficient of Skewness (Sk) } = \frac{M_3}{M_2^{\frac{3}{2}}}=\) [[input:anssk]][[validation:anssk]][[feedback:prtsk]]</p>
<p>(e)  \(\displaystyle  \text{Moment Coefficient of Kurtosis (Ku) } = \frac{M_4}{M_2^2} - 3=\) [[input:anskt]][[validation:anskt]][[feedback:prtkt]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b><h3>Worked Solution</h3></b></p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">IPHE values CI(\(\%\))</th>
    <th width="15%", style="text-align:center">IPHE boundaries CI(\(\%\))</th>
     <th width="25%", style="text-align:center">Midpoint of intervals \((X_i)\)</th>
     <th width="25%", style="text-align:center">Number of children \((f)\)</th>
    <th width="30%", style="text-align:center">Cumulative frequency \((cf)\)</th>
     <th width="15%", style="text-align:center">\(fX_i\)</th>
     <th width="15%", style="text-align:center">\(f(X_i-\mu)^2\)</th>
     <th width="60%", style="text-align:center">\(fd^3\)</th>
     <th width="60%", style="text-align:center">\(fd^4\)</th>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range1@}-{@upper_range1@}\)</td>
    <td style="text-align:center">\({@float(lower_bound1)@}-{@float(upper_bound1)@}\)</td>
    <td style="text-align:center">\({@midpoints[1]@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf[1]@}\)</td>
    <td style="text-align:center">\({@fxi[1]@}\)</td>
    <td style="text-align:center">\({@mean_fdiffs[1]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[1]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[1]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range2@}-{@upper_range2@}\)</td>
    <td style="text-align:center">\({@float(lower_bound2)@}-{@float(upper_bound2)@}\)</td>
     <td style="text-align:center">\({@midpoints[2]@}\)</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">\({@cf[2]@}\)</td>
    <td style="text-align:center">\({@fxi[2]@}\)</td>
    <td style="text-align:center">\({@mean_fdiffs[2]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[2]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[2]@}</td>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_range3@}-{@upper_range3@}\)</td>
    <td style="text-align:center">\({@float(lower_bound3)@}-{@float(upper_bound3)@}\)</td>
    <td style="text-align:center">\({@midpoints[3]@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">\({@cf[3]@}\)</td>
    <td style="text-align:center">\({@fxi[3]@}\)</td>
    <td style="text-align:center">\({@mean_fdiffs[3]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[3]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[3]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range4@}-{@upper_range4@}\)</td>
    <td style="text-align:center">\({@float(lower_bound4)@}-{@float(upper_bound4)@}\)</td>
    <td style="text-align:center">\({@midpoints[4]@}\)</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">\({@cf[4]@}\)</td>
    <td style="text-align:center">\({@fxi[4]@}\)</td>
    <td style="text-align:center">\({@mean_fdiffs[4]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[4]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[4]@}</td>
  </tr>
  <tr>
   <td style="text-align:center">\({@lower_range5@}-{@upper_range5@}\)</td>
    <td style="text-align:center">\({@float(lower_bound5)@}-{@float(upper_bound5)@}\)</td>
    <td style="text-align:center">\({@midpoints[5]@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">\({@cf[5]@}\)</td>
    <td style="text-align:center">\({@fxi[5]@}\)</td>
    <td style="text-align:center">\({@mean_fdiffs[5]@}\)</td>
    <td style="text-align:center"> {@mean_fdiffss3[5]@}</td>
    <td style="text-align:center"> {@mean_fdiffss4[5]@}</td>
  </tr>
    <tr>
   <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center">\(\sum f={@n@}\)</td>
    <td style="text-align:center"></td>
    <td style="text-align:center">\(\sum fX_i={@sum_fxi@}\)</td>
    <td style="text-align:center">\(\sum{fd^2}={@mean_squared_diffs@}\)</td>
    <td style="text-align:center">\(\sum{fd^3}={@sum_m3@}\) </td>
    <td style="text-align:center">\(\sum{fd^4}={@sum_m4@}\) </td>
  </tr>
</table>
</p>
<p>Using the information in the table above, we will compute the following statistics:</p>
<br>\(\textbf{Part A}\)
<p>(a) Mean \( \displaystyle \bar{x}=\frac{\sum{fX}}{\sum{f}}=\frac{{@sum_fxi@}}{{@n@}}={@mean@}\)</p>

<p>(b) Median</p>
<p>The median class is that class containing the \( \displaystyle  \left( \frac{{@total_f@}}{2}\right)\) item \(={@median_cf@}^{\text{th}}\)item.</p>
<p>Therefore, the median class is \({@float(lower_bound)@}-{@float(upper_bound)@}\).</p>
<p>Thus \(B_L={@float(lower_bound)@}, \, \displaystyle  \left( \frac{n}{2}\right)={@median_cf@}, \, f_c={@f_median@}, \, i={@float(upper_bound)@}-{@float(lower_bound)@}={@width@}, \, cf_b={@cf_lower@} \) </p>
<p>Therefore,</p><p> \( \begin{align} \displaystyle md &=B_L+  \left( \frac{\frac{N}{2}-cf_b}{f_c}\right) \cdot i \\
                                                                             &={@float(lower_bound)@}+  \left( \frac{\frac{{@total_f@}}{2}-{@cf_lower@}}{{@f_median@}}\right) \cdot  {@width@} \\
                                                           &={@float(lower_bound)@}+{@md1@} \\
                                                           &={@ta@} \end{align}\)</p>
<p>The median immunity percentage in home exposure is the value \({@ta@}\).</p>

<p>(c) Mode</p>
<p>When dealing with grouped data, the mode represents the value or values in a dataset that have the highest frequency. </p>
<p>To calculate the mode for our grouped data, we will use the following formula:</p>

<p>\(\displaystyle \text{Mode}=L+\left( \frac{f_1 - f_0}{(f_1 - f_0)+(f_1 - f_2)}\right )\cdot h\)</p>

<p>Where:</p>
<p>-\( L={@float(Lm)@}\): Lower boundary of the modal class (class with the highest frequency)</p>
<p>- \(f_1={@f1m@}\): Frequency of the modal class</p>
<p>- \(f_0={@f0m@}\): Frequency of the class preceding the modal class</p>
<p>- \(f_2={@f2m@}\): Frequency of the class succeeding the modal class</p>
<p>- \(h={@hm@}\): Class width (the difference between upper and lower boundaries of each class)</p>
<p>\(\displaystyle \text{Mode}={@float(Lm)@}+\left( \frac{{@f1m@}- {@f0m@}}{({@f1m@}- {@f0m@})+({@f1m@} - {@f2m@})}\right )\cdot {@hm@}={@Mode@}\)</p>
<p>This formula allows us to estimate the mode for grouped data by taking into account the relative frequencies of the modal class and its neighboring classes.</p>

<p>It's important to note that the mode for grouped data is an estimate and may not represent an exact value within the dataset.</p>

<p>(d) Standard deviation \( \displaystyle  \sigma = \sqrt{\frac{\sum{f(x-\bar{x})^2}}{\sum f}}= \sqrt{\frac{{@mean_squared_diffs@}}{{@n@}}}={@standard_deviation@}\) </p>

<p>(e) Variance \( \displaystyle { \sigma^2}=\frac{{@mean_squared_diffs@}}{{@n@}}={@variance@}\)</p>
<br>\(\textbf{Part B}\)
<p>To calculate the moments, skewness, and kurtosis, we first need to find the central moment of the distribution. The central moment is given by the formula:</p>
<p>
\(\displaystyle
M_n = \frac{1}{N} \sum_{i=1}^{N} (x_i - \bar{x})^n \cdot f_i
\)
</p>

<p>Where:</p>
<ul>
  <li>\(M_n\) is the central moment of order \(n\).</li>
  <li>\(N\) is the total number of data points (in this case, the sum of all frequencies).</li>
  <li>\(x_i\) is the midpoint of each class interval.</li>
  <li>\(\bar{x}\) is the assumed mean ({@mean@} in this case).</li>
  <li>\(n\) is the order of the moment.</li>
  <li>\(f_i\) is the frequency of the class interval.</li>
</ul>

<p>Let's calculate the moments, skewness, and moment coefficient of kurtosis for the given data:</p>

<ul>
  <li>
    <p>Calculate the central moments:</p>
    <p>
    \( \displaystyle 
\begin{align}
   (a)\,  M_2 &= \frac{1}{N} \sum_{i=1}^{5} (x_i - \bar{x})^2 \cdot f_i \\
             &=\frac{{@mean_squared_diffs@}}{{@n@}} \\
        &={@decimalplaces(M2,2)@} \\
\\ 
   (b)\,    M_3 &= \frac{1}{N} \sum_{i=1}^{5} (x_i - \bar{x})^3 \cdot f_i \\
             &=\frac{{@sum_m3@}}{{@n@}} \\
        &={@decimalplaces(M3,2)@} \\
\\ 
     (c)\,  M_4 &= \frac{1}{N} \sum_{i=1}^{5} (x_i - \bar{x})^4 \cdot f_i \\
              &=\frac{{@sum_m4@}}{{@n@}} \\
        &={@decimalplaces(M4,2)@}
    \end{align} \)
    </p>
  </li>
<p>Now, you can calculate \(Sk\) and \(Ku\) using the values of \(M_2\), \(M_3\), and \(M_4\).</p>
<p>(d) \( \displaystyle \text{Moment Coefficient of Skewness (Sk) } = \frac{M_3}{M_2^{\frac{3}{2}}}={@Sk@}\) </p>
<p>(e)  \(\displaystyle  \text{Moment Coefficient of Kurtosis (Ku) } = \frac{M_4}{M_2^2} - 3={@Ku@}\) </p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[class_widths:rand(10)+2;

/*randomly generating class intervals*/
aa:rand(20)+10;

lower_range1: aa+ class_widths;  /*Adding 5 on each variable to ensure the subsequent class width is the same*/
upper_range1: lower_range1 + class_widths;

lower_range2: upper_range1 + 1;
upper_range2: lower_range2 + class_widths;

lower_range3: upper_range2 + 1;
upper_range3: lower_range3 + class_widths;

lower_range4: upper_range3 + 1;
upper_range4: lower_range4 + class_widths;

lower_range5: upper_range4 + 1;
upper_range5: lower_range5 + class_widths;

error_bound:((lower_range2-upper_range1)/2);

lower_bound1: lower_range1-error_bound;  /*Adding 0.5 and subtracting 0.5 as the correcting factor for class boundaries*/
upper_bound1: upper_range1+error_bound;

lower_bound2:  lower_range2-error_bound;
upper_bound2: upper_range2+error_bound;

lower_bound3:  lower_range3-error_bound;
upper_bound3: upper_range3+error_bound;

lower_bound4:  lower_range4-error_bound;
upper_bound4: upper_range4+error_bound;

lower_bound5:  lower_range5-error_bound;
upper_bound5: upper_range5+error_bound;

lower_boundaries: [lower_bound1, lower_bound2, lower_bound3, lower_bound4, lower_bound5];
upper_boundaries: [upper_bound1, upper_bound2, upper_bound3,upper_bound4, upper_bound5];
class_widths: upper_boundaries - lower_boundaries;

midpoints: (float(lower_boundaries + class_widths / 2));

frequencies:makelist(rand_with_step(2,20,2),5);/*create a list with 5 elements randomly picked between 2-20 for even N*/
f1:frequencies[1];    
f2:frequencies[2]+6; 
f3:frequencies[3]+2; 
f4:frequencies[4]+4; 
f5:frequencies[5];    
frequencies:[f1,f2,f3,f4,f5]; 

/*cumulative frequencies*/
cf: makelist(sum(frequencies[i], i, 1, j), j, 1, length(frequencies));
/*sum(frequencies);*/
n:sum(frequencies[i],i,1,5); 
total_f:n;

midpointss: lower_boundaries + class_widths / 2;
midpoints1:decimalplaces(midpointss,2);
mptd1:round(midpoints1);
mptd_1:is(abs((midpoints1-mptd1)[1])=0.0);
midpoints: if mptd_1=true then mptd1 else midpoints1;

fxis:(midpoints * frequencies);
fxis1:decimalplaces(fxis,2);
fxis11:round(fxis);
fxis_1:is(abs((fxis1-fxis11)[1])=0.0);
fxi: if fxis_1=true then fxis11 else fxis1;
sum_fxi:sum(fxi[i],i,1,5);

/* Mean */
sum_fx:sum((midpoints * frequencies)[i], i, 1, 5);
means: sum_fx / n;
meanss:decimalplaces(means,2);
mean1:round(mean);
mean_1:is(abs((meanss-mean1))=0.0);
mean: if mean_1=true then mean1 else meanss;

/*median*/
median_cf: n / 2;

i: 1;
while cf[i] < median_cf do i: i + 1;

median_cf;
lower_bound:lower_boundaries[i];/*L is the lower limit of the median class.*/
upper_bound: upper_boundaries[i];
cf_lower: cf[i-1]; /*cf is the cumulative frequency of the class before the median class.*/
f_median: frequencies[i]; /*f is the frequency of the median class.*/
width: upper_range2 - upper_range1; /*w is the interval width (the difference between the upper and lower limits of each class).*/

median_height: lower_boundaries[i] + (median_cf - cf[i-1]) * class_widths[i] / frequencies[i];
md1: float( (median_cf - cf[i-1]) * class_widths[i] / frequencies[i]);

/*closest median height integer*/
median_height1:decimalplaces(median_height,2);
ta1:round(median_height1);
ta_1:is(abs(round(median_height1)-median_height1)=0.0);
ta: if ta_1=true then ta1 else median_height1;

/*mode*/
max_freq: apply('max, frequencies);
modal_class_index: 1;
max_value: frequencies[1];

for i: 2 thru length(frequencies) do (
    if frequencies[i] > max_value then (
        max_value: frequencies[i],
        modal_class_index: i
    )
);

Lm: lower_boundaries[modal_class_index];
f1m: frequencies[modal_class_index];
f0m: frequencies[modal_class_index-1];
f2m: frequencies[modal_class_index+1];
hm: upper_boundaries[modal_class_index] - lower_boundaries[modal_class_index];

Mode: Lm + ((f1m - f0m) / ((f1m - f0m) + (f1m - f2m))) * hm;

Modes: Mode;
Modess:decimalplaces(Modes,2);
Modes1:round(Modes);
Modes_1:is(abs((Modess-Modes1))=0.0);
Mode: if Modes_1=true then Modes1 else Modess;

mean_fdiffss1:(midpoints - mean) * frequencies;
mean_fdiffss:(midpoints - mean)^2 * frequencies;
mean_fdiffss3:(midpoints - mean)^3 * frequencies;

mean_fdiffss3s: mean_fdiffss3;
mean_fdiffss3ss:decimalplaces(mean_fdiffss3s,2);
mean_fdiffss31:round(mean_fdiffss3ss);
mean_fdiffss3_1:is(abs((mean_fdiffss3ss-mean_fdiffss31))=0.0);
mean_fdiffss3: if mean_fdiffss3_1=true then mean_fdiffss31 else mean_fdiffss3ss;

mean_fdiffss4:(midpoints - mean)^4 * frequencies;

mean_fdiffss4s: mean_fdiffss4;
mean_fdiffss4ss:decimalplaces(mean_fdiffss4s,2);
mean_fdiffss41:round(mean_fdiffss4ss);
mean_fdiffss4_1:is(abs((mean_fdiffss4ss-mean_fdiffss41))=0.0);
mean_fdiffss4: if mean_fdiffss4_1=true then mean_fdiffss41 else mean_fdiffss4ss;

M2: sum((midpoints[i] - mean)^2 * frequencies[i], i, 1, length(midpoints)) / n; /*2nd moment*/
M2s: M2;
M2ss:decimalplaces(M2s,2);
M21:round(M2s);
M2_1:is(abs((M2ss-M21))=0.0);
M2: if M2_1=true then M21 else M2ss;

sum_m2:M2*n;
M3: sum((midpoints[i] - mean)^3 * frequencies[i], i, 1, length(midpoints)) / n; /*3rd moment*/
M3s: M3;
M3ss:decimalplaces(M3s,2);
M31:round(M3s);
M3_1:is(abs((M3ss-M31))=0.0);
M3: if M3_1=true then M31 else M3ss;

sum_m3:M3*n;
sum_m3s: sum_m3;
sum_m3ss:decimalplaces(sum_m3s,2);
sum_m31:round(sum_m3s);
sum_m3_1:is(abs((sum_m3ss-sum_m31))=0.0);
sum_m3: if sum_m3_1=true then sum_m31 else sum_m3ss;

M4: sum((midpoints[i] - mean)^4 * frequencies[i], i, 1, length(midpoints)) / n; /*4th moment*/
M4s: M4;
M4ss:decimalplaces(M4s,2);
M41:round(M4s);
M4_1:is(abs((M4ss-M41))=0.0);
M4: if M4_1=true then M41 else M4ss;

sum_m4:M4*n;
sum_m4s: sum_m4;
sum_m4ss:decimalplaces(sum_m4s,2);
sum_m41:round(sum_m4ss);
sum_m4_1:is(abs((sum_m4ss-sum_m41))=0.0);
sum_m4: if sum_m4_1=true then sum_m41 else sum_m4ss;

Sk: M3 / (M2^(3/2));
Sks: Sk;
Skss:decimalplaces(Sks,2);
Sk1:round(Sks);
Sk_1:is(abs((Skss-Sk1))=0.0);
Sk: if Sk_1=true then Sk1 else Skss;


Ku: M4 / (M2^2) - 3;
Kus: Ku;
Kuss:decimalplaces(Kus,2);
Ku1:round(Kus);
Ku_1:is(abs((Kuss-Ku1))=0.0);
Ku: if Ku_1=true then Ku1 else Kuss;

mean_fdiffs1:decimalplaces(mean_fdiffss,2);
mean_fdiffs11:round(mean_fdiffss);
mean_fdiffs_1:is(abs((mean_fdiffs1-mean_fdiffs11)[1])=0.0);
mean_fdiffs: if mean_fdiffs_1=true then mean_fdiffs11 else mean_fdiffs1;

/* Variance */
mean_squared_diffs: sum((mean_fdiffs)[i],i,1,5);

mean_squared_diffss: mean_squared_diffs;
mean_squared_diffsss:decimalplaces(mean_squared_diffss,2);
mean_squared_diffs1:round(mean_squared_diffsss);
mean_squared_diffs_1:is(abs((mean_squared_diffsss-mean_squared_diffs1))=0.0);
mean_squared_diffs: if mean_squared_diffs_1=true then mean_squared_diffs1 else mean_squared_diffsss;

variances: mean_squared_diffs / n;
variancess:decimalplaces(variances,2);
variance1:round(variances);
variance_1:is(abs((variancess-variance1))=0.0);
variance: if variance_1=true then variance1 else variancess;

/* Standard Deviation */
standard_deviations: sqrt(variance);
standard_deviationss:decimalplaces(standard_deviations,2);
standard_deviation1:round(standard_deviations);
standard_deviation_1:is(abs((standard_deviationss-standard_deviation1))=0.0);
standard_deviation: if standard_deviation_1=true then standard_deviation1 else standard_deviationss;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Number of children \(={@n@}\)</p>
<p>\(\textbf{Part A}\)</p>
<p>(a) Mean \( \displaystyle \bar{x}=\frac{\sum{fX}}{\sum{f}}={@mean@}\)</p>

<p>(b) Median \(={@ta@}\)</p>

<p>(c) Mode\(={@Mode@}\)</p>

<p>(d) Standard deviation \(   \sigma ={@standard_deviation@}\) </p>

<p>(e) Variance \({ \sigma^2}={@variance@}\)</p>
<p>\(\textbf{Part B}\)</p>
    <p>
    \( \displaystyle 
\begin{align}
   (a)\,  M_2  &={@decimalplaces(M2,2)@} \\
   (b)\,    M_3 &= {@decimalplaces(M3,2)@} \\
     (c)\,  M_4 &= {@decimalplaces(M4,2)@}
    \end{align} \)
    </p>
<p>(d) \( Sk  ={@Sk@}\) </p>
<p>(e)  \( Ku={@Ku@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire.</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description"> Part A. Finding the mean, median, mode, sd, variance and filling in the table for a frequency table with \(5\) classes and Part B. The question on finding the skewness and kurtosis from a given data in table and the first,second and third moments</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>fxi[1]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>fxi[5]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>n</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>sum_fxi</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>algebraic</type>
      <tans>mean_squared_diffs</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans14</name>
      <type>algebraic</type>
      <tans>mean</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans15</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans16</name>
      <type>algebraic</type>
      <tans>Mode</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans17</name>
      <type>algebraic</type>
      <tans>standard_deviation</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans18</name>
      <type>algebraic</type>
      <tans>variance</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>midpoints[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>mean_fdiffs[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3b</name>
      <type>algebraic</type>
      <tans>mean_fdiffss3[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3d</name>
      <type>algebraic</type>
      <tans>mean_fdiffss3[4]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>cf[3]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4b</name>
      <type>algebraic</type>
      <tans>mean_fdiffss4[2]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>fxi[3]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>mean_fdiffs[3]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>midpoints[4]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>fxi[4]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>cf[5]</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>anskt</name>
      <type>algebraic</type>
      <tans>Ku</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansm2</name>
      <type>algebraic</type>
      <tans>M2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansm3</name>
      <type>algebraic</type>
      <tans>M3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansm4</name>
      <type>algebraic</type>
      <tans>M4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>anssk</name>
      <type>algebraic</type>
      <tans>Sk</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>fxi[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>fxi[5]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt11</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>n</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt11-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt11-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt12</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>sum_fxi</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt12-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt12-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt13</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>mean_squared_diffs</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt13-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt13-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt14</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans14</sans>
        <tans>mean</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt14-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt14-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt15</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans15</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt15-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt15-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt16</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans16</sans>
        <tans>Mode</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt16-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt16-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt17</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans17</sans>
        <tans>standard_deviation</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt17-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt17-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt18</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans18</sans>
        <tans>variance</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt18-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt18-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>midpoints[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mean_fdiffs[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3b</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3b</sans>
        <tans>mean_fdiffss3[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3b-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3b-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3d</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3d</sans>
        <tans>mean_fdiffss3[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3d-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3d-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>cf[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4b</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4b</sans>
        <tans>mean_fdiffss4[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4b-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4b-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>fxi[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>mean_fdiffs[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>midpoints[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>fxi[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>cf[5]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtkt</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>anskt</sans>
        <tans>Ku</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtkt-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtkt-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtm2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ansm2</sans>
        <tans>M2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtm2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtm2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtm3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ansm3</sans>
        <tans>M3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtm3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtm3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtm4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ansm4</sans>
        <tans>M4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtm4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtm4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtsk</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>anssk</sans>
        <tans>Sk</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtsk-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtsk-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1254099694</deployedseed>
    <deployedseed>361267886</deployedseed>
    <deployedseed>736450061</deployedseed>
    <deployedseed>1137644940</deployedseed>
    <deployedseed>1814204003</deployedseed>
    <deployedseed>871774896</deployedseed>
    <deployedseed>727738617</deployedseed>
    <deployedseed>105017837</deployedseed>
    <deployedseed>481094135</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>fxi[1]</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>fxi[5]</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>n</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>sum_fxi</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>mean_squared_diffs</value>
      </testinput>
      <testinput>
        <name>ans14</name>
        <value>mean</value>
      </testinput>
      <testinput>
        <name>ans15</name>
        <value>ta</value>
      </testinput>
      <testinput>
        <name>ans16</name>
        <value>Mode</value>
      </testinput>
      <testinput>
        <name>ans17</name>
        <value>standard_deviation</value>
      </testinput>
      <testinput>
        <name>ans18</name>
        <value>variance</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>midpoints[2]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>mean_fdiffs[2]</value>
      </testinput>
      <testinput>
        <name>ans3b</name>
        <value>mean_fdiffss3[2]</value>
      </testinput>
      <testinput>
        <name>ans3d</name>
        <value>mean_fdiffss3[4]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>cf[3]</value>
      </testinput>
      <testinput>
        <name>ans4b</name>
        <value>mean_fdiffss4[2]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>fxi[3]</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>mean_fdiffs[3]</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>midpoints[4]</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>fxi[4]</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>cf[5]</value>
      </testinput>
      <testinput>
        <name>anskt</name>
        <value>Ku</value>
      </testinput>
      <testinput>
        <name>ansm2</name>
        <value>M2</value>
      </testinput>
      <testinput>
        <name>ansm3</name>
        <value>M3</value>
      </testinput>
      <testinput>
        <name>ansm4</name>
        <value>M4</value>
      </testinput>
      <testinput>
        <name>anssk</name>
        <value>Sk</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt10</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt10-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt11</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt11-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt12</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt12-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt13</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt13-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt14</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt14-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt15</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt15-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt16</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt16-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt17</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt17-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt18</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt18-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3b</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3b-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3d</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3d-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4b</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4b-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtkt</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtkt-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtm2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtm2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtm3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtm3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtm4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtm4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtsk</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtsk-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
