<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Spearman's Rank Correlation and Coefficient of Determination Analysis</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the Spearman's ranks correlation coefficient between the two sets of variables given below and interpret the results</p>
<table border="solid">
  <tr>
      <th width="16%", style="text-align:center">X</th>
       <th width="10%", style="text-align:center">\({@x1@}\)</th>
       <th width="10%", style="text-align:center">\({@x2@}\)</th>
       <th width="10%", style="text-align:center">\({@x3@}\)</th>
      <th width="10%", style="text-align:center">\({@x4@}\)</th>
       <th width="10%", style="text-align:center">\({@x5@}\)</th>
      <th width="10%", style="text-align:center">\({@x6@}\)</th>
       <th width="10%", style="text-align:center">\({@x7@}\)</th>
       <th width="10%", style="text-align:center">\({@x8@}\)</th>
  </tr>
   <tr>
      <th width="16%", style="text-align:center">Y</th>
       <th width="10%", style="text-align:center">\({@y1@}\)</th>
       <th width="10%", style="text-align:center">\({@y2@}\)</th>
       <th width="10%", style="text-align:center">\({@y3@}\)</th>
      <th width="10%", style="text-align:center">\({@y4@}\)</th>
       <th width="10%", style="text-align:center">\({@y5@}\)</th>
      <th width="10%", style="text-align:center">\({@y6@}\)</th>
       <th width="10%", style="text-align:center">\({@y7@}\)</th>
       <th width="10%", style="text-align:center">\({@y8@}\)</th>
  </tr>
  
 </table>
<p>Sperman rank correlation  coefficient \(r=\) [[input:ans1]]. [[validation:ans1]][[feedback:prt1]]</p>
<p>From the sperman rank correlation  coefficient we can say that [[input:ans3]].[[validation:ans3]][[feedback:prt3]]</p>
<p><b>Note:</b> Give your answer correct to \(3 \) decimal places.</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<table border="solid">
  <tr>
      <th width="10%", style="text-align:center">X</th>
       <th width="10%", style="text-align:center">Y</th>
       <th width="15%", style="text-align:center">Rank X</th>
       <th width="15%", style="text-align:center">Rank Y</th>
      <th width="40%", style="text-align:center">d \(=\) Rank X\(-\)Rank Y</th>
       <th width="10%", style="text-align:center">\(\text{d}^2\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x1@}\)</th>
    <th width="10%", style="text-align:center">\({@y1@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[1]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[1]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[1]-rank_y[1]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[1]-rank_y[1])^2@}\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x2@}\)</th>
    <th width="10%", style="text-align:center">\({@y2@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[2]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[2]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[2]-rank_y[2]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[2]-rank_y[2])^2@}\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x3@}\)</th>
    <th width="10%", style="text-align:center">\({@y3@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[3]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[3]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[3]-rank_y[3]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[3]-rank_y[3])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x4@}\)</th>
       <th width="10%", style="text-align:center">\({@y4@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[4]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[4]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[4]-rank_y[4]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[4]-rank_y[4])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x5@}\)</th>
       <th width="10%", style="text-align:center">\({@y5@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[5]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[5]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[5]-rank_y[5]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[5]-rank_y[5])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x6@}\)</th>
       <th width="10%", style="text-align:center">\({@y6@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[6]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[6]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[6]-rank_y[6]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[6]-rank_y[6])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x7@}\)</th>
       <th width="10%", style="text-align:center">\({@y7@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[7]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[7]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[7]-rank_y[7]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[7]-rank_y[7])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x8@}\)</th>
       <th width="10%", style="text-align:center">\({@y8@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[8]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[8]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[8]-rank_y[8]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[8]-rank_y[8])^2@}\)</th>
  </tr>
  
 </table>
<p>\(\displaystyle \sum d^{2}={@total_d_squared@}\)</p>

<p>\(n=\) no. of ranks \(=8\)</p>
<p>Spearman rank correlation  coefficient \(\displaystyle r=1- \frac{ 6\sum d^2}{n(n^2-1)}\)</p>
<p>\( \begin{align} \displaystyle r &=1- \frac{6({@total_d_squared@})}{8(64-1)}\\
&=1-\frac{{@6*total_d_squared@}}{504}=1-{@(6*total_d_squared)/504@}\\
&={@ta1@}\\
\end{align} 
\)</p>
<p>Correct to three decimal places, the value of the coefficient is \({@ta1@}\). So we can say that {@mcq_correct(ta3)[1]@} .</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[x1:rand_with_step(30,80,1);
x2:rand_with_prohib(30,80,[x1]);
x3:rand_with_prohib(30,80,[x1,x2]);
x4:rand_with_prohib(30,80,[x1,x2,x3]);
x5:rand_with_prohib(30,80,[x1,x2,x3,x4]);
x6:rand_with_prohib(30,80,[x1,x2,x3,x4,x5]);
x7:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6]);
x8:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6,x7]);
x_coord_list:[x1,x2,x3,x4,x5,x6,x7,x8];
x_coord_list_ord:sort(x_coord_list);

rank_x:makelist(0,i,1,8);
 
for i:1 thru 8 do (
    for j:1 thru 8 do (
        if x_coord_list[i]=x_coord_list_ord[j] then rank_x[i]:j
    )
);
y1:rand_with_step(30,60,1);
y2:rand_with_prohib(30,60,[y1]);
y3:rand_with_prohib(30,60,[y1,y2]);
y4:rand_with_prohib(30,60,[y1,y2,y3]);
y5:rand_with_prohib(30,60,[y1,y2,y3,y4]);
y6:rand_with_prohib(30,60,[y1,y2,y3,y4,y5]);
y7:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6]);
y8:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6,y7]);
y_coord_list:[y1,y2,y3,y4,y5,y6,y7,y8];
y_coord_list_ord:sort(y_coord_list);
rank_y:makelist(0,i,1,8);
 
for i:1 thru 8 do (
    for j:1 thru 8 do (
        if y_coord_list[i]=y_coord_list_ord[j] then rank_y[i]:j
    )
);
r_1:(rank_x[1]-rank_y[1])^2;
r_2:(rank_x[2]-rank_y[2])^2;
r_3:(rank_x[3]-rank_y[3])^2;
r_4:(rank_x[4]-rank_y[4])^2;
r_5:(rank_x[5]-rank_y[5])^2;
r_6:(rank_x[6]-rank_y[6])^2;
r_7:(rank_x[7]-rank_y[7])^2;
r_8:(rank_x[8]-rank_y[8])^2;

d_squared:[r_1,r_2,r_3,r_4,r_5,r_6,r_7,r_8]; 


/*sum(d_squared)*/
n:sum(d_squared[i],i,1,8); 
total_d_squared:n;
ta1:decimalplaces(1-((6*total_d_squared)/504),3);
ta3:random_permutation([[1,true,"There is a moderate negative association between the two sets of variables."], [2,false, "There is a weak positive association between the two sets of variables."], [3,false, "There is no significant correlation between the two sets of variables.
"], [4,false, "There is a strong positive association between the two sets of variables."]]);

/*Section B*/
ta_21: if ta1 > 0 then "There is a positive association between the two sets of variables." elseif ta1 < 0 then "There is a negative association between the two sets of variables." elseif ta1 = 0 then "No correlation" else "None of the options apply";
ta_22: if ta1 > 0 then " There is a negative association between the two sets of variables." elseif ta1 < 0 then "TThere is a positive association between the two sets of variables." else "None of the options apply";
 
ta3:random_permutation([[ta_21, true],[ta_22, false],["No correlation", false]]);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Sperman rank correlation  coefficient \(r={@ta1@}\)</p>
<p>The coefficient of determination \(={@ta2@}\) </p>
<p>Correct to three decimal places, the value of the coefficient is \({@ta1@}\). So we can say that {@mcq_correct(ta3)[1]@} .</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">This question requires the student to calculate the Spearman's rank correlation coefficient between two sets of variables (X and Y) and determine the coefficient of determination. The question also asks the student to interpret the results in terms of the ranks of the variables.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>847698618</deployedseed>
    <deployedseed>1741347603</deployedseed>
    <deployedseed>99543001</deployedseed>
    <deployedseed>149227509</deployedseed>
    <deployedseed>1783172902</deployedseed>
    <deployedseed>1384446889</deployedseed>
    <deployedseed>1204413521</deployedseed>
    <deployedseed>37891233</deployedseed>
    <deployedseed>1848571744</deployedseed>
    <deployedseed>728545959</deployedseed>
    <deployedseed>59099582</deployedseed>
    <deployedseed>2073129416</deployedseed>
    <deployedseed>369193162</deployedseed>
    <deployedseed>1854134409</deployedseed>
    <deployedseed>1558826663</deployedseed>
    <deployedseed>1158648808</deployedseed>
    <deployedseed>1559058954</deployedseed>
    <deployedseed>1666402289</deployedseed>
    <deployedseed>1359390318</deployedseed>
    <deployedseed>78386672</deployedseed>
    <deployedseed>367466531</deployedseed>
    <deployedseed>962199914</deployedseed>
    <deployedseed>732506414</deployedseed>
    <deployedseed>640874612</deployedseed>
    <deployedseed>2004304425</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
