<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Study of Scatteredness in Data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Study of scatteredness of observations is known as</p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]

<p>
The term "Measure of Dispersion" refers to the study of the scatteredness or spread of observations within a data set. It is used to quantify how data points in a dataset are dispersed or spread out from the central tendency, typically represented by measures such as the mean or median.
</p>
[[/ if ]]
[[ if test="member(2, ta_index)" ]]

<p>The study of the scatteredness of observations is known as "Measure of Spread." This is a concept in statistics that describes how data points in a dataset are spread out or dispersed. The most common measures of spread include the range, variance, standard deviation, and interquartile range. These measures help to quantify the extent to which data points deviate from the central tendency (such as the mean or median) and provide insights into the variability or dispersion of the data.</p>

[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "Measure of dispersion"],
[2, true, "Measure of spread"]
], 1);
incorrect_answers:rand_selection([
[3, false, "Measure of shape"],
[4, false, "Standard deviation"],
[5, false, "Measure of central tendency"],
[6, false, "Measure of association"],
[7, false, "Interquartile range"]
], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[[[ if test="member(1, ta_index)" ]]

<p>
The term "Measure of Dispersion" refers to the study of the scatteredness or spread of observations within a data set. It is used to quantify how data points in a dataset are dispersed or spread out from the central tendency, typically represented by measures such as the mean or median.
</p>
[[/ if ]]
[[ if test="member(2, ta_index)" ]]

<p>The study of the scatteredness of observations is known as "Measure of Spread." This is a concept in statistics that describes how data points in a dataset are spread out or dispersed. The most common measures of spread include the range, variance, standard deviation, and interquartile range. These measures help to quantify the extent to which data points deviate from the central tendency (such as the mean or median) and provide insights into the variability or dispersion of the data.</p>

[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This question requires the student to identify the term that represents the study of scatteredness of observations among the given options</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>378249949</deployedseed>
    <deployedseed>2040647807</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
