<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Analyzing Independence: Assessing the Relationship Between Events A and B</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>If \(P[A]={@p_a@},P[B]={@p_b@}\) and \(P[A \cup B]={@p_aub@}\), are events \(A\) and \(B\) independent?</p>
<p>[[input:ans1]][[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>If two events \(A\) and \(B\) are independent, then we must have \(P[A \cap B] = P[A] \times P[B]\). However,</p>
<p>\(P[A \cap B]=P[A]+P[B]-P[A \cup B]={@p_a@}+{@p_b@}-{@p_aub@}={@decimalplaces((p_a+p_b-p_aub),2)@}\)</p>
<p>\(P[A] \times P[B]={@p_a@} \times {@p_b@}={@decimalplaces(p_a*p_b,4)@} \neq P[A \cap B]\), therefore they cannot be independent.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a1:rand_with_step(10,20,1);
a2:rand_with_step(30,50,1);
a3:rand_with_prohib(10,20,[a1]);
a4:100-(a1+a2+a3);
a5:a1+a2+a3+a4;
p_a:float((a2+a3)/100);
p_b:float((a3+a4)/100);
p_anb:float(a3/100);
p_aub:p_a+p_b-p_anb;
ta:[[1,false, "Yes"], [2,true, "No"]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(P[A]={@p_a@},P[B]={@p_b@}\) and \(P[A \cup B]={@p_aub@}\)</p>
<p>\(P[A \cap B]={@decimalplaces((p_a+p_b-p_aub),2)@}\)</p>
<p>\(P[A] \times P[B]={@decimalplaces(p_a*p_b,4)@} \neq P[A \cap B]\)]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/1aoDEbgrn3qrBCGbplfnTzW1Xc2g0B1p3/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question requires the student to determine the independence of events A and B based on their respective probabilities (P(A), P(B)), and the probability of their union (P(A U B)).</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>221184769</deployedseed>
    <deployedseed>921389764</deployedseed>
    <deployedseed>1320309695</deployedseed>
    <deployedseed>405223147</deployedseed>
    <deployedseed>1568432044</deployedseed>
    <deployedseed>519795669</deployedseed>
    <deployedseed>1564023649</deployedseed>
    <deployedseed>730916647</deployedseed>
    <deployedseed>1753443869</deployedseed>
    <deployedseed>2028918577</deployedseed>
    <deployedseed>582101159</deployedseed>
    <deployedseed>1035500302</deployedseed>
    <deployedseed>1081604856</deployedseed>
    <deployedseed>659223438</deployedseed>
    <deployedseed>429088553</deployedseed>
    <deployedseed>656105717</deployedseed>
    <deployedseed>1402385727</deployedseed>
    <deployedseed>270535751</deployedseed>
    <deployedseed>2046675910</deployedseed>
    <deployedseed>1282294330</deployedseed>
    <deployedseed>862778067</deployedseed>
    <deployedseed>83179052</deployedseed>
    <deployedseed>1716738812</deployedseed>
    <deployedseed>79865727</deployedseed>
    <deployedseed>992215048</deployedseed>
    <deployedseed>32998762</deployedseed>
    <deployedseed>1898889498</deployedseed>
    <deployedseed>1797960358</deployedseed>
    <deployedseed>1475621022</deployedseed>
    <deployedseed>1734059752</deployedseed>
    <deployedseed>499953495</deployedseed>
    <deployedseed>263681918</deployedseed>
    <deployedseed>1148892851</deployedseed>
    <deployedseed>1875308042</deployedseed>
    <deployedseed>1646018622</deployedseed>
    <deployedseed>22484248</deployedseed>
    <deployedseed>529000187</deployedseed>
    <deployedseed>1966619975</deployedseed>
    <deployedseed>1764043579</deployedseed>
    <deployedseed>1896239010</deployedseed>
    <deployedseed>1305129479</deployedseed>
    <deployedseed>307236633</deployedseed>
    <deployedseed>25635076</deployedseed>
    <deployedseed>771588854</deployedseed>
    <deployedseed>158847582</deployedseed>
    <deployedseed>559861939</deployedseed>
    <deployedseed>854600742</deployedseed>
    <deployedseed>1249778319</deployedseed>
    <deployedseed>113695771</deployedseed>
    <deployedseed>338649976</deployedseed>
    <deployedseed>689656904</deployedseed>
    <deployedseed>616433330</deployedseed>
    <deployedseed>1762132224</deployedseed>
    <deployedseed>1675819587</deployedseed>
    <deployedseed>231943531</deployedseed>
    <deployedseed>1886193506</deployedseed>
    <deployedseed>493653221</deployedseed>
    <deployedseed>859032028</deployedseed>
    <deployedseed>740806427</deployedseed>
    <deployedseed>110603633</deployedseed>
    <deployedseed>2067506818</deployedseed>
    <deployedseed>1988872697</deployedseed>
    <deployedseed>804723095</deployedseed>
    <deployedseed>1638738438</deployedseed>
    <deployedseed>1725730760</deployedseed>
    <deployedseed>838117723</deployedseed>
    <deployedseed>547327348</deployedseed>
    <deployedseed>421821129</deployedseed>
    <deployedseed>1106371446</deployedseed>
    <deployedseed>728154951</deployedseed>
    <deployedseed>1998565262</deployedseed>
    <deployedseed>2048611587</deployedseed>
    <deployedseed>962811580</deployedseed>
    <deployedseed>1151679745</deployedseed>
    <deployedseed>343996348</deployedseed>
    <deployedseed>1101826301</deployedseed>
    <deployedseed>42237553</deployedseed>
    <deployedseed>182516354</deployedseed>
    <deployedseed>723975004</deployedseed>
    <deployedseed>460296045</deployedseed>
    <deployedseed>1828282066</deployedseed>
    <deployedseed>530226122</deployedseed>
    <deployedseed>808842419</deployedseed>
    <deployedseed>166694470</deployedseed>
    <deployedseed>538759626</deployedseed>
    <deployedseed>1399863656</deployedseed>
    <deployedseed>236896816</deployedseed>
    <deployedseed>221179767</deployedseed>
    <deployedseed>806011662</deployedseed>
    <deployedseed>1119585600</deployedseed>
    <deployedseed>350288392</deployedseed>
    <deployedseed>208876171</deployedseed>
    <deployedseed>694336609</deployedseed>
    <deployedseed>2083550294</deployedseed>
    <deployedseed>794285011</deployedseed>
    <deployedseed>1881582294</deployedseed>
    <deployedseed>1103580873</deployedseed>
    <deployedseed>114518740</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>