<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Approximate percentile of a distribution presented as a histogram (STACK)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The histogram below shows the heights of plants measured in an experiment.</p>

[[jsxgraph width="400px" height="400px"]]

var xmin = {#xmin#};
var xmax = {#xmax#};
var bounds = {#bounds#};
var heights = {#plotheights#};

/* Leave half a class width between the histogram and each side. */
var xleft  = xmin - 0.5;
var xright = xmax + 0.5;

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [xleft - 0.15, 6, xright + 0.15, -0.9],
    axis: false,
    keepaspectratio: true,
    showCopyright: false,
    showNavigation: false
});


/* Horizontal guide lines. */
for (var y = 1; y <= 5; y++) {
    board.create('segment',
        [[xleft, y], [xright, y]],
        {
            strokeWidth: 1,
            strokeOpacity: 0.35,
            fixed: true,
            highlight: false
        }
    );
}


/* Draw histogram bars. */
for (var i = 0; i < 6; i++) {

    var vertices = [
        [bounds[i], 0],
        [bounds[i + 1], 0],
        [bounds[i + 1], heights[i]],
        [bounds[i], heights[i]]
    ];

    board.create('polygon', vertices, {
        fillOpacity: 0.5,
        fixed: true,
        highlight: false,

        vertices: {
            visible: false,
            fixed: true,
            highlight: false
        },

        borders: {
            strokeWidth: 2,
            fixed: true,
            highlight: false
        }
    });
}


/* Horizontal axis. */
board.create('segment',
    [[xleft, 0], [xright, 0]],
    {
        strokeWidth: 2,
        fixed: true,
        highlight: false
    }
);


/* Vertical axis, separated from the first bar. */
board.create('segment',
    [[xleft, 0], [xleft, 5.5]],
    {
        strokeWidth: 2,
        fixed: true,
        highlight: false
    }
);


/* Tick marks and labels at the class boundaries. */
for (var i = 0; i < bounds.length; i++) {

    var x = bounds[i];

    board.create('segment',
        [[x, -0.08], [x, 0.08]],
        {
            strokeWidth: 1,
            fixed: true,
            highlight: false
        }
    );

    board.create('text',
        [x, -0.18, x.toFixed(1)],
        {
            anchorX: 'middle',
            anchorY: 'top',
            fontSize: 16,
            fixed: true,
            highlight: false
        }
    );
}


/* Horizontal axis label. */
board.create('text',
    [(xleft + xright) / 2, -0.58, 'Height (cm)'],
    {
        anchorX: 'middle',
        anchorY: 'top',
        fontSize: 16,
        fixed: true,
        highlight: false
    }
);

[[/jsxgraph]]

<p>Estimate the \(75\)th percentile of the plant heights: [[input:ans1]] [[validation:ans1]] cm [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The \(75\)th percentile is the height below which approximately \(75\%\) of the observations lie.</p>

<p>Starting from the left of the histogram, consider the area of each bar and accumulate these areas until you reach about \(75\%\) of the total area. The \(75\)th percentile lies within the bar where this occurs. Estimate how far through that bar you need to go to reach \(75\%\).</p>

<p>For this histogram, the \(75\)th percentile is approximately \({@dispdp(ta,1)@}\) cm.</p>]]></text>
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Function to round values to a given number of decimal places. */
rnd(x,dp):=float(round(x*10^dp)/10^dp);


/* Randomise the first bin centre.
   The remaining five centres increase by 1. */
xstart:rand(8)+4;
xvals:makelist(xstart+i,i,0,5);


/* Six equal-width classes, each of width 1. */
xmin:xstart-1/2;
xmax:xstart+11/2;
bounds:makelist(xmin+i,i,0,6);


/* Place the normal distribution exactly in the middle
   of the displayed histogram range. */
mean:xstart+5/2;


/* Keep the standard deviation fairly small so that
   very few observations fall outside the displayed range. */
sd:1.2;


/* Generate 100 normally distributed observations,
   rounded to 3 decimal places. */
n:100;

dun:makelist(
    rnd(float(quantile_normal(rand(1.0),mean,sd)),3),
    i,1,n
);


/* Clamp any observations which fall outside the
   displayed histogram range. */
data:map(
    lambda([x],
        if x<xmin then xmin
        elseif x>xmax then xmax
        else x
    ),
    dun
);


/* Sort the observations for calculating the percentile. */
sdata:sort(data);


/*
   Estimate the sample 75th percentile by interpolation.

   Using the position:
       1 + (n-1)*0.75

   For n=100 this is 75.25, so Q3 lies one quarter
   of the way between observations 75 and 76.
*/
qpos:1+(n-1)*3/4;
qlo:floor(qpos);
qfrac:qpos-qlo;

ta:sdata[qlo]+qfrac*(sdata[qlo+1]-sdata[qlo]);


/* Count the number of observations in each histogram class. */
barheights:makelist(
    length(
        sublist(
            data,
            lambda([x],
                if i<6 then
                    is(x>=bounds[i] and x<bounds[i+1])
                else
                    is(x>=bounds[i] and x<=bounds[i+1])
            )
        )
    ),
    i,1,6
);

/* Find the maximum bar height. */
maxheight:barheights[1];

for i:2 thru 6 do (
    if barheights[i]>maxheight then
        maxheight:barheights[i]
);

/* Scale the tallest bar to height 5. */
plotheights:makelist(
    5*barheights[i]/maxheight,
    i,1,6
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>xmin = {@xmin@}, bar heights = {@barheights@}, correct answer = {@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.04</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Great estimation.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.07</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Good estimation, but you could be improved.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your estimate is not close enough.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1540749599</deployedseed>
    <deployedseed>1780942287</deployedseed>
    <deployedseed>591784999</deployedseed>
    <deployedseed>1943870494</deployedseed>
    <deployedseed>1166670793</deployedseed>
    <deployedseed>591410521</deployedseed>
    <deployedseed>340438751</deployedseed>
    <deployedseed>1063080676</deployedseed>
    <deployedseed>440731585</deployedseed>
    <deployedseed>1030173075</deployedseed>
    <deployedseed>798266329</deployedseed>
    <deployedseed>434715193</deployedseed>
    <deployedseed>486681146</deployedseed>
    <deployedseed>1014073919</deployedseed>
    <deployedseed>94930815</deployedseed>
    <deployedseed>748855066</deployedseed>
    <deployedseed>1035899052</deployedseed>
    <deployedseed>2002464454</deployedseed>
    <deployedseed>1126278837</deployedseed>
    <deployedseed>1414664510</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>