<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Assessing Product Success: Probability Analysis and Review Evaluation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Customers are used to evaluate preliminary product designs. In the past  \({@A@} \%\) of highly successful products recived good reviews \(, {@B@} \%\) of moderately successful products recived good reviews, and \({@C@} \%\) of poor products recived good reviews. In addition, \({@product_A@}\% \) of the products have been highly successful \(, {@product_B@}\%\) have been moderately successful and \({@product_C@}\%\) have been poor products. 
<p><b>Note:</b> Give your answer to 3 decimal places.</p>
<p>a) What is the probability that a product attain good review?</p>
<p>\(P(G)=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b) If a new design attain a good review, what is the probability that it will be a highly successful product?</p>
<p>\(P(H|G)=\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>c)  If a product does not attain a good review, what is the probability that it will be highly successful?
<p>\(P(H|G^{\prime})=\)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>a)  \(P(G)=P(H)P(G|H)+P(M)P(G|M)+P(P)P(G|P)\)</p>
<p> \(P(G)=({@prob_A@} \times {@product_A_1@})+({@prob_B@} \times {@product_B_1@})+({@prob_C@} \times {@product_C_1@})\)</p>
<p> \(P(G)={@ta_1@}\)</p>
<p>b) \(\displaystyle P(H|G)=\frac{ P(G|H) P(H)}{ P(G)}\) </p>
<p> \(\displaystyle P(H|G)=\frac{ {@prob_A@}({@product_A_1@})}{ {@ta_1@}}\) </p>
<p>\(\displaystyle P(H|G)={@ta_2@}\) </p>

<p>c) \(\displaystyle P(H|G^{\prime})=\frac{ P(G^{\prime}|H) P(H)}{ P(G^{\prime})}\) </p>
<p> \(\displaystyle P(H|G^{\prime})=\frac{ {@1-prob_A@}({@product_A_1@})}{ 1-{@ta_1@}}\) </p>
<p>\(\displaystyle P(H|G^{\prime})={@ta_3@}\) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>C:rand_with_step(10,30,5);/*poor products receiving good news */
prob_C:decimalplaces((C/100),2);
B:rand_with_step(10,30,5)+C;/*moderately successful products receiving good news*/
prob_B:decimalplaces((B/100),2);
A:rand_with_step(10,30,5)+B;/*Highly succesful products receiving good news*/
prob_A:decimalplaces((A/100),2);
product_A:rand(30)+10;/*products that have been highly succesful*/
product_B:rand(30)+10;/*products that have been moderately succesful*/
product_C:100-(product_A+product_B);/*products that have been poor*/


product_A_1:decimalplaces((product_A/100),2);
product_B_1:decimalplaces((product_B/100),2);
product_C_1:decimalplaces((product_C/100),2);
prob_good_review:(prob_A*product_A_1)+(prob_B*product_B_1)+(prob_C*product_C_1);
prob_good_success:(prob_A*product_A_1)/prob_good_review;
prob_poor_success:((1-prob_A)*product_A_1)/(1-prob_good_review);
ta_1:decimalplaces(prob_good_review,3);
ta_2:decimalplaces(prob_good_success,3);

ta_3:decimalplaces(prob_poor_success,3);
/*mistakes*/
ta_11:round(ta_1);
ta_21:round(decimalplaces(prob_good_success,3));

ta_31:round(decimalplaces(prob_poor_success,3));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(P(H)={@prob_A@}, P(M)={@prob_B@}, P(P)={@prob_C@}\)</p>
<p>\(P(G|H)={@product_A_1@}, P(G|M)={@product_B_1@}\) and \(P(G|P)={@product_C_1@}\)</p>
<p>\(P(G)={@ta_1@}\)</p>

<p>\(P(H|G)={@ta_2@}\)</p>

<p>\(P(H|G^{\prime})={@ta_3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/1Y_962hRAywGPoE-glatkmbcuOOKxz9UP/edit</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question requires the student to analyze the probabilities of product success based on preliminary reviews, considering varying success levels and review outcomes.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer with many decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have given your answer without considering the expected number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer with many decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have given your answer without considering the expected number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer with many decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have given your answer without considering the expected number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1843742595</deployedseed>
    <deployedseed>1460324934</deployedseed>
    <deployedseed>1394251052</deployedseed>
    <deployedseed>579462721</deployedseed>
    <deployedseed>558652644</deployedseed>
    <deployedseed>450250329</deployedseed>
    <deployedseed>36757050</deployedseed>
    <deployedseed>640565937</deployedseed>
    <deployedseed>152045102</deployedseed>
    <deployedseed>714322453</deployedseed>
    <deployedseed>442508298</deployedseed>
    <deployedseed>37940691</deployedseed>
    <deployedseed>1249513398</deployedseed>
    <deployedseed>248781555</deployedseed>
    <deployedseed>1982584264</deployedseed>
    <deployedseed>585510248</deployedseed>
    <deployedseed>1759317786</deployedseed>
    <deployedseed>1846505790</deployedseed>
    <deployedseed>237694389</deployedseed>
    <deployedseed>2039541696</deployedseed>
    <deployedseed>370673313</deployedseed>
    <deployedseed>242195716</deployedseed>
    <deployedseed>1414718753</deployedseed>
    <deployedseed>1456248028</deployedseed>
    <deployedseed>773237383</deployedseed>
    <deployedseed>1226722004</deployedseed>
    <deployedseed>2117370418</deployedseed>
    <deployedseed>1861531146</deployedseed>
    <deployedseed>933445499</deployedseed>
    <deployedseed>324328847</deployedseed>
    <deployedseed>516345038</deployedseed>
    <deployedseed>70471170</deployedseed>
    <deployedseed>2084767908</deployedseed>
    <deployedseed>1874969458</deployedseed>
    <deployedseed>556219020</deployedseed>
    <deployedseed>1708331096</deployedseed>
    <deployedseed>940687247</deployedseed>
    <deployedseed>85450003</deployedseed>
    <deployedseed>473766137</deployedseed>
    <deployedseed>1776841477</deployedseed>
    <deployedseed>860972382</deployedseed>
    <deployedseed>1978910870</deployedseed>
    <deployedseed>479112847</deployedseed>
    <deployedseed>1442476124</deployedseed>
    <deployedseed>683268370</deployedseed>
    <deployedseed>1878954746</deployedseed>
    <deployedseed>983965521</deployedseed>
    <deployedseed>2036227117</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>