<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Basic Demographic Measures_Q4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Identify methods that can be used to obtain Vital Statistics, select all that apply:</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>Population registers - Population registers are official records maintained by governments that contain information about individuals within a specific geographic area. They often include details like births, deaths, marriages, and other demographic events. This method is effective for obtaining vital statistics.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>Population census - A population census is a complete enumeration of a population, typically conducted at regular intervals. It collects a wide range of demographic information, including births, deaths, and other vital events. Population censuses are crucial for obtaining comprehensive vital statistics.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Sample surveys involve collecting data from a subset or sample of the population rather than the entire population. By carefully selecting a representative sample, researchers can extrapolate the findings to make inferences about the entire population. Sample surveys can be used to gather information on vital events, including births and deaths, and can provide estimates of vital statistics.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a:2;
correct_answer:rand_selection([
[1, true, "Population registers"],
[2, true, "Population census"],
[3, true, "Sample survey"]],a);

b:2;
incorrect_answers:rand_selection([
[4, false, "Fecundity"],
[5, false, "Migration"],
[6, false, "Calendars "],
[7, false, "Questionnaires"]], b);
ta:random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@mcq_correct(ta)@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire<p>
<p id="concept">https://docs.google.com/document/d/1mDXYsnqCXrevmSMkXVQb77HkuurpnSzBneht0i0Gl8g/edit#heading=h.j3chgt1zcto0</p>
<p id="reviewer">1</p>
<p id="reviewer">2</p>
<p id="description">The question focuses on Identifying methods that can be used to obtain Vital Statistics</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1763870324</deployedseed>
    <deployedseed>533066621</deployedseed>
    <deployedseed>1554131498</deployedseed>
    <deployedseed>140717616</deployedseed>
    <deployedseed>1379532310</deployedseed>
    <deployedseed>2125295570</deployedseed>
    <deployedseed>1928711442</deployedseed>
    <deployedseed>136278076</deployedseed>
    <deployedseed>1565933778</deployedseed>
    <deployedseed>361920987</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>