<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>BASIC MEASURES OF FERTILITY_q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>a) Complete the table below:</p>
<p><b>Note:</b> Give your answers correct to \(4\) decimal places.</p>
<p>
<table border="solid">
 
  <tr>
    <th width="15%" style="text-align:center">Age group
</th>
    <th width="15%" style="text-align:center">No. of women \( (‘000’) \)</th>
    <th width="15%" style="text-align:center">No. of live births</th>
 <th width="15%" style="text-align:center">Age specific fertility rate</th>

  </tr>
  <tr>
    <td style="text-align:center">{@age_group1@}-{@age_group2@}</td>
    <td style="text-align:center">{@fem_pop1@}</td>
    <td style="text-align:center">{@fem_birth1@}</td>
 <td style="text-align:center">{@ASFR_1@}</td>

  </tr>
  <tr>
    <td style="text-align:center">{@age_group2@}-{@age_group3@}</td>
    <td style="text-align:center">{@fem_pop2@}</td>
    <td style="text-align:center">{@fem_birth2@}</td>
<td style="text-align:center">{@ASFR_2@}</td>
  
  </tr>
  <tr>
    <td style="text-align:center">{@age_group3@}-{@age_group4@}</td>
    <td style="text-align:center">{@fem_pop3@}</td>
    <td style="text-align:center">{@fem_birth3@}</td>
 <td style="text-align:center">[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</td>
  
  </tr>
  <tr>
    <td style="text-align:center">{@age_group4@}-{@age_group5@}</td>
    <td style="text-align:center">{@fem_pop4@}</td>
    <td style="text-align:center">{@fem_birth4@}</td>
 <td style="text-align:center">{@ASFR_4@}</td>
    
  </tr>
  <tr>
    <td style="text-align:center">{@age_group5@}-{@age_group6@}</td>
    <td style="text-align:center">{@fem_pop5@}</td>
    <td style="text-align:center">{@fem_birth5@}</td>
 <td style="text-align:center">[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</td>
 
  </tr>
  <tr>
    <td style="text-align:center">{@age_group6@}-{@age_group7@}</td>
    <td style="text-align:center">{@fem_pop6@}</td>
    <td style="text-align:center">{@fem_birth6@}</td>
 <td style="text-align:center">{@ASFR_6@}</td>

  </tr>
  <tr>
    <td style="text-align:center">{@age_group7@}-{@age_group8@}</td>
    <td style="text-align:center">{@fem_pop7@}</td>
    <td style="text-align:center">{@fem_birth7@}</td>
 <td style="text-align:center">{@ASFR_7@}</td>
      
</th>
  </tr>
 <tr>
   <th width="15%" style="text-align:center">Total
</th>
    <td style="text-align:center">{@total_fem_pop@}</td>
    <td style="text-align:center">{@total_fem_birth@}</td>
     
</th>
  </tr>
</table>
</p>
<p>It is given that out of \({@total_fem_birth@}\) the number of female births was \({@num_female_birth@}\).</p>

<p>b).General Fertility Rate\(\textbf{(GFR)}= \) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<p>c). Total Fertility Rate \(= \) [[input:ans4]] [[validation:ans4]][[feedback:prt4]]</p>
<p>d). Gross Reproductive Rate \(= \) [[input:ans5]] [[validation:ans5]][[feedback:prt5]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

a)
<table border="solid">

  <tr>
    <th width="15%" style="text-align:center">Age group
</th>
    <th width="15%" style="text-align:center">No. of women</th>
    <th width="15%" style="text-align:center">No. of live births</th>
  <th width="15%" style="text-align:center">Age specific fertility rate per woman (ASFR)</th>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group1@}-{@age_group2@}</td>
    <td style="text-align:center">{@fem_popt1@}</td>
    <td style="text-align:center">{@fem_birth1@}</td>
<td style="text-align:center">\( \displaystyle \frac{{@fem_birth1@}}{{@fem_popt1@}} \times 1000 = {@ASFR_1@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group2@}-{@age_group3@}</td>
    <td style="text-align:center">{@fem_popt2@}</td>
    <td style="text-align:center">{@fem_birth2@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth2@}}{{@fem_popt2@}} \times 1000 = {@ASFR_2@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group3@}-{@age_group4@}</td>
    <td style="text-align:center">{@fem_popt3@}</td>
    <td style="text-align:center">{@fem_birth3@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth3@}}{{@fem_popt3@}}  \times 1000 = {@ASFR_3@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group4@}-{@age_group5@}</td>
    <td style="text-align:center">{@fem_popt4@}</td>
    <td style="text-align:center">{@fem_birth4@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth4@}}{{@fem_popt4@}}  \times 1000 = {@ASFR_4@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group5@}-{@age_group6@}</td>
    <td style="text-align:center">{@fem_popt5@}</td>
    <td style="text-align:center">{@fem_birth5@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth5@}}{{@fem_popt5@}}  \times 1000 = {@ASFR_5@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group6@}-{@age_group7@}</td>
    <td style="text-align:center">{@fem_popt6@}</td>
    <td style="text-align:center">{@fem_birth6@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth6@}}{{@fem_popt6@}}  \times 1000 = {@ASFR_6@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center">{@age_group7@}-{@age_group8@}</td>
    <td style="text-align:center">{@fem_popt7@}</td>
    <td style="text-align:center">{@fem_birth7@}</td>
    <td style="text-align:center">\( \displaystyle \frac{{@fem_birth7@}}{{@fem_popt7@}}  \times 1000 = {@ASFR_7@} \)</td>
  </tr>
  <tr>
    <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center"></td>
    <td style="text-align:center">\( \displaystyle \sum ASFR ={@sum_ASFR@}\)</td>
  </tr>
</table>
<p>b) General Fertility Rate \(\displaystyle \frac{\text {number of live births}}{ \text{number of women age}     {@age_group1@}-{@age_group8@}} \times 1000=\frac{{@total_fem_birth@}}{{@total_fem_pop*100@}} \times 1000={@ta1@}\)</p>

<p>c)</p> 
<table border="solid">

  <tr>
    <th width="15%" style="text-align:center">Exact Age (years)
</th>
    <th width="15%" style="text-align:center">\(ASFR \times {@age@}\)</th>
   
  <th width="15%" style="text-align:center">Total births per \(1000\) females aged \({@age_group1@}\) by stated ages</th>
 
  </tr>
  <tr>
    <td style="text-align:center">{@age_group1@}</td>
    <td style="text-align:center">\(0\)</td>
    <td style="text-align:center">\(0\)</td>

  </tr>
  <tr>
    <td style="text-align:center">{@age_group2@}</td>
    <td style="text-align:center">\({@ASFR_1@} \times {@age@}={@ASFR_s1@}\)</td>
    <td style="text-align:center">{@aa@}</td>
    
  </tr>
  <tr>
    <td style="text-align:center">{@age_group3@}</td>
    <td style="text-align:center">\({@ASFR_2@} \times {@age@}={@ ASFR_s2@}\)</td>
    <td style="text-align:center">{@bb@}</td>
    
  </tr>
  <tr>
    <td style="text-align:center">{@age_group4@}</td>
    <td style="text-align:center">\({@ASFR_3@} \times {@age@}={@ASFR_s3@}\)</td>
    <td style="text-align:center">{@cc@}</td>
    
  </tr>
  <tr>
    <td style="text-align:center">{@age_group5@}</td>
    <td style="text-align:center">\({@ASFR_4@} \times {@age@}={@ASFR_s4@}\)</td>
    <td style="text-align:center">{@dd@}</td>
  
  </tr>
  <tr>
    <td style="text-align:center">{@age_group6@}</td>
    <td style="text-align:center">\({@ASFR_5@} \times {@age@}={@ASFR_s5@}\)</td>
    <td style="text-align:center">{@ee@}</td>
 
  </tr>
  <tr>
    <td style="text-align:center">{@age_group7@}</td>
    <td style="text-align:center">\({@ASFR_6@} \times {@age@}={@ASFR_s6@}\)</td>
    <td style="text-align:center">{@ff@}</td>
    
  </tr>
<tr>
    <td style="text-align:center">{@age_group8+1@}</td>
    <td style="text-align:center">\({@ASFR_7@} \times {@age@}={@ASFR_s7@}\)</td>
    <td style="text-align:center">{@gg@}</td>
    
  </tr>
 
</table>
  <p>Total Fertility Rate \( =\displaystyle \frac {\sum ASFR}{1000} =\frac{{@gg@}}{1000}={@ta2@}\)
  <p>d) GRR \(= \frac{\text{Number of female births}}{\text{Total births}} \times \text{Total Fertility Rate} =\displaystyle \frac{{@num_female_birth@}}{{@total_fem_birth@}} \times {@ta2@} ={@ta3@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*general and fertility rates*/
age:rand(3)+4;
age_group1:rand_with_step(13,17,1);
age_group2:age_group1+age;
age_group3:age_group2+age;
age_group4:age_group3+age;
age_group5:age_group4+age;
age_group6:age_group5+age;
age_group7:age_group6+age;
age_group8:age_group7+age;

/*Female population (‘000’)*/
fem_pop1:rand_with_step(10,100,5);
fem_pop2:rand_with_prohib(10,100,[fem_pop1]);
fem_pop3:rand_with_prohib(10,100,[fem_pop1,fem_pop2]);
fem_pop4:rand_with_prohib(10,100,[fem_pop1,fem_pop2,fem_pop3]);
fem_pop5:rand_with_prohib(10,100,[fem_pop1,fem_pop2,fem_pop3,fem_pop4]);
fem_pop6:rand_with_prohib(10,100,[fem_pop1,fem_pop2,fem_pop3,fem_pop4,fem_pop5]);
fem_pop7:rand_with_prohib(10,100,[fem_pop1,fem_pop2,fem_pop3,fem_pop4,fem_pop5,fem_pop6]);
total_fem_pop:fem_pop1+fem_pop2+fem_pop3+fem_pop4+fem_pop5+fem_pop6+fem_pop7;
 /*Total fem pop*/
fem_popt1:fem_pop1*1000;
fem_popt2:fem_pop2*1000;
fem_popt3:fem_pop3*1000;
fem_popt4:fem_pop4*1000;
fem_popt5:fem_pop5*1000;
fem_popt6:fem_pop6*1000;
fem_popt7:fem_pop7*1000;

/*Number of live births*/
fem_birth1:rand_with_step(100,7000,50);
fem_birth2:rand_with_prohib(100,7000,[fem_birth1]);
fem_birth3:rand_with_prohib(100,7000,[fem_birth1,fem_birth2]);
fem_birth4:rand_with_prohib(100,7000,[fem_birth1,fem_birth2,fem_birth3]);
fem_birth5:rand_with_prohib(100,7000,[fem_birth1,fem_birth2,fem_birth3,fem_birth4]);
fem_birth6:rand_with_prohib(100,7000,[fem_birth1,fem_birth2,fem_birth3,fem_birth4,fem_birth5]);
fem_birth7:rand_with_prohib(100,7000,[fem_birth1,fem_birth2,fem_birth3,fem_birth4,fem_birth5,fem_birth6]);
total_fem_birth: fem_birth1+fem_birth2+fem_birth3+fem_birth4+fem_birth5+fem_birth6+fem_birth7;



/*worked solution variables*/
ASFR_1:decimalplaces(((fem_birth1/fem_popt1)*1000),4);
ASFR_2:decimalplaces(((fem_birth2/fem_popt2)*1000),4);
ASFR_3:decimalplaces(((fem_birth3/fem_popt3)*1000),4);
ASFR_4:decimalplaces(((fem_birth4/fem_popt4)*1000),4);
ASFR_5:decimalplaces(((fem_birth5/fem_popt5)*1000),4);
ASFR_6:decimalplaces(((fem_birth6/fem_popt6)*1000),4);
ASFR_7:decimalplaces(((fem_birth7/fem_popt7)*1000),4);
sum_ASFR:decimalplaces((ASFR_1+ASFR_2+ASFR_3+ASFR_4+ASFR_5+ASFR_6+ASFR_7),2);

/*ASFR*Survival rate(S)*/
ASFR_s1:decimalplaces((ASFR_1*age),4);
ASFR_s2:decimalplaces((ASFR_2*age),4);
ASFR_s3:decimalplaces((ASFR_3*age),4);
ASFR_s4:decimalplaces((ASFR_4*age),4);
ASFR_s5:decimalplaces((ASFR_5*age),4);
ASFR_s6:decimalplaces((ASFR_6*age),4);
ASFR_s7:decimalplaces((ASFR_7*age),4);
aa:ASFR_s1;bb:ASFR_s1+ASFR_s3;cc:ASFR_s1+ASFR_s2+ASFR_s3;dd:ASFR_s1+ASFR_s2+ASFR_s3+ASFR_s4; ee:ASFR_s1+ASFR_s2+ASFR_s3+ASFR_s4+ASFR_s5;ff:ASFR_s1+ASFR_s2+ASFR_s3+ASFR_s4+ASFR_s5+ASFR_s6;gg:ASFR_s1+ASFR_s2+ASFR_s3+ASFR_s4+ASFR_s5+ASFR_s6+ASFR_s7;
sum_ASFR_s:aa+bb+cc+dd+ee+ff+gg;
num_female_birth:round(total_fem_birth*(0.4+rand(0.2)));
/*model answers*/

ta1:decimalplaces((total_fem_birth/(total_fem_pop*100))*1000,4);/*Gross Reproduction Rate (GRR)*/
ta2:decimalplaces((gg/1000),4);
ta3:decimalplaces(((num_female_birth/total_fem_birth)*ta2),4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Total number of live births \(={@total_fem_birth@}\)</p>
<p>Number of female birth \(={@num_female_birth@}\)</p>
<p>General Fertility Rate \(={@ta1@} \)</p>
<p>Total Fertility Rate \(={@ta2@} \)</p>
<p>Gross Reproductive Rate \(={@ta3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer"></p>
<p id="description">This question requires the student to calculate the General Fertility Rate, Total Fertility Rate and Gross Reproductive Rate using provided data on female population across age groups and number of female birth.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ASFR_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ASFR_5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ASFR_3</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for number of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ASFR_3</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not given your answer to the right number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ASFR_5</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for number of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ASFR_5</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not given your answer to the right number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for number of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta1</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not given your answer to the right number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for number of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta2</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not given your answer to the right number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for number of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>ta3</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have not given your answer to the right number of decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1620449862</deployedseed>
    <deployedseed>1775967909</deployedseed>
    <deployedseed>1363120725</deployedseed>
    <deployedseed>1248135504</deployedseed>
    <deployedseed>1270718576</deployedseed>
    <deployedseed>382499623</deployedseed>
    <deployedseed>991441889</deployedseed>
    <deployedseed>973018585</deployedseed>
    <deployedseed>120201689</deployedseed>
    <deployedseed>335796161</deployedseed>
    <deployedseed>1253948950</deployedseed>
    <deployedseed>415930698</deployedseed>
    <deployedseed>499037337</deployedseed>
    <deployedseed>2126963698</deployedseed>
    <deployedseed>1776322144</deployedseed>
    <deployedseed>1267141851</deployedseed>
    <deployedseed>1801797291</deployedseed>
    <deployedseed>1048774902</deployedseed>
    <deployedseed>1652258526</deployedseed>
    <deployedseed>250065166</deployedseed>
    <deployedseed>1376433220</deployedseed>
    <deployedseed>1449692201</deployedseed>
    <deployedseed>314208793</deployedseed>
    <deployedseed>1366948100</deployedseed>
    <deployedseed>964386449</deployedseed>
    <deployedseed>2030809906</deployedseed>
    <deployedseed>1658446659</deployedseed>
    <deployedseed>1473632925</deployedseed>
    <deployedseed>948147864</deployedseed>
    <deployedseed>1531379508</deployedseed>
    <deployedseed>1833004933</deployedseed>
    <deployedseed>460974929</deployedseed>
    <deployedseed>1394720159</deployedseed>
    <deployedseed>1023784190</deployedseed>
    <deployedseed>1738660545</deployedseed>
    <deployedseed>1947482790</deployedseed>
    <deployedseed>1131438476</deployedseed>
    <deployedseed>1721060006</deployedseed>
    <deployedseed>1380096308</deployedseed>
    <deployedseed>1204210367</deployedseed>
    <deployedseed>1130178576</deployedseed>
    <deployedseed>895053580</deployedseed>
    <deployedseed>123182048</deployedseed>
    <deployedseed>1732711209</deployedseed>
    <deployedseed>505911473</deployedseed>
    <deployedseed>1094828977</deployedseed>
    <deployedseed>57473650</deployedseed>
    <deployedseed>1100930998</deployedseed>
    <deployedseed>1691041414</deployedseed>
    <deployedseed>625070447</deployedseed>
    <deployedseed>592833646</deployedseed>
    <deployedseed>800283679</deployedseed>
    <deployedseed>1759542523</deployedseed>
    <deployedseed>84390160</deployedseed>
    <deployedseed>571183706</deployedseed>
    <deployedseed>122110804</deployedseed>
    <deployedseed>143997257</deployedseed>
    <deployedseed>120221426</deployedseed>
    <deployedseed>90835780</deployedseed>
    <deployedseed>1940253709</deployedseed>
    <deployedseed>773446707</deployedseed>
    <deployedseed>916507904</deployedseed>
    <deployedseed>1409826617</deployedseed>
    <deployedseed>1454027667</deployedseed>
    <deployedseed>1618024659</deployedseed>
    <deployedseed>690824500</deployedseed>
    <deployedseed>1307641364</deployedseed>
    <deployedseed>1375759724</deployedseed>
    <deployedseed>299252547</deployedseed>
    <deployedseed>1550187390</deployedseed>
    <deployedseed>2111291908</deployedseed>
    <deployedseed>2008682894</deployedseed>
    <deployedseed>716292020</deployedseed>
    <deployedseed>10387439</deployedseed>
    <deployedseed>1546116880</deployedseed>
    <deployedseed>1471746408</deployedseed>
    <deployedseed>796059494</deployedseed>
    <deployedseed>552256258</deployedseed>
    <deployedseed>1125499886</deployedseed>
    <deployedseed>1812060948</deployedseed>
    <deployedseed>10224197</deployedseed>
    <deployedseed>372771531</deployedseed>
    <deployedseed>425843329</deployedseed>
    <deployedseed>1489407134</deployedseed>
    <deployedseed>698220446</deployedseed>
    <deployedseed>700348352</deployedseed>
    <deployedseed>805049074</deployedseed>
    <deployedseed>1164958348</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ASFR_2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ASFR_5</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>