<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Basic Demographic Measures_Q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Death rates have decreased due to: </p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>Death rates have decreased due to <i>Better health care</i> - The improvement in healthcare has been a significant factor contributing to the decline in death rates over the years. Better healthcare encompasses a range of advancements, practices, and policies that collectively work to enhance the overall well-being of individuals and address various health challenges.</p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>Death rates have decreased due to <i>Better diet</i> - The idea that a better diet can contribute to a decrease in death rates is rooted in the understanding that nutritional habits play a crucial role in overall health and well-being.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Death rates have decreased due to <i>Better Education</i> - Education often leads to increased awareness of health-related issues. Well-educated individuals are more likely to understand the importance of preventive measures, healthy lifestyles, and regular medical check-ups. This knowledge can contribute to early detection and management of health conditions, ultimately reducing mortality rates.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "Better health care"],
[2, true, "Better diet"],
[3, true, "Better Education"]],1);
incorrect_answers:rand_selection([
[4, false, "Fat people"],
[5, false, "Less murders"],
[6, false, "No wars"],
[7, false, "Control of CORONA virus"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Death rates have decreased due to: </p>

[[ if test="member(1, ta_index)" ]]
<p>{@ta@}</p>
<p>Better health care.</p>
[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>{@ta@}</p>
<p>Better diet.</p>

[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>{@ta@}</p>
<p>Better Education.</p>
[[/ if ]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat<p>
<p id="concept">https://docs.google.com/document/d/1XGyaU3wIcl2by1bN8UzAA_0O7CoOy--N9bJALFsrvOE/edit#heading=h.j3chgt1zcto0</p>
<p id="reviewer">Mary</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question requires the student to examine and elucidate the various factors contributing to the observed decrease in death rates, offering a comprehensive analysis of the multifaceted influences shaping mortality trends.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2110384822</deployedseed>
    <deployedseed>617513444</deployedseed>
    <deployedseed>1473694765</deployedseed>
    <deployedseed>691684705</deployedseed>
    <deployedseed>1097936011</deployedseed>
    <deployedseed>2000112007</deployedseed>
    <deployedseed>401255606</deployedseed>
    <deployedseed>111966407</deployedseed>
    <deployedseed>561520915</deployedseed>
    <deployedseed>52374172</deployedseed>
    <deployedseed>144293223</deployedseed>
    <deployedseed>1480481053</deployedseed>
    <deployedseed>3559999</deployedseed>
    <deployedseed>1060267022</deployedseed>
    <deployedseed>1749798921</deployedseed>
    <deployedseed>601837857</deployedseed>
    <deployedseed>790457585</deployedseed>
    <deployedseed>2067089445</deployedseed>
    <deployedseed>1844034647</deployedseed>
    <deployedseed>643250559</deployedseed>
    <deployedseed>199377492</deployedseed>
    <deployedseed>73531361</deployedseed>
    <deployedseed>782022479</deployedseed>
    <deployedseed>1794501044</deployedseed>
    <deployedseed>1833521282</deployedseed>
    <deployedseed>2111095296</deployedseed>
    <deployedseed>1438285154</deployedseed>
    <deployedseed>2137697204</deployedseed>
    <deployedseed>1926105434</deployedseed>
    <deployedseed>141917993</deployedseed>
    <deployedseed>1642372168</deployedseed>
    <deployedseed>1034192740</deployedseed>
    <deployedseed>2082433009</deployedseed>
    <deployedseed>1257127020</deployedseed>
    <deployedseed>740380622</deployedseed>
    <deployedseed>1037159335</deployedseed>
    <deployedseed>48425710</deployedseed>
    <deployedseed>1314583902</deployedseed>
    <deployedseed>804474963</deployedseed>
    <deployedseed>133952572</deployedseed>
    <deployedseed>1891876108</deployedseed>
    <deployedseed>1068240252</deployedseed>
    <deployedseed>1978768881</deployedseed>
    <deployedseed>8481930</deployedseed>
    <deployedseed>477398782</deployedseed>
    <deployedseed>625946847</deployedseed>
    <deployedseed>1297967070</deployedseed>
    <deployedseed>608117791</deployedseed>
    <deployedseed>1223736687</deployedseed>
    <deployedseed>1183072208</deployedseed>
    <deployedseed>512757727</deployedseed>
    <deployedseed>1518120857</deployedseed>
    <deployedseed>784331425</deployedseed>
    <deployedseed>1700686138</deployedseed>
    <deployedseed>2084292763</deployedseed>
    <deployedseed>1021930687</deployedseed>
    <deployedseed>34783297</deployedseed>
    <deployedseed>1747430972</deployedseed>
    <deployedseed>1029672525</deployedseed>
    <deployedseed>727602924</deployedseed>
    <deployedseed>1995370609</deployedseed>
    <deployedseed>1804065531</deployedseed>
    <deployedseed>2068220761</deployedseed>
    <deployedseed>1604968201</deployedseed>
    <deployedseed>561116846</deployedseed>
    <deployedseed>1808339480</deployedseed>
    <deployedseed>1706485892</deployedseed>
    <deployedseed>1230591343</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>