<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Calculate the corrected mean and standard deviation values (STACK)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A scientist is weighing each of {@N@} fish. She obtains a mean of
{@m@} g and a standard deviation of {@s@} g.
After completing the weighing, she finds that the scale was misaligned,
and always under-reported every weight by {@e@} g, i.e. a fish that
really weighed {@x@} g was reported to weigh {@x_reported@} g.
</p>

<p>
What is the mean and standard deviation after correcting for the error
in the scale?
</p>

<p>
[Hint: recall that the mean measures central tendency and the standard
deviation measures spread.]
</p>

<p>
Corrected mean: [[input:ans1]] [[validation:ans1]] g [[feedback:prt1]]
</p>

<p>
Corrected standard deviation: [[input:ans2]] [[validation:ans2]] g [[feedback:prt2]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>
The scale under-reported every weight by {@e@} g, so {@e@} g must be added to each recorded weight.
Therefore, the corrected mean is
\[
{@m@}+{@e@}={@ta1@}\text{ g}.
\]
</p>

<p>
Adding the same value to every observation shifts all the weights by the same amount but does not change their spread.
Therefore, the standard deviation remains unchanged:
\[
{@ta2@}\text{ g}.
\]
</p>

<p>
Hence, the corrected mean is {@ta1@} g and the corrected standard deviation is {@ta2@} g.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Number of fish: 30 ± 6 */
N: 30 + rand(13) - 6;

/* Mean of the recorded weights, in grams */
m: rand(31) + 30;

/* Standard deviation of the recorded weights, in grams */
s: rand(4) + 2;

/* Amount by which the scale under-reports each weight, in grams */
e: rand(4) + 1;

/* Example actual weight used to illustrate the scale error */
x: rand(10) + 40;

/* Reported value for the illustrative fish */
x_reported: x - e;

/* Correct answers */
ta1: m + e;
ta2: s;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>Systematic scale error with N = {@N@}, m = {@m@}, s = {@s@}, e = {@e@}, x = {@x@}.</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1889811620</deployedseed>
    <deployedseed>1542625021</deployedseed>
    <deployedseed>719118506</deployedseed>
    <deployedseed>1374999696</deployedseed>
    <deployedseed>1431072476</deployedseed>
    <deployedseed>178506651</deployedseed>
    <deployedseed>1685233594</deployedseed>
    <deployedseed>1692940304</deployedseed>
    <deployedseed>1432753297</deployedseed>
    <deployedseed>1133868943</deployedseed>
    <deployedseed>1553071731</deployedseed>
    <deployedseed>478548969</deployedseed>
    <deployedseed>836506817</deployedseed>
    <deployedseed>1659730050</deployedseed>
    <deployedseed>936865597</deployedseed>
    <deployedseed>1821116419</deployedseed>
    <deployedseed>1198389768</deployedseed>
    <deployedseed>1148857414</deployedseed>
    <deployedseed>1525234148</deployedseed>
    <deployedseed>488353695</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>