<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Climate Data</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The data for this question are monthly total sunshine hours for 30 years from a single station and from satellite estimates close to that station. There are four columns in the data: <code>year</code>, <code>month</code>, <code>sunh</code> and <code>sunh_sat</code>.</p>
<p><code>sunh</code> are the station values and <code>sunh_sat</code> are the satellite values.</p>
<p>Click the button to download the data in CSV format. Then import the data into R-Instat (or your favourite statistics package).</p>
<button onclick="exportToCsv()">Download Data</button>
<script>
var Results= {#data_set_wheader#};

exportToCsv = function(file_name) {
  var CsvString = "";
  Results.forEach(function(RowItem, RowIndex) {
    RowItem.forEach(function(ColItem, ColIndex) {
             CsvString += ColItem;
             if( ColIndex < RowItem.length - 1){
                 CsvString += ',';
                } 
    });
    CsvString += "\r\n";
  });
  CsvString = "data:application/csv," + encodeURIComponent(CsvString);
 var x = document.createElement("A");
 x.setAttribute("href", CsvString );
 x.setAttribute("download","sunshine_month.csv");
 document.body.appendChild(x);
 x.click();
}
</script>
<p>Calculate the mean error (bias) for each month.</p>
<p>1. What is the mean error in <code>{@month_ran_text@}</code>?</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>
<p>2. In which month does the mean bias indicate that the satellite values are giving the closest estimates to the station data on average?</p>
<p>[[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[first_year:1983;
num_years:30;
row_1:["year", "month", "sunh", "sunh_sat"];
months_text: ["January", "February", "March", "April", "May", "June", "July", "August", "September", "October", "November", "December"];
month_ran: rand(12)+1;
month_ran_text: months_text[month_ran];

month_means: [45.8, 66.0, 136.0, 205.1, 273.6, 256.8, 271.2, 239.5, 172.4, 119.6, 55.2, 38.5];
month_sds: [17.6, 19.4, 34.1, 54.7, 55.2, 51.2, 56.3, 43.0, 42.1, 25.2, 20.1, 13.6] * 0.5;
diff_means: [0.3596, -2.1915, -8.0770, -0.6041, -13.7988, -10.5792, -9.5156, -5.7240, -3.6311, -2.7347, -0.4352, 2.7373];
diff_sds: [6.87, 8.78, 18.43, 12.49, 13.25, 12.64, 8.54, 10.58, 9.86, 8.56, 6.33, 7.22];
month_data: makelist(makelist([y, m, float(quantile_normal(rand(1.0),month_means[m],month_sds[m])), diff_means[m], diff_sds[m]], m, 1, 12), y, first_year, first_year+num_years-1);
data_set:apply(append, month_data);
data_set:makelist(endcons(data_set[i][3] + float(quantile_normal(rand(1.0),data_set[i][4],data_set[i][5])), rest(data_set[i], -2)), i, 1, 30*12);
data_set_wheader:append([row_1],data_set);

station_mean: [0,0,0,0,0,0,0,0,0,0,0,0];
satellite_mean: [0,0,0,0,0,0,0,0,0,0,0,0];

for m: 1 thru 12 do /*
  */ for r: 1 thru length(data_set) do /*
         */ (if data_set[r][2] = m then station_mean[m]: station_mean[m] + data_set[r][3], if data_set[r][2] = m then satellite_mean[m]: satellite_mean[m] + data_set[r][4])

bias: makelist(satellite_mean[i]-station_mean[i], i, 1, 12);
bias: bias/num_years;
ta1: bias[month_ran];
ta1: float(ta1);
ta1_dp: dispdp(ta1,2);
abs_bias: abs(bias);
min_bias: lmin(abs_bias);
min_bias_ind: sublist_indices(bias, lambda([x1], abs(x1) = min_bias))[1];
ta2_correct: months_text[min_bias_ind];
func(x) := is(x = ta2_correct);
ta2: maplist(func, months_text);
ta2:zip_with("[",months_text, ta2);

max_bias: lmax(abs_bias);
max_bias_ind: sublist_indices(bias, lambda([x1], abs(x1) = max_bias))[1];
ta2_wrong: months_text[max_bias_ind];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1_dp</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>-ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Not quite! Your answer has the right value but the wrong sign, so you did the right calculation but probably with the columns the wrong way round.</p>
<p>Normally, mean error is defined as: <code>mean(estimates - observations)</code>.</p>
<p>So you should have got {@ta1_dp@} as your answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2_correct</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2_wrong</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You might need to check how to interpret the mean error. You chose the month where the mean bias was furthest from {@0@}, so this is the month where the satellite values are furthest from the station values. When the satellite values are closest to the station values, the mean error is close to {@0@}.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1052727116</deployedseed>
  </question>

</quiz>