<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Counting Pair Combinations for Mixed Doubles Tennis Tournament Selection</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A team of \({@pairs@}\) pairs is to be selected from \({@men@}\) men and \({@women@}\) women, each pair consisting of man and woman, for a tennis tournament. In how many ways can this be done? </p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the number of ways we multiply the individual counts.</p>
<p>\(^{@men@}C_{@pairs@} \times  ^{@women@}C_{@pairs@} \times {@pairs@}!\)</p>
<p>\(^{@men@}C_{@pairs@}\) represents the number of ways to choose \({@pairs@}\) items out of \({@men@}\) available items</p>
<p>\(^{@women@}C_{@pairs@}\) represents the number of ways to choose \({@pairs@}\) items out of \({@women@}\) available items.</p>
<p>\({@pairs@}!\) represents the number of ways the \({@pairs@}\) selected women can be paired up with the \({@pairs@}\) selected men.</p>
<p>\(^{@men@}C_{@pairs@} \times  ^{@women@}C_{@pairs@} \times {@pairs@}!={@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*prob*/
pairs:rand(4)+2;
men:pairs+1+rand(3);
women:pairs+1+rand(3);

/*worked solution variables*/
sol1:binomial(men,pairs);
sol2:binomial(women,pairs);
sol3:sol1*sol2*factorial(pairs);

mistake1: sol1*sol2;

ta:sol3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(^{@men@}C_{@pairs@} \times  ^{@women@}C_{@pairs@} \times {@pairs@}={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary</p>
<p id="concept">https://docs.google.com/document/d/1F6nZWnrWsFo4LYyD2W4wmGSeiLTeJX8bHCVVIlcV2WQ/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question involves finding the number of ways that a given number of choosing a team member from a team of men and women.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>forgot factorial</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer is almost correct. You forgot to account for the number of ways that the selected men and women can be paired up.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1874115254</deployedseed>
    <deployedseed>1730660230</deployedseed>
    <deployedseed>1359840916</deployedseed>
    <deployedseed>653868574</deployedseed>
    <deployedseed>3958876</deployedseed>
    <deployedseed>793813183</deployedseed>
    <deployedseed>579441191</deployedseed>
    <deployedseed>460305280</deployedseed>
    <deployedseed>1043819341</deployedseed>
    <deployedseed>168003304</deployedseed>
    <deployedseed>741499816</deployedseed>
    <deployedseed>1772967351</deployedseed>
    <deployedseed>369267363</deployedseed>
    <deployedseed>710683362</deployedseed>
    <deployedseed>1997484134</deployedseed>
    <deployedseed>420846561</deployedseed>
    <deployedseed>2145174133</deployedseed>
    <deployedseed>1966033292</deployedseed>
    <deployedseed>1361122255</deployedseed>
    <deployedseed>433104146</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>