<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Counting Words with Specific Letter Constraints</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>1. Conside the word {@word1@}.</p>
<p>Assuming that any arrangement of the letters in {@word1@} forms a word, how many different words can we make with the letters of {@word1@}?</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>2. Consider the word {@word2@}.</p>
<p>Assuming that any arrangement of the letters in {@word2@} forms a word, how many different words can we make with the letters of {@word2@}?</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>3. Consider the word {@word3@}.</p>
<p>Assuming that any arrangement of the letters in {@word3@} forms a word, how many different words can we make with the letters of {@word3@}?</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>\(1.\)</p>
<p>The word {@word1@} has {@wordlength1@} letters, all of which are different. Therefore, the number of words we can make is the number of possible arrangements (or permutations) of the letters of the word {@word1@}. This is given by \({@wordlength1@}! = {@wordlength1!@}\).</p>
<br>
<p>\(2.\)</p>
<p>The word {@word2@} has {@wordlength2@} letters. Therefore, the number of words we can make is the number of possible arrangements (or permutations) of the letters of the word {@word2@}. This is given by \({@wordlength2@}! = {@wordlength2!@}\).</p>
<p>However, the letter {@letter2@} appears twice in the word {@word2@} and many options in the number above will be repeated. As {@letter2@} appears twice, we consider the number of ways in which these two letters could be arranged. We have 2 options for the first one and 1 option for the second one. So there are \(2\cdot 1 = 2! = 2\) ways to arrange the two letters {@letter2@}.</p>
<p>Therefore, the number of words, without repetition, is given by \(\displaystyle \frac{{@wordlength2@}!}{2!} = \frac{{@wordlength2!@}}{2} = {@ta2@}\).</p></p>
<br>
<p>\(3.\)</p>
<p>The word {@word3@} has {@wordlength3@} letters. Therefore, the number of words we can make is the number of possible arrangements (or permutations) of the letters of the word {@word3@}. This is given by \({@wordlength3@}! = {@wordlength3!@}\).</p>
<p>However, the letter {@letter3@} appears three times in the word {@word3@} and many options in the number above will be repeated. As {@letter3@} appears three times, we consider the number of ways in which these three letters could be arranged. We have 3 options for the first one, 2 options for the second one and 1 option for the third one. So there are \(3\cdot 2\cdot 1 = 3! = 6\) ways to arrange the three letters {@letter3@}.</p>
<p>Therefore, the number of words, without repetition, is given by \(\displaystyle \frac{{@wordlength3@}!}{3!} = \frac{{@wordlength3!@}}{6} = {@ta3@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[wordlength1:rand([6,7,8]);
word1:if wordlength1 = 6 then rand([friday,family,silver,monday]) elseif wordlength1=7 then rand([departs,country,history,kitchen]) else rand([thousand,children,equation,champion]);
ta1:if wordlength1 = 6 then 6! elseif wordlength1 = 7 then 7! else 8!;

wordlength2:rand([6,7,8]);
word2:if wordlength2 = 6 then rand([twelve,thirty,future,africa]) elseif wordlength2=7 then rand([perfect,america,freedom,seventy]) else rand([fourteen,american,mountain,strength]);
letter2:if word2 = fourteen then "e" elseif word2 = american then "a" elseif word2 = mountain then "n" elseif word2 = strength then "t" elseif word2 = perfect then "e" elseif word2 = america then "a" elseif word2 = freedom then "e" elseif word2 = seventy then "e" elseif word2 = twelve then "e" elseif word2 = thirty then "t" elseif word2 = future then "u" else "a";
ta2:if wordlength2 = 6 then 6!/2 elseif wordlength2 = 7 then 7!/2 else 8!/2;

wordlength3:rand([6,7,8]);
word3:if wordlength3 = 6 then rand([eleven,canada,bubble,mirror]) elseif wordlength3=7 then rand([warrior,weekend,between,goggles]) else rand([business,syllable,virginia,envelope]);
letter3:if word3 = eleven then "e" elseif word3 = canada then "a" elseif word3 = bubble then "b" elseif word3 = mirror then "r" elseif word3 = warrior then "r" elseif word3 = weeked then "e" elseif word3 = between then "e" elseif word3 = goggles then "g" elseif word3 = business then "s" elseif word3 = syllable then "l" elseif word3 = virginia then "i" else "e";
ta3:if wordlength3 = 6 then 6!/6 elseif wordlength3 = 7 then 7!/6 else 8!/6;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>1. Conside the word {@word1@}={@ta1@}</p>
<p>2. Consider the word {@word2@}={@ta2@}</p>
<p>3. Consider the word {@word3@}={@ta3@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayun</p>
<p id="concept">NA</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">This question requires the student to explore the possibilities of forming different words from given sets of letters in the given words. Students are tasked with determining the total number of distinct words that can be created by rearranging the letters in each word.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer ta1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if the answer is 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer ta2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if the answer is 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if the answer is 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>871504297</deployedseed>
    <deployedseed>1228822263</deployedseed>
    <deployedseed>1801399087</deployedseed>
    <deployedseed>1628948374</deployedseed>
    <deployedseed>1451564185</deployedseed>
    <deployedseed>344196696</deployedseed>
    <deployedseed>651093488</deployedseed>
    <deployedseed>341794772</deployedseed>
    <deployedseed>1218630652</deployedseed>
    <deployedseed>1194214044</deployedseed>
    <deployedseed>928915434</deployedseed>
    <deployedseed>1343255008</deployedseed>
    <deployedseed>850812045</deployedseed>
    <deployedseed>2006426789</deployedseed>
    <deployedseed>1864518263</deployedseed>
    <deployedseed>516011511</deployedseed>
    <deployedseed>8136125</deployedseed>
    <deployedseed>1837701023</deployedseed>
    <deployedseed>1638248379</deployedseed>
    <deployedseed>1764779937</deployedseed>
    <deployedseed>246344646</deployedseed>
    <deployedseed>780197701</deployedseed>
    <deployedseed>141140527</deployedseed>
    <deployedseed>1748344167</deployedseed>
    <deployedseed>71217724</deployedseed>
    <deployedseed>1764767720</deployedseed>
    <deployedseed>1874714532</deployedseed>
    <deployedseed>998444436</deployedseed>
    <deployedseed>1469268986</deployedseed>
    <deployedseed>1671321224</deployedseed>
    <deployedseed>922240545</deployedseed>
    <deployedseed>254898194</deployedseed>
    <deployedseed>1773555826</deployedseed>
    <deployedseed>2107210850</deployedseed>
    <deployedseed>891085159</deployedseed>
    <deployedseed>1208128539</deployedseed>
    <deployedseed>1372961272</deployedseed>
    <deployedseed>280023534</deployedseed>
    <deployedseed>954886710</deployedseed>
    <deployedseed>488127246</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>