<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Determining general fertility and total fertility rate for women</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>From the figures given below calculate the female gross reproduction rate if the ratio of male and female children is \({@male_ratio@}:{@female_ratio@}\)</p>
<p>
<table border="solid">
  <tr>
    <th width="30%", style="text-align:center">Age Group</th>
    <th width="80%", style="text-align:center">Number of children born to {@total_women@} women</th>

  </tr>
  <tr>
    <td style="text-align:center">\({@lower_age_groups[1]@}-{@upper_age_groups[1]@}\)</td>
    <td style="text-align:center">\({@count_1@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[2]@}-{@upper_age_groups[2]@}\)</td>
    <td style="text-align:center">\({@count_2@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[3]@}-{@upper_age_groups[3]@}\)</td>
    <td style="text-align:center">\({@count_3@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[4]@}-{@upper_age_groups[4]@}\)</td>
    <td style="text-align:center">\({@count_4@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[5]@}-{@upper_age_groups[5]@}\)</td>
    <td style="text-align:center">\({@count_5@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[6]@}-{@upper_age_groups[6]@}\)</td>
    <td style="text-align:center">\({@count_6@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[7]@}-{@upper_age_groups[7]@}\)</td>
    <td style="text-align:center">\({@count_7@}\)</td>


  </tr>
 <tr>
    <td style="text-align:center">\(\sum\)</td>
    <td style="text-align:center">\({@total_children@}\)</td>


  </tr>
</table>
</p>
<p><strong>Note:</strong> It is assumed that the number of children born (given) are the total children born during the time {@total_women@}  women passed through the age group, i.e in years)</p>
<p><strong>Note:</strong> Give your answer correct to \(4\) decimal places</p>
<p>Find the Gross Reproduction Rate \(\textbf{(GRR)}\)  </p>
<p>\(\textbf{GRR}= \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Calculation of Gross Reproduction Rate \(\textbf{(GRR)}\)</p>
<p>From the figures given below we calculate the <b>female gross reproduction rate</b> given the ratio of male and female children as \({@male_ratio@}:{@female_ratio@}\)</p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Age Group</th>
    <th width="15%", style="text-align:center">Number of children born to {@total_women@} women</th>
    <th width="15%", style="text-align:center">Number of female children born</th>
  </tr>
  <tr>
    <td style="text-align:center">\({@lower_age_groups[1]@}-{@upper_age_groups[1]@}\)</td>
    <td style="text-align:center">\({@count_1@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_1@}\times {@female_ratio@}}{100}={@female_children_count[1]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[2]@}-{@upper_age_groups[2]@}\)</td>
    <td style="text-align:center">\({@count_2@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_2@}\times {@female_ratio@}}{100}={@female_children_count[2]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[3]@}-{@upper_age_groups[3]@}\)</td>
    <td style="text-align:center">\({@count_3@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_3@}\times {@female_ratio@}}{100}={@female_children_count[3]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[4]@}-{@upper_age_groups[4]@}\)</td>
    <td style="text-align:center">\({@count_4@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_4@}\times {@female_ratio@}}{100}={@female_children_count[4]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[5]@}-{@upper_age_groups[5]@}\)</td>
    <td style="text-align:center">\({@count_5@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_5@}\times {@female_ratio@}}{100}={@female_children_count[5]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[6]@}-{@upper_age_groups[6]@}\)</td>
    <td style="text-align:center">\({@count_6@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_6@}\times {@female_ratio@}}{100}={@female_children_count[6]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\({@lower_age_groups[7]@}-{@upper_age_groups[7]@}\)</td>
    <td style="text-align:center">\({@count_7@}\)</td>
    <td style="text-align:center">\(\displaystyle \frac{{@count_7@}\times {@female_ratio@}}{100}={@female_children_count[7]@}\)</td>

  </tr>
 <tr>
    <td style="text-align:center">\(\sum\)</td>
    <td style="text-align:center">\({@total_children@}\)</td>
    <td style="text-align:center">\({@total_female_children_count@}\)</td>

  </tr>
</table>
</p>
<br>
<p>Finding the Gross Reproduction Rate \(\textbf{(GRR)}\)  </p>
<p> \(\textbf{GRR}=\displaystyle \frac{\text{Total births to}\,  {@total_women@} \, \text{Women}}{{@total_women@}}=\frac{{@total_female_children_count@}}{{@total_women@}} \approx {@grr@} \) </p>
<p><strong>Note:</strong> It is assumed that the number of children born (given) are the total children born during the time {@total_women@}  women passed through the age group, i.e in years)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*age group intervals*/ 
age:5; 
age_group1:rand_with_step(14,16,1); 
age_group2:age_group1+age; 
age_group3:age_group2+age; 
age_group4:age_group3+age; 
age_group5:age_group4+age; 
age_group6:age_group5+age;
age_group7:age_group6+age; 
lower_age_groups:[age_group1,age_group2,age_group3,age_group4,age_group5,age_group6,age_group7];
upper_age_groups:lower_age_groups+(age-1);
percentage:100;

male_ratio:rand_with_step(40,60,1); /*randomly generate this ratios*/
female_ratio:percentage-male_ratio;

total_women:1000;

/* Define the number of children born in each age group */ 
children_counts:makelist((rand_with_step(10,500,10)),i,1,7);
count_1: children_counts[1]; 
count_2: children_counts[2]; 
count_3: children_counts[3]; 
count_4: children_counts[4]; 
count_5: children_counts[5]; 
count_6: children_counts[6]; 
count_7: children_counts[7]; 

total_children: sum(children_counts[i],i,1,length(children_counts));

female_children_count:decimalplaces(children_counts*(female_ratio/percentage),4);
/*female_children_counts:decimalplaces(female_children_countz,2); 
female_children_count1:round(female_children_countz); 
female_children_counts_1:is(abs((female_children_counts-female_children_count1))=0.0);
female_children_count: if female_children_counts_1=true then female_children_count1 else female_children_counts;*/
total_female_children_count:sum(female_children_count[i],i,1,length(female_children_count));

/* Calculate the female gross reproduction rate (GRR) */
grr:decimalplaces(total_female_children_count / total_women,4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(\textbf{GRR}=\displaystyle \frac{{@total_female_children_count@}}{{@total_women@}} \approx {@grr@} \) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/1n30yhK0rU91G1jMt-Hy2YTu1Gn4El0y2ih8-MRGs3jE/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">Given a table with age group,number of women and  Age specific fertility rate  (per 1000), the question requires you to find General fertility rate and Total fertility Rate for n women.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>grr</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer for GRR </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>grr</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done for giving your answer correct to \(4\) decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer for GRR  to 4dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>total_female_children_count / total_women</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done for giving your answer in exact form</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for an error margin of 0.00001 to the sa</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>grr</tans>
        <testoptions>0.00001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1620449862</deployedseed>
    <deployedseed>1775967909</deployedseed>
    <deployedseed>1363120725</deployedseed>
    <deployedseed>1248135504</deployedseed>
    <deployedseed>1270718576</deployedseed>
    <deployedseed>382499623</deployedseed>
    <deployedseed>973018585</deployedseed>
    <deployedseed>120201689</deployedseed>
    <deployedseed>335796161</deployedseed>
    <deployedseed>1253948950</deployedseed>
    <deployedseed>499037337</deployedseed>
    <deployedseed>2126963698</deployedseed>
    <deployedseed>1776322144</deployedseed>
    <deployedseed>1267141851</deployedseed>
    <deployedseed>1048774902</deployedseed>
    <deployedseed>1652258526</deployedseed>
    <deployedseed>1376433220</deployedseed>
    <deployedseed>1449692201</deployedseed>
    <deployedseed>314208793</deployedseed>
    <deployedseed>1366948100</deployedseed>
    <deployedseed>964386449</deployedseed>
    <deployedseed>2030809906</deployedseed>
    <deployedseed>1473632925</deployedseed>
    <deployedseed>1531379508</deployedseed>
    <deployedseed>1833004933</deployedseed>
    <deployedseed>460974929</deployedseed>
    <deployedseed>1394720159</deployedseed>
    <deployedseed>1023784190</deployedseed>
    <deployedseed>1738660545</deployedseed>
    <deployedseed>1947482790</deployedseed>
    <deployedseed>1131438476</deployedseed>
    <deployedseed>1721060006</deployedseed>
    <deployedseed>1380096308</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>grr</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>