<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Estimating the combined mean and standard deviation for two datasets presented as dotplots</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The dot plots below show the weights of samples of male and female birds of the same species. Can you estimate the mean and standard deviation if the male and female birds were combined?

<p>
[[jsxgraph]] 
// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true
});

//Create points of Data Set 2
var lx2 = {#dorx2#};
var ly2 = {#dory2#};
var i = 0;
for (i = 0; i <= lx2.length - 1; i++) {
  p[i] = board.create('point', [lx2[i], ly2[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 2
var txt = board.create('text', [{#xmin#}+1, {#maxfreq2#}/2+2, "Male"], {
  fontSize: 30,fixed: true
});

//Create points of Data Set 1
var lx1 = {#dorx1#};
var ly1 = {#dory1#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 1
var txt = board.create('text', [{#xmin#}+1, {#maxfreq1#}/2+{#maxfreq2#} + 5, "Female"], {
  fontSize: 30,fixed: true
});
[[/jsxgraph]]</p>


<p>Estimated combined mean of the male and female birds? [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p>Estimated combined standard deviation  of  male and female birds? [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>To estimate the combined mean, imagine merging the two dotplots into a single distribution. Look for a plausible value for the mean by considering how far the observations lie above and below it. Observations further from the mean should have more influence on your estimate than observations close to it.</p>

<p>To estimate the combined standard deviation, consider the distances of all the observations from your estimated combined mean. Look for a value that represents a typical spread around the mean, remembering that observations far from the mean have a greater effect on the standard deviation.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Generate the means and standard deviations for the two datasets.
   The second mean is shifted from the first, while its standard
   deviation is a multiple of the first. */
mean1:rand(101)+100;
aa:rand([-1,0,1]);
mean2:mean1+(rand(11)+5)*aa;
sd1:rand(5)+5;
bb:rand([1/3,1/2,1,2,3]);
sd2:sd1*bb;

/* Generate two datasets of equal size. */
ds1:rand(70)+30;
ds2:ds1;

/* Generate normally distributed observations, rounded to integers,
   and sort them for construction of the dotplots. */
dun1:makelist(round(float(quantile_normal(rand(1.0),mean1,sd1))),i,1,ds1);
dorx1:sort(dun1);
dun2:makelist(round(float(quantile_normal(rand(1.0),mean2,sd2))),i,1,ds2);
dorx2:sort(dun2);

/* Find the maximum frequency of any value in each dataset.
   These are used to determine the vertical spacing of the dotplots. */
maxfreq1:0;
k:1;
for i:2 thru ds1 do
    if dorx1[i]=dorx1[i-1]
    then (k:k+1, if maxfreq1 < k then maxfreq1:k)
    else k:1;

maxfreq2:0;
k:1;
for i:2 thru ds2 do
    if dorx2[i]=dorx2[i-1]
    then (k:k+1, if maxfreq2 < k then maxfreq2:k)
    else k:1;

/* Set common plotting limits so that both dotplots use the same scale. */
xmin:if dorx1[1]<=dorx2[1] then dorx1[1]-2 else dorx2[1]-2;
xmax:if dorx1[ds1]>=dorx2[ds2] then dorx1[ds1]+2 else dorx2[ds2]+2;
ymin:-4;
ymax:maxfreq1+maxfreq2+5;

/* Generate vertical coordinates for the first dotplot.
   It is placed above the second dotplot with sufficient separation. */
j:1;
dory1:makelist(0,i,1,ds1);
dory1[1]:1+maxfreq2+3;
for i:2 thru ds1 do
    if dorx1[i]=dorx1[i-1]
    then (j:j+1, dory1[i]:j+maxfreq2+3)
    else (j:1, dory1[i]:j+maxfreq2+3);

/* Generate vertical coordinates for the second dotplot. */
j:1;
dory2:makelist(0,i,1,ds2);
dory2[1]:1;
for i:2 thru ds2 do
    if dorx2[i]=dorx2[i-1]
    then (j:j+1, dory2[i]:j)
    else (j:1, dory2[i]:j);

/* Combine the two datasets and calculate the teacher answers. */
combined:append(dorx1,dorx2);
ta1:float(mean(combined));
ta2:float(std(combined));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>sd1 = {@sd1@}, sd2 = {@sd2@}. mean1 = {@mean1@}, mean2 = {@mean2@}. ta1= {@ta1@}, ta2 = {@ta2@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, your estimate is within 10% of the combined mean.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.25</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Good estimate. You are within 25% of the combined mean. See if you can refine your estimate.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your estimate is not close enough. Look again at the values in both groups and consider how all the observations contribute to the combined mean.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Well done, your estimate is within 10% of the combined standard deviation.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumRelative</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.25</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Good estimate. You are within 25% of the combined standard deviation. See if you can refine your estimate.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your estimate is not close enough. Consider how far observations in both groups lie from the combined mean.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1930926414</deployedseed>
    <deployedseed>1851519359</deployedseed>
    <deployedseed>841796194</deployedseed>
    <deployedseed>2004887973</deployedseed>
    <deployedseed>711299222</deployedseed>
    <deployedseed>110287764</deployedseed>
    <deployedseed>417007919</deployedseed>
    <deployedseed>1753100980</deployedseed>
    <deployedseed>1926537662</deployedseed>
    <deployedseed>459257524</deployedseed>
    <deployedseed>1497828040</deployedseed>
    <deployedseed>507761624</deployedseed>
    <deployedseed>708146209</deployedseed>
    <deployedseed>941754810</deployedseed>
    <deployedseed>1403654721</deployedseed>
    <deployedseed>2084484821</deployedseed>
    <deployedseed>1646495901</deployedseed>
    <deployedseed>694635780</deployedseed>
    <deployedseed>1538028623</deployedseed>
    <deployedseed>1631652458</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>