<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>FERTILITY (Factors and Determinants)_q3</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>What does Maternal Mortality Rate mean? Select all that apply:</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>This is death that occurs to expectant mothers, because of complications in pregnancy</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p> Death to a mother can also happen while giving birth as a result of complications she had during pregnancy.
</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Death to a mother can also happen while giving birth as a result of infections she got during pregnancy</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a:rand(3)+1;
correct_answer:rand_selection([
[1, true, " The risk of dying as a result of complications of pregnancy.
"],
[2, true, "The risk of dying as a result of complications of child birth."],
[3, true, "The risk of a woman’s death as a result of infections during pregnancy"]],a);

b:4-a;
incorrect_answers:rand_selection([
[4, false, "Dying while giving birth
"],
[5, false, " A woman’s death during pregnancy
"],
[6, false, "The risk of a woman dying per 1000 women"],
[7, false, "Ability of a woman to die while giving birth"]], b);

ta:random_permutation(append(correct_answer,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>What does Maternal Mortality Rate mean?<p/>
<p> {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat<p>
<p id="concept">https://docs.google.com/document/d/1A8eJg87rqt5t25oRCycjEunKtesHkXcWO9QFM2BU1l4/edit</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Motognon Wastalas</p>
<p id="description">This question requires the student to grasp the concept of Maternal Mortality Rate.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1332696332</deployedseed>
    <deployedseed>1834284750</deployedseed>
    <deployedseed>1173818126</deployedseed>
    <deployedseed>652192032</deployedseed>
    <deployedseed>1900730405</deployedseed>
    <deployedseed>335417982</deployedseed>
    <deployedseed>785039154</deployedseed>
    <deployedseed>1631552070</deployedseed>
    <deployedseed>1186495663</deployedseed>
    <deployedseed>1705700301</deployedseed>
    <deployedseed>2013468497</deployedseed>
    <deployedseed>317095047</deployedseed>
    <deployedseed>615607131</deployedseed>
    <deployedseed>1477410302</deployedseed>
    <deployedseed>2142041351</deployedseed>
    <deployedseed>947860721</deployedseed>
    <deployedseed>2008714625</deployedseed>
    <deployedseed>1034560332</deployedseed>
    <deployedseed>1967432734</deployedseed>
    <deployedseed>1409107824</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>