<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Finding the total number of different combinations to form a group</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A work group for a project is composed of {@men@} men and {@women@} women. They have to give a presentation to their employers on the progress of the project and only {@numtot@} team members can take part in the presentation.</p>
<p>How many different groups can be made for the presentation if:</p>
<p>1. Exactly 1 man can be in the group?</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>2. Exactly 2 men can be in the group?</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>3. At most 3 men can be in the group?</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
<p>4. At least 2 men and 2 women have to be in the group?</p>
<p>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>men:rand(5)+6;
women:rand(5)+6;
total:men+women;
nummen:rand(2)+3;
numtot:nummen+2+rand(2);
ta1:binomial(men,1)*binomial(women,numtot-1);
ta2:binomial(men,2)*binomial(women,numtot-2);
ta3:binomial(women,numtot)+binomial(men,1)*binomial(women,numtot-1)+binomial(men,2)*binomial(women,numtot-2)+binomial(men,3)*binomial(women,numtot-3);
ta4:binomial(women,2)*binomial(men,2)*binomial(total-4,numtot-4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Men: {@men@}, Women: {@women@}, Total: {@total@}, Subgroup: {@numtot@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Worked Solution:</p>
<p>We need exactly one man in a team of {@numtot@} people. We have to choose 1 man from a possible {@men@} men in the working group and the remaining {@numtot-1@} members of the presentation team have to be women from a possible total of {@women@} women in the working group.</p>
<p>The number of possibilities for the man is \(^{{@men@}}C_{1} = {@binomial(men,1)@}\).</p>
<p>The number of possibilities for the women is \(^{{@women@}}C_{{@numtot-1@}} = {@binomial(women,numtot-1)@}\).</p>
<p>Therefore, the total number of different teams is \(^{{@men@}}C_{1} \cdot ^{{@women@}}C_{{@numtot-1@}} = {@binomial(men,1)@} \cdot {@binomial(women,numtot-1)@} = {@ta1@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Worked Solution:</p>
<p>We need exactly two men in a team of {@numtot@} people. We have to choose 2 men from a possible {@men@} men in the working group and the remaining {@numtot-2@} members of the presentation team have to be women from a possible total of {@women@} women in the working group.</p>
<p>The number of possibilities for the men is \(^{{@men@}}C_{2} = {@binomial(men,2)@}\).</p>
<p>The number of possibilities for the women is \(^{{@women@}}C_{{@numtot-2@}} = {@binomial(women,numtot-2)@}\).</p>
<p>Therefore, the total number of different teams is \(^{{@men@}}C_{2} \cdot ^{{@women@}}C_{{@numtot-2@}} = {@binomial(men,2)@} \cdot {@binomial(women,numtot-2)@} = {@ta2@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Worked Solution version 1:</p>
<p>In this case, we have three different options for the team:</p>
<p>1. The team has no men at all and is made up of only women. We need to choose {@numtot@} women from a total of {@women@} to choose from. This can be done in \(^{{@women@}}C_{{@numtot@}} = {@binomial(women,numtot)@}\) ways.</p>
<p>2. The team has exactly one man. This is the same as part 1 and the number of possibilities is {@ta1@}.</p>
<p>3. The team has exactly two men. This is the same as part 2 and the number of possibilities is {@ta2@}.</p>
<p>4. The team has exactly three men. We calculate the number of possibilities in the same way as before:</p>
<p>We need exactly three men in a team of {@numtot@} people. We have to choose 3 men from a possible {@men@} men in the working group and the remaining {@numtot-2@} members of the presentation team have to be women from a possible total of {@women@} women in the working group.</p>
<p>The number of possibilities for the men is \(^{{@men@}}C_{3} = {@binomial(men,3)@}\).</p>
<p>The number of possibilities for the women is \(^{{@women@}}C_{{@numtot-3@}} = {@binomial(women,numtot-3)@}\).</p>
<p>Therefore, the total number of different teams is \(^{{@men@}}C_{3} \cdot ^{{@women@}}C_{{@numtot-3@}} = {@binomial(men,3)@} \cdot {@binomial(women,numtot-3)@} = {@binomial(men,3)*binomial(women,numtot-3)@}\).</p>
<p>As the four options don't overlap, we now add the three results and get a total of \({@binomial(women,numtot)@} + {@ta1@}+{@ta2@}+{@binomial(men,3)*binomial(women,numtot-3)@} = {@ta3@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Worked Solution:</p>
<p>We need at least two women in the presentation team and we have {@women@} to choose from. The number of ways in which we can do so is \(^{{@women@}}C_{2} = {@binomial(women,2)@}\).</p>
<p>Similarly, we need at least two men in the presentation team and we have {@men@} to choose from. The number of ways in which we can do so is \(^{{@men@}}C_{2} = {@binomial(men,2)@}\).</p>
<p>The gender of the remaining {@numtot-4@} members of the presentation team is indifferent and we have {@women - 2@} women and {@men - 2@} men to choose from. So, we have to choose {@numtot-4@} from a possible {@total-4@} people, which can be done in \(^{{@total-4@}}C_{{@numtot-4@}} = {@binomial(total-4,numtot-4)@}\) ways.</p>
<p>Therefore, the total number of ways in which the presentation team can be formed is \(^{{@women@}}C_{2} \cdot ^{{@men@}}C_{2} \cdot ^{{@total-4@}}C_{{@numtot-4@}} = {@binomial(women,2)@} \cdot {@binomial(men,2)@} \cdot {@binomial(total-4,numtot-4)@} = {@binomial(women,2)*binomial(men,2)*binomial(total-4,numtot-4)@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1427417649</deployedseed>
    <deployedseed>1720638024</deployedseed>
    <deployedseed>68057378</deployedseed>
    <deployedseed>1641111400</deployedseed>
    <deployedseed>1532750021</deployedseed>
    <deployedseed>828024878</deployedseed>
    <deployedseed>1464616844</deployedseed>
    <deployedseed>23857013</deployedseed>
    <deployedseed>773748948</deployedseed>
    <deployedseed>1487342865</deployedseed>
    <deployedseed>1519310880</deployedseed>
    <deployedseed>1687679289</deployedseed>
    <deployedseed>1148717400</deployedseed>
    <deployedseed>1818054341</deployedseed>
    <deployedseed>1876033549</deployedseed>
    <deployedseed>2027880997</deployedseed>
    <deployedseed>158993201</deployedseed>
    <deployedseed>1287783048</deployedseed>
    <deployedseed>1211606934</deployedseed>
    <deployedseed>987117710</deployedseed>
    <deployedseed>1715318572</deployedseed>
    <deployedseed>2062430952</deployedseed>
    <deployedseed>103368343</deployedseed>
    <deployedseed>1178377826</deployedseed>
    <deployedseed>1471310730</deployedseed>
    <deployedseed>733770855</deployedseed>
    <deployedseed>1719638112</deployedseed>
    <deployedseed>392213853</deployedseed>
    <deployedseed>9948618</deployedseed>
    <deployedseed>1423214395</deployedseed>
    <deployedseed>299064809</deployedseed>
    <deployedseed>319676430</deployedseed>
  </question>

</quiz>