<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Find percentiles from frequency table</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Find \(P_{@p1@}\) and \(P_{@p2@}\) for the following frequency distribution, where
\(P_n\) denotes the \(n^{th}\) percentile.
</p>

<p>
First complete the cumulative frequency column.
</p>

<table border="solid">
  <tr>
    <th width="30%" style="text-align:center">Value</th>
    <th width="30%" style="text-align:center">Frequency</th>
    <th width="40%" style="text-align:center">Cumulative frequency</th>
  </tr>

  <tr>
    <td style="text-align:center">\({@x1@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf1@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x2@}\)</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x3@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x4@}\)</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x5@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x6@}\)</td>
    <td style="text-align:center">\({@f6@}\)</td>
    <td style="text-align:center">[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</td>
  </tr>
</table>

<p>
Using the completed cumulative frequency distribution, find:
</p>

<p>
\(P_{@p1@}=\) [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]
</p>

<p>
\(P_{@p2@}=\) [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>

<p>
First calculate the cumulative frequencies by successively adding the frequencies:
</p>

<table border="solid">
  <tr>
    <th width="30%" style="text-align:center">Value</th>
    <th width="30%" style="text-align:center">Frequency</th>
    <th width="40%" style="text-align:center">Cumulative frequency</th>
  </tr>

  <tr>
    <td style="text-align:center">\({@x1@}\)</td>
    <td style="text-align:center">\({@f1@}\)</td>
    <td style="text-align:center">\({@cf1@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x2@}\)</td>
    <td style="text-align:center">\({@f2@}\)</td>
    <td style="text-align:center">\({@cf2@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x3@}\)</td>
    <td style="text-align:center">\({@f3@}\)</td>
    <td style="text-align:center">\({@cf3@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x4@}\)</td>
    <td style="text-align:center">\({@f4@}\)</td>
    <td style="text-align:center">\({@cf4@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x5@}\)</td>
    <td style="text-align:center">\({@f5@}\)</td>
    <td style="text-align:center">\({@cf5@}\)</td>
  </tr>

  <tr>
    <td style="text-align:center">\({@x6@}\)</td>
    <td style="text-align:center">\({@f6@}\)</td>
    <td style="text-align:center">\({@cf6@}\)</td>
  </tr>
</table>

<p>
There are \(N={@N@}\) observations in total.
</p>

<hr>

<p>
The position of \(P_{@p1@}\) is
</p>

[[if test="intpos1"]]

<p>
\[
\frac{@p1@}{100}\times {@N@}
= {@ppos1@}.
\]
</p>

<p>
Therefore, \(P_{@p1@}\) corresponds to the
\({@ppos1@}^{\text{th}}\) observation.
</p>

<p>
From the cumulative frequencies, this observation has the value
\({@ta6@}\). Therefore,
</p>

[[else]]

<p>
\[
\frac{@p1@}{100}\times {@N@}
= {@dispdp(ppos1,2)@}.
\]
</p>

<p>
Therefore, \(P_{@p1@}\) lies between the
\({@lowerpos1@}^{\text{th}}\) and
\({@upperpos1@}^{\text{th}}\) observations.
</p>

<p>
From the cumulative frequencies, both of these observations have the value
\({@ta6@}\). Therefore,
</p>

[[/if]]

<p>
\[
\boxed{P_{@p1@}={@ta6@}}.
\]
</p>

<hr>

<p>
The position of \(P_{@p2@}\) is
</p>

[[if test="intpos2"]]

<p>
\[
\frac{@p2@}{100}\times {@N@}
= {@ppos2@}.
\]
</p>

<p>
Therefore, \(P_{@p2@}\) corresponds to the
\({@ppos2@}^{\text{th}}\) observation.
</p>

<p>
From the cumulative frequencies, this observation has the value
\({@ta7@}\). Therefore,
</p>

[[else]]

<p>
\[
\frac{@p2@}{100}\times {@N@}
= {@dispdp(ppos2,2)@}.
\]
</p>

<p>
Therefore, \(P_{@p2@}\) lies between the
\({@lowerpos2@}^{\text{th}}\) and
\({@upperpos2@}^{\text{th}}\) observations.
</p>

<p>
From the cumulative frequencies, both of these observations have the value
\({@ta7@}\). Therefore,
</p>

[[/if]]

<p>
\[
\boxed{P_{@p2@}={@ta7@}}.
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Generate six consecutive data values */
x1: rand(11)+10;
x2: x1+1;
x3: x1+2;
x4: x1+3;
x5: x1+4;
x6: x1+5;

xvals: [x1,x2,x3,x4,x5,x6];


/* Generate six distinct frequencies between 5 and 14 */
f1: rand(10)+5;
f2: rand_with_prohib(5,14,[f1]);
f3: rand_with_prohib(5,14,[f1,f2]);
f4: rand_with_prohib(5,14,[f1,f2,f3]);
f5: rand_with_prohib(5,14,[f1,f2,f3,f4]);
f6: rand_with_prohib(5,14,[f1,f2,f3,f4,f5]);

freqs: [f1,f2,f3,f4,f5,f6];


/* Calculate cumulative frequencies */
cf1: f1;
cf2: cf1+f2;
cf3: cf2+f3;
cf4: cf3+f4;
cf5: cf4+f5;
cf6: cf5+f6;

cumfreqs: [cf1,cf2,cf3,cf4,cf5,cf6];


/* Total number of observations */
N: cf6;


/*
   Select two target rows in ascending order.

   row1 can be 2, 3 or 4.
   row2 is guaranteed to be strictly greater than row1,
   with a maximum value of 5.
*/
row1: rand(3)+2;
row2: row1 + rand(5-row1) + 1;


/*
   Cumulative frequencies immediately before and
   at the end of each selected row.
*/
cf_before1: cumfreqs[row1-1];
cf_after1: cumfreqs[row1];

cf_before2: cumfreqs[row2-1];
cf_after2: cumfreqs[row2];


/*
   Generate an integer observation position strictly
   inside each selected cumulative-frequency block.
*/
pos1: cf_before1 + 1 + rand(cf_after1-cf_before1-1);
pos2: cf_before2 + 1 + rand(cf_after2-cf_before2-1);


/*
   Convert each selected position into an integer
   percentile rank to display in the question.
*/
p1: round(100*pos1/N);
p2: round(100*pos2/N);


/*
   Recalculate the percentile positions from the
   percentile ranks actually shown to the student.
*/
ppos1: p1*N/100;
ppos2: p2*N/100;


/*
   Record whether the displayed percentile produces
   an exact integer observation position.
*/
intpos1: integerp(ppos1);
intpos2: integerp(ppos2);


/*
   Positions immediately below and above each
   calculated percentile position.

   These are used when the position is not an integer.
*/
lowerpos1: floor(ppos1);
upperpos1: ceiling(ppos1);

lowerpos2: floor(ppos2);
upperpos2: ceiling(ppos2);


/*
   Teacher answers.

   ta1-ta5: missing cumulative frequencies.
   ta6-ta7: values of the two requested percentiles.
*/
ta1: cf2;
ta2: cf3;
ta3: cf4;
ta4: cf5;
ta5: cf6;

ta6: xvals[row1];
ta7: xvals[row2];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>Cumulative frequencies: \({@cf1@}, {@cf2@}, {@cf3@}, {@cf4@}, {@cf5@}, {@cf6@}\). Percentiles: \(P_{@p1@}={@ta6@}\), \(P_{@p2@}={@ta7@}\). {@intpos1@}, {@intpos2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2006124738</deployedseed>
    <deployedseed>479113790</deployedseed>
    <deployedseed>350868234</deployedseed>
    <deployedseed>605766956</deployedseed>
    <deployedseed>718967969</deployedseed>
    <deployedseed>327101754</deployedseed>
    <deployedseed>1855618969</deployedseed>
    <deployedseed>358879926</deployedseed>
    <deployedseed>848855284</deployedseed>
    <deployedseed>1677009791</deployedseed>
    <deployedseed>1885479513</deployedseed>
    <deployedseed>503642728</deployedseed>
    <deployedseed>1620028556</deployedseed>
    <deployedseed>1409512788</deployedseed>
    <deployedseed>129233444</deployedseed>
    <deployedseed>2144506222</deployedseed>
    <deployedseed>2138103208</deployedseed>
    <deployedseed>1887962067</deployedseed>
    <deployedseed>667573085</deployedseed>
    <deployedseed>557538297</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt3-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt4-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt5-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt6-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt7-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>