<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Household Appliance Survey: Calculating Probabilities of Electric Devices</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A survey is made to determine the number of households having electric appliances in a certain city. It is found that \({@r@} \%\) have radios \((\textbf{R})\), \({@iron@} \%\) have irons \((\textbf{I})\), \({@t@} \%\) have electric toasters \((\textbf{T})\), \({@ir@} \%\) have irons and radio \((\textbf{IR})\), \({@it@} \%\) have irons and electric toasters \((\textbf{IR})\), \({@rt@} \%\) have radio and electric toasters \((\textbf{RT})\) and \({@irt@} \%\) have all three.</p>
<p>Find the probability that a household has at least one of the appliances</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the probability that each household has at least one of the appliances (radio, iron, or electric toaster), we can use the complementary probability and proceed as shown below.</p>
<p>The probability that a household has at least one of the appliances</p>
<p>\(P[R \cup I \cup T]=P[R]+P[I]+P[T]-P[R \cap I]-P[R \cap T]-P[I \cap T]+P[R \cap T\cap I]\)</p>
<p>\(P[R \cup I \cup T]={@float(r/100)@}+{@float(iron/100)@}+{@float(t/100)@}-{@float(ir/100)@}-{@float(rt/100)@}-{@float(it/100)@}+{@float(irt/100)@}={@ta1@}\)</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>non:rand_with_step(3,10,1);
irt:rand_with_prohib(6,10,[non]);
i_on:rand_with_prohib(6,15,[non,irt]);
r_on:rand_with_prohib(6,15,[non,irt,i_on]);
t_on:rand_with_prohib(6,15,[non,irt,i_on,r_on]);
ir_on:rand_with_prohib(6,15,[non,irt,i_on,r_on,t_on]);
it_on :rand_with_prohib(6,15,[non,irt,i_on,r_on,t_on,ir_on]);
rt_on:100-(non+irt+i_on+r_on+t_on+ir_on+it_on);
ir:ir_on+irt;
it:it_on+irt;
rt:rt_on+irt;
/*probabilities*/
iron:i_on+ir_on+it_on+irt;
r:r_on+ir_on+rt_on+irt;
t:t_on+it_on+rt_on+irt;

/*model answers*/
ta:r+iron+t-ir-rt-it+irt; /*should be the same as 1-non*/
ta1:float(ta/100);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \({@r@} \%\) have radios \(\textbf{R}\)</p>
<p>\({@iron@} \%\) have irons \(\textbf{I}\)</p>
<p>\({@t@} \%\) have electric toasters \( \textbf{T} \)</p>
<p>\({@ir@} \%\) have irons and radio \(\textbf{IR}\)</p>
<p>\({@rt@} \%\) have electric toasters and radio \(\textbf{IR}\)</p>
<p>\({@it@} \%\) have irons and electric toasters \( \textbf{IT}\) and \({@irt@} \%\) have all three</p>
<p>\(P[R \cup I \cup T]={@float(r/100)@}+{@float(iron/100)@}+{@float(t/100)@}-{@float(ir/100)@}-{@float(rt/100)@}-{@float(it/100)@}+{@float(irt/100)@}={@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1cUGpx6r8Sq8RUiP3oNxfeuFeLnuWZ4u5/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Georg Osang</p>
<p id="description">In this question a survey is made to determine  number of households having electric appliances in a certain city. It is found that is a varied percentage of people with radio,electric toasters, iron, and those with double items. The question is to find the probability that household has at least one of these appliances</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You gave your answer in percent, but you were supposed to give it as a normalized probability (between 0 and 1).</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1377404703</deployedseed>
    <deployedseed>1535380681</deployedseed>
    <deployedseed>950937430</deployedseed>
    <deployedseed>1316989182</deployedseed>
    <deployedseed>676401844</deployedseed>
    <deployedseed>1784656833</deployedseed>
    <deployedseed>1786305296</deployedseed>
    <deployedseed>236751065</deployedseed>
    <deployedseed>1920427895</deployedseed>
    <deployedseed>496681913</deployedseed>
    <deployedseed>1890449478</deployedseed>
    <deployedseed>469685668</deployedseed>
    <deployedseed>615992219</deployedseed>
    <deployedseed>1050114</deployedseed>
    <deployedseed>993005632</deployedseed>
    <deployedseed>648243322</deployedseed>
    <deployedseed>79648844</deployedseed>
    <deployedseed>1180083319</deployedseed>
    <deployedseed>1156253481</deployedseed>
    <deployedseed>304787516</deployedseed>
    <deployedseed>858545664</deployedseed>
    <deployedseed>887913171</deployedseed>
    <deployedseed>1132905322</deployedseed>
    <deployedseed>345872880</deployedseed>
    <deployedseed>1776370968</deployedseed>
    <deployedseed>175082907</deployedseed>
    <deployedseed>366676417</deployedseed>
    <deployedseed>997866962</deployedseed>
    <deployedseed>76558166</deployedseed>
    <deployedseed>1757983918</deployedseed>
    <deployedseed>49711969</deployedseed>
    <deployedseed>464195459</deployedseed>
    <deployedseed>498094815</deployedseed>
    <deployedseed>1596231951</deployedseed>
    <deployedseed>1620944183</deployedseed>
    <deployedseed>750281961</deployedseed>
    <deployedseed>607287834</deployedseed>
    <deployedseed>704428087</deployedseed>
    <deployedseed>1293841888</deployedseed>
    <deployedseed>1940039933</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>