<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Infant Mortality and its Measures_q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given below is the data relating to infant deaths and births in Guatemala, taking the unit of measurement as \('000'\) for live births and deaths .</p>
<br>
<p>
<table border="solid">
  <tr>
    <th>1959</th>
    <th>1960</th>
  </tr>
  <tr>
    <td>{@d2@} (Late neonatal deaths in \(1959\))</td>
    <td>{@d4@} (Late neonatal deaths in \(1960\))</td>
  </tr>
  <tr>
    <td>{@d1@} (Early neonatal deaths in \(1959\))</td>
    <td>{@d3@} (Early neonatal deaths in \(1960\))</td>
  </tr>
 
</table>
</p>
<p></p>
<p>Given the total number of live births in Guatemala in the year \(1959\) and \(1960\) as {@b1@} and {@b2@} respectfully, answer the questions below.</p>
<br>
<p>Note: Leave your answer to correct \(2\) dp</p>
<p></p>
<p>a). Determine Infant mortality rate by numerator separation method \(=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>b). Determine Infant mortality rate by adjustment formulae separation factor \(=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p></p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>
<table border="solid">
  <tr>
    <th>1959</th>
    <th>1960</th>
  </tr>
  <tr>
    <td>\({@d2@} \,\, D_{2}\) (Late neonatal deaths in \(1959\) )</td>
    <td>\({@d4@} \,D_{4}\) (Late neonatal deaths in \(1960\) )</td>
  </tr>
  <tr>
    <td>\({@d1@} \,D_{1}\) (Early neonatal deaths in \(1959\))</td>
    <td>\({@d3@} \,D_{3}\) (Early neonatal deaths in \(1960\))</td>
  </tr>
 
</table>
</p>

<p>a). Determining Infant mortality rate by numerator separation method </p>
<p>To determine Infant mortality rate by numerator separation method,we use the formula \(\displaystyle \left [ \frac{D_4}{B_2}+\frac{D_3}{B_1} \right ] \times 1000\) </p>
<p>Applying the formula </p>
<p>
\(\begin{align}
\displaystyle \left [ \frac{D_4}{B_2}+\frac{D_3}{B_1} \right ] \times 1000&=\displaystyle \left [ \frac{{@d4@}}{{@b2@}}+\frac{{@d3@}}{{@b1@}} \right ] \times 1000\\
&={@sol3@}
\end{align}\)
</p>
<br>
<p>b). Determining Infant mortality rate by adjustment formulae separation factor</p>
<p>Infant mortality rate by adjustment formulae separation factor is determined as below</p>
<p>Separation factor (f) is given by \(=\displaystyle \left [ \frac{D_3}{D_3+D_4}\right ]={@decimalplaces(sep_factor,4)@}\)</p>
<p>Adjustment formula,</p>
<p>
\(\begin{align}
\displaystyle \left [ \frac{D_3+D_4}{fB_1+(1-f)B_2}\right ] \times 1000&=\displaystyle \left [ \frac{{@D_3@}+{@D_4@}}{{@decimalplaces(sep_factor,4)@}\times {@b1@}+(1-{@decimalplaces(sep_factor,4)@}) \times{@b2@}}\right ] \times 1000\\
&=\displaystyle \left [ \frac{{@d3+d4@}}{{@(1-decimalplaces(sep_factor,4))*b2@}}\right ] \times 1000\\
&={@sol4@}\\
\end{align}\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*topic 8*/
deaths:makelist((rand_with_step(5000,5500,100)),i,1,2);
d1:deaths[1];
d3:deaths[2]+rand(500);
/*d2 and d4*/
deaths_2:makelist((rand_with_step(10000,12000,1000)),i,1,2);
d2:deaths_2[1];
d4:deaths_2[2]+rand(500);
/*live births*/
births:makelist((rand_with_step(180000,200000,1000)),i,1,2);
b1:births[1];
b2:births[2]+rand(1000);
/*worked solution variables*/
/*separation*/
sol3:decimalplaces(((d4/b2)+(d3/b1))*1000,2); 
/*adjustment formula*/
sep_factor:(d3/(d3+d4));
s1:sep_factor*b1+(1-sep_factor)*b2;
sol4:decimalplaces(((d3+d4)/(s1))*1000,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>a). Determine Infant mortality rate by numerator separation method \(={@sol3@}\) </p>
<p>b). Determine Infant mortality rate by adjustment formulae separation factor \(={@sol4@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question requires the student to calculate the crude infant mortality rate  and Determine Infant mortality rate by product formula given the deaths and live births for years \(1959\) and \(1960\)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>sol3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>sol4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sol3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>sol3</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct anser</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sol4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for no of decimal places</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>sol4</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not given to the right number of decimal places</p></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>514072334</deployedseed>
    <deployedseed>319487233</deployedseed>
    <deployedseed>1096983354</deployedseed>
    <deployedseed>22001896</deployedseed>
    <deployedseed>2016547698</deployedseed>
    <deployedseed>503156368</deployedseed>
    <deployedseed>1681456936</deployedseed>
    <deployedseed>603086195</deployedseed>
    <deployedseed>1274943153</deployedseed>
    <deployedseed>573812401</deployedseed>
    <deployedseed>1356143622</deployedseed>
    <deployedseed>1643305617</deployedseed>
    <deployedseed>1725926493</deployedseed>
    <deployedseed>696142532</deployedseed>
    <deployedseed>1192265850</deployedseed>
    <deployedseed>842172355</deployedseed>
    <deployedseed>298762209</deployedseed>
    <deployedseed>1764382946</deployedseed>
    <deployedseed>516321142</deployedseed>
    <deployedseed>78320962</deployedseed>
    <deployedseed>143376629</deployedseed>
    <deployedseed>182803710</deployedseed>
    <deployedseed>197224446</deployedseed>
    <deployedseed>487804621</deployedseed>
    <deployedseed>998781849</deployedseed>
    <deployedseed>797959870</deployedseed>
    <deployedseed>1921241475</deployedseed>
    <deployedseed>1331359256</deployedseed>
    <deployedseed>2102799671</deployedseed>
    <deployedseed>722605718</deployedseed>
    <deployedseed>1841457142</deployedseed>
    <deployedseed>1457065273</deployedseed>
    <deployedseed>703593449</deployedseed>
    <deployedseed>1468084452</deployedseed>
    <deployedseed>818971851</deployedseed>
    <deployedseed>1474230118</deployedseed>
    <deployedseed>1673082078</deployedseed>
    <deployedseed>442274151</deployedseed>
    <deployedseed>286681305</deployedseed>
    <deployedseed>1866423657</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>sol3</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>sol4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>