<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Interpret the text and use the information to calculate conditional probabilities</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>During an epidemic in a town, {@aa@}% of its inhabitants became sick. Of any {@100@} sick persons, {@cc@} will need to be admitted to an emergency ward. What is the probability that a randomly chosen person from this town will be admitted to an emergency ward?
[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>Give answer correct to 2 decimal places</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution</b></p>
<p>Let</p>
<p>\(A=\) the person is healthy</p>
<p>and</p>
<p>\(B=\) the person is admitted to an emergency  ward</p>
<p>It is given</p>
<p>\(P(A^{C})={@ee@}\)</p>
<p>Hence,</p>
<p>\(P(A)={@dd_1@}\)</p>
<p>We want to find \(P(B)\). Now \(P(B|A)=0\),  because a heathy person will not be admitted to an emergency ward. Also,</p>
<p>\(P(B|A^{C})=\frac{{@cc@}}{100}={@ta_1@}\)</p>
<p>Hence, by the total probability rule,</p>
<p>
\(
\begin{align}
  P(B) &=P(A)P(B|A)+P(A^C)P(B|A^C) \\
  &={@dd_1@}\cdot 0+{@ee@}\cdot {@ff_1@} \\
  &= {@(dd_1*0)+(ee*ff_1)@}\\
 &={@ta@}
\end{align}
\)]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Question Variables*/
aa:rand(70)+20;/*percentage of  inhabitants becoming sick*/
bb:100;
cc:rand(20)+10;/*persons admitted to an emergency ward*/
dd:(100-aa)/100;/*P(A)*/
dd_1:decimalplaces(float(dd),4);
ee_1:1-dd_1;/*P(A^C)*/
ee:decimalplaces(ee_1,2);
ff:cc/bb;/*P(B|A^C)*/
ff_1:decimalplaces(float(ff),4);
gg:dd*0+ee*ff;/*P(B)*/
/*Model answer*/
ta:decimalplaces(float(gg),2);
ta_1:float(cc/100);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2: mistake:Interchanging the P(A) with P(A^C)*/
mistake_2_1:aa/100;/*P(A)*/
mistake_2_2:decimalplaces(float(mistake_2_1),2);
mistake_2_3:dd_1;/*P(A^C)*/
mistake_2_4:mistake_2_1*0+mistake_2_3*ff;
mistake_2:decimalplaces(mistake_2_4,2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have interchanged the  P(A) with P(A^C)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>