<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>LIFE TABLES (Single Decrement Tables)_q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Which of the following <strong>CANNOT</strong> be determined using a population Life table?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>The correct answer is <strong> The probability of literacy from any given age to any other given age</strong>.</p>
<p>The probability of literacy from any given age to any other given age has nothing to do with deaths or longevity of a human being </p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>The correct answer is <strong>The probability of illness for any given number of persons at any age</strong> .</p>
<p>The probability of illness for any given number of persons at any age has nothing to do with deaths or longevity of a human being as it deals with epidemiology, specifically with the study of disease occurrence and distribution within populations.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>The correct answer is <strong>The probability of transmitting a disease from any given age to any other given age </strong>.</p>
<p>The probability of transmitting a disease from any given age to any other given age has nothing to do with deaths or longevity of a human being as it deals with epidemiology, specifically with the study of disease occurrence and distribution within populations.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[correct_answer:rand_selection([
[1, true, "The probability of literacy from any given age to any other given age."],
[2, true, "The probability of illness for any given number of persons at any age."],
[3, true, "The probability of transmitting a disease from any given age to any other given age.
"]],1);

incorrect_answers:rand_selection([
[4, false, "The probability of dying within one year of any person at each age.
"],
[5, false, "The average number of years of life remaining for a person of any age."],
[6, false, "The average number of years a newborn infant can expect to live.
"],
[7, false, "The number of persons at a particular age above or below a particular age or between any two ages under conditions of a stationary population.
"]], 3);
ta: random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));
ta_largest: last(ta_index);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni<p>
<p id="concept">https://docs.google.com/document/d/1yKMASRNBrxv56EctlLkgvdHRMASolQhCBNM67lq8h9M/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">The question focuses on outlining the major concerns of a demographer</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct option</description>
        <answertest>AlgEquiv</answertest>
        <sans>[ans1]</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2064413255</deployedseed>
    <deployedseed>1505987292</deployedseed>
    <deployedseed>248992741</deployedseed>
    <deployedseed>1888676474</deployedseed>
    <deployedseed>832679805</deployedseed>
    <deployedseed>696606810</deployedseed>
    <deployedseed>300179244</deployedseed>
    <deployedseed>839931534</deployedseed>
    <deployedseed>1580596430</deployedseed>
    <deployedseed>153034525</deployedseed>
    <deployedseed>1042166734</deployedseed>
    <deployedseed>231555006</deployedseed>
    <deployedseed>700511361</deployedseed>
    <deployedseed>238971822</deployedseed>
    <deployedseed>965790117</deployedseed>
    <deployedseed>1361340779</deployedseed>
    <deployedseed>12805689</deployedseed>
    <deployedseed>1100156112</deployedseed>
    <deployedseed>647508334</deployedseed>
    <deployedseed>1251369272</deployedseed>
    <deployedseed>1146041233</deployedseed>
    <deployedseed>595894462</deployedseed>
    <deployedseed>18868006</deployedseed>
    <deployedseed>1551224673</deployedseed>
    <deployedseed>1889872727</deployedseed>
    <deployedseed>1162419845</deployedseed>
    <deployedseed>1742633348</deployedseed>
    <deployedseed>296118395</deployedseed>
    <deployedseed>1967299012</deployedseed>
    <deployedseed>1034272311</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>