<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>LIFE TABLES (Single Decrement Tables)_q4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Use the Australian Males Life Table \((1960-62)-\text{A}^{\text{M}61}\) given below to answer the question that follows.</p>
<p><strong style="color: blue;">
 <a href="https://drive.google.com/uc?export=download&id=1HLhNdPh1r0467c_FfyHkov1UPUZk1ZE1">Download Australian Males Life Table Data </a>
</strong></p>
<p>Using the life tables for Australian Males, estimate the following measures.</p>
<p>What is the expected age at death assuming \(A^{m61}\) mortality of three groups of persons now aged \({@age_index1@}\), \({@age_index2@}\) and \({@age_index3@}\) respectively.</p>
<p><b>Note:</b> Give your answers correct to \(2\) decimal places.</p>
<p>\((a)\,\) Persons now aged \({@age_index1@}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>\((b)\,\) Persons now aged \({@age_index2@}=\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>\((c)\,\) Persons now aged \({@age_index3@}=\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Persons now aged \({@age_index1@}\) on the average are expected to live further \(e^0{@age_index1@}\) years and hence their average age at death is;</p>
<p>\((a). \, {@age_index1@} + \frac{T_{{@age_index1@}}}{l_{{@age_index1@}}}={@age_index1@}+\frac{{@T_num_1@}}{{@l_num_1@}}={@age_index1@}+{@decimalplaces((T_num_1/l_num_1),2)@}={@ta1@}\)</p>
<p>\((b). \, {@age_index2@} + \frac{T_{{@age_index2@}}}{l_{{@age_index2@}}}={@age_index2@}+\frac{{@T_num_2@}}{{@l_num_2@}}={@age_index2@}+{@decimalplaces((T_num_2/l_num_2),2)@}={@ta2@}\)</p>
<p>\((c). \, {@age_index3@} + \frac{T_{{@age_index3@}}}{l_{{@age_index3@}}}={@age_index3@}+\frac{{@T_num_3@}}{{@l_num_3@}}={@age_index3@}+{@decimalplaces((T_num_3/l_num_3),2)@}={@ta3@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*LIFE TABLES*/
/* Define the data points */ 
	age: [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 102, 103, 104, 105, 106, 107, 108, 109]; 
	val_lx: [100000, 97761, 97584, 97467, 97379, 97315, 97259, 97206, 97154, 97105, 97082, 97022, 96981, 96936, 96885, 96825, 96752, 96660, 96541, 96384, 96215, 96049, 95886, 95728, 95577, 95432, 95292, 95154, 95014, 94871, 94726, 94577, 94425, 94267, 94123, 93931, 93749, 93554, 93349, 93112, 92859, 92580, 92374, 91938, 91569, 91165, 90723, 90238, 89705, 89118, 88473, 87762, 86879, 86119, 85175, 84142, 83015, 81790, 80459, 79017, 77456, 75771, 73954, 72002, 69915, 67699, 65361, 62910, 60353, 57696, 54944, 52100, 49168, 46160, 43092, 39984, 36860, 33745, 30661, 27629, 24669, 21803, 19054, 16448, 14008, 11758, 9716, 7897, 6306, 4943, 3800, 2862, 2111, 1524, 1076, 743, 500, 329, 211, 132, 80, 47, 27, 15, 8, 4, 2, 1, 0, 0];  
       val_Tx: [6791451, 6623018, 6535363, 6497838, 6400415, 6303088, 6206781, 6108548,5011358, 5914239, 5817155, 5720113,5623112, 5526153, 5429243, 5332388, 5235599, 5138893, 5042293, 4945830, 4849531, 4753399, 4657431, 4581624, 4465972, 4370467, 4275105, 4179882, 4054798, 3989856, 3895057, 3800405, 3705905, 3611559, 3517374, 3423357, 3329517, 3235353, 3142417, 3049199, 2958204, 2903484, 2771057, 2678951, 2587198, 2485831, 2404887, 2314406, 2224435, 2135023, 2045228, 1958110, 1870740, 1784191, 1698544, 1513885, 1530307, 1447904, 1366780, 1287042, 1208805, 1132192, 1057329, 984351, 913393, 844586, 778056, 713920, 652289, 593264, 536944, 483422, 432788, 385124, 340498, 298960, 260538, 225236, 193033, 163888, 137739, 114503, 94074, 76323, 61095, 48212, 37475, 28669, 21567, 15943, 11571, 8240, 5754, 3936, 2636, 1727, 1106, 692, 422, 250, 144, 81, 44, 23, 11, 5, 2, 1, 0, 0];  

/* Given data */
index_num:rand(40)+1;
age_index1:age[index_num];
age_index2:age_index1+20;
age_index3:age_index2+20;

T_num_1:val_Tx[age_index1+1]; /*T value at the first given age*/
T_num_2:val_Tx[age_index2+1]; /*T value at the second given age*/
T_num_3:val_Tx[age_index3+1]; /*T value at the third given age*/

l_num_1:val_lx[age_index1+1]; /*l value at the first given age*/
l_num_2:val_lx[age_index2+1]; /*l value at the second given age*/
l_num_3:val_lx[age_index3+1]; /*l value at the third given age*/
/*Model answers*/
ta1:decimalplaces((age_index1+(T_num_1/l_num_1)),2);
ta2:decimalplaces((age_index2+(T_num_2/l_num_2)),2);
ta3:decimalplaces((age_index3+(T_num_3/l_num_3)),2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\((a)\,\) Persons now aged \({@age_index1@}={@decimalplaces(ta1,2)@}\) </p>
<p>\((b)\,\) Persons now aged \({@age_index2@}={@decimalplaces(ta2,2)@}\)</p>
<p>\((c)\,\) Persons now aged \({@age_index2@}={@decimalplaces(ta3,2)@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat<p>
<p id="concept">https://docs.google.com/document/d/1l78jWPiFZz6jKY9qpa41nXpADuUquauIihBbw6oRdbc/edit#heading=h.eddp73mos3v3</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question requires the student to utilize the Australian Males Life Table \((1960-62)-\text{A}^{\text{M}61}\) to estimate the expected age at death for three groups of individuals </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer for the first group of persons*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the first group of persons</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct answer for the first group of persons</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer for the second group of persons*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the second group of persons</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct answer for the second group of personsChecks for the correct answer for the second group of persons</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer for the third group of persons*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer for the third group of persons</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct answer for the third group of persons</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>523012097</deployedseed>
    <deployedseed>1844895447</deployedseed>
    <deployedseed>1825335369</deployedseed>
    <deployedseed>256953688</deployedseed>
    <deployedseed>1688686217</deployedseed>
    <deployedseed>845635175</deployedseed>
    <deployedseed>2090945600</deployedseed>
    <deployedseed>2114294538</deployedseed>
    <deployedseed>1555549107</deployedseed>
    <deployedseed>1545137487</deployedseed>
    <deployedseed>910257802</deployedseed>
    <deployedseed>1037735081</deployedseed>
    <deployedseed>187057667</deployedseed>
    <deployedseed>1488105468</deployedseed>
    <deployedseed>786085208</deployedseed>
    <deployedseed>1844081360</deployedseed>
    <deployedseed>188866049</deployedseed>
    <deployedseed>289004141</deployedseed>
    <deployedseed>926507749</deployedseed>
    <deployedseed>71617664</deployedseed>
    <deployedseed>78008412</deployedseed>
    <deployedseed>1226688866</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>