<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>LINEAR REGRESSION-Q4</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Ten boys compete in throwing a cricket ball and the table below shows the height of each boy (x) measured to the nearest centimeter and the distance (y) to which he can throw the ball
</p>
<table border="solid">
<tr>
      <th width="28%", style="text-align:center">Boy</th>
       <th width="6%", style="text-align:center">\(A\)</th>
       <th width="6%", style="text-align:center">\(B\)</th>
       <th width="6%", style="text-align:center">\(C\)</th>
      <th width="6%", style="text-align:center">\(D\)</th>
       <th width="6%", style="text-align:center">\(E\)</th>
      <th width="6%", style="text-align:center">\(F\)</th>
       <th width="6%", style="text-align:center">\(G\)</th>
       <th width="6%", style="text-align:center">\(H\)</th>
        <th width="6%", style="text-align:center">\(I\)</th>
       <th width="6%", style="text-align:center">\(J\)</th>
    
  </tr>
  <tr>
      <th width="28%", style="text-align:center">Height of each boy (x)</th>
       <th width="6%", style="text-align:center">\({@x1@}\)</th>
       <th width="6%", style="text-align:center">\({@x2@}\)</th>
       <th width="6%", style="text-align:center">\({@x3@}\)</th>
      <th width="6%", style="text-align:center">\({@x4@}\)</th>
       <th width="6%", style="text-align:center">\({@x5@}\)</th>
      <th width="6%", style="text-align:center">\({@x6@}\)</th>
       <th width="6%", style="text-align:center">\({@x7@}\)</th>
       <th width="6%", style="text-align:center">\({@x8@}\)</th>
       <th width="6%", style="text-align:center">\({@x9@}\)</th>
      <th width="6%", style="text-align:center">\({@x10@}\)</th>
       
  </tr>
   <tr>
      <th width="28%", style="text-align:center">Distance (y)</th>
       <th width="6%", style="text-align:center">\({@y1@}\)</th>
       <th width="6%", style="text-align:center">\({@y2@}\)</th>
       <th width="6%", style="text-align:center">\({@y3@}\)</th>
      <th width="6%", style="text-align:center">\({@y4@}\)</th>
       <th width="6%", style="text-align:center">\({@y5@}\)</th>
      <th width="6%", style="text-align:center">\({@y6@}\)</th>
       <th width="6%", style="text-align:center">\({@y7@}\)</th>
       <th width="6%", style="text-align:center">\({@y8@}\)</th>
       <th width="6%", style="text-align:center">\({@y9@}\)</th>
      <th width="6%", style="text-align:center">\({@y10@}\)</th>
     
  </tr>
 </table>
<br>
<p><b>Note:</b> Give your answers correct to \(2\) decimal places</p>
<p>Calculate the:</p>
<ol>
<li>
<p>Product moment correlation coefficient \(\left(r\right)\)</p>
<p>\(r= \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
</li>
<li>
<p>Spearman’s coefficient of rank correlation \(\left( \rho \right)\)</p>
<p>\(\rho= \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
</li>
<li>
<p>Equations of regression line \(\left(y= mx + c \right)\) .</p>
<p>\(y=\) [[input:ans3]] x \(+\) [[input:ans4]] [[validation:ans3]][[validation:ans4]][[feedback:prt3]][[feedback:prt4]]</p>
</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To solve this problem you'll need to create a table which looks like this as the first step.</p>
<table border="solid">
  <tr>
      <th width="17%", style="text-align:center">\(x\)</th>
       <th width="22%", style="text-align:center">\(x^2\)</th>
       <th width="17%", style="text-align:center">\(y\)</th>
       <th width="22%", style="text-align:center">\(y^2\)</th>
      <th width="22%", style="text-align:center">\(x y\)</th>
   <th width="15%", style="text-align:center">Rank(x) </th>
       <th width="15%", style="text-align:center">Rank(y)</th>
<th width="40%", style="text-align:center">d \(=\) Rank(x)\(-\)Rank(y)</th>
       <th width="10%", style="text-align:center">\(\text{d}^2\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x1@}\)</th>
       <th width="22%", style="text-align:center">\({@x1^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y1@}\)</th>
       <th width="22%", style="text-align:center">\({@y1^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x1*y1@}\)</th>
   <th width="15%", style="text-align:center">\({@rank_x[1]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[1]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[1]-rank_y[1]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[1]-rank_y[1])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x2@}\)</th>
       <th width="22%", style="text-align:center">\({@x2^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y2@}\)</th>
       <th width="22%", style="text-align:center">\({@y2^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x2*y2@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[2]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[2]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[2]-rank_y[2]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[2]-rank_y[2])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x3@}\)</th>
       <th width="22%", style="text-align:center">\({@x3^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y3@}\)</th>
       <th width="22%", style="text-align:center">\({@y3^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x3*y3@}\)</th>

<th width="15%", style="text-align:center">\({@rank_x[3]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[3]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[3]-rank_y[3]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[3]-rank_y[3])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x4@}\)</th>
       <th width="22%", style="text-align:center">\({@x4^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y4@}\)</th>
       <th width="22%", style="text-align:center">\({@y4^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x4*y4@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[4]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[4]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[4]-rank_y[4]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[4]-rank_y[4])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x5@}\)</th>
       <th width="22%", style="text-align:center">\({@x5^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y5@}\)</th>
       <th width="22%", style="text-align:center">\({@y5^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x5*y5@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[5]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[5]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[5]-rank_y[5]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[5]-rank_y[5])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x6@}\)</th>
       <th width="22%", style="text-align:center">\({@x6^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y6@}\)</th>
       <th width="22%", style="text-align:center">\({@y6^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x6*y6@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[6]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[6]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[6]-rank_y[6]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[6]-rank_y[6])^2@}\)</th>
  </tr>
   <tr>
      <th width="17%", style="text-align:center">\({@x7@}\)</th>
       <th width="22%", style="text-align:center">\({@x7^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y7@}\)</th>
       <th width="22%", style="text-align:center">\({@y7^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x7*y7@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[7]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[7]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[7]-rank_y[7]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[7]-rank_y[7])^2@}\)</th>
  </tr>
   <tr>
      <th width="17%", style="text-align:center">\({@x8@}\)</th>
       <th width="22%", style="text-align:center">\({@x8^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y8@}\)</th>
       <th width="22%", style="text-align:center">\({@y8^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x8*y8@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[8]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[8]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[8]-rank_y[8]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[8]-rank_y[8])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x9@}\)</th>
       <th width="22%", style="text-align:center">\({@x9^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y9@}\)</th>
       <th width="22%", style="text-align:center">\({@y9^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x9*y9@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[9]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[9]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[9]-rank_y[9]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[9]-rank_y[9])^2@}\)</th>
  </tr>
  <tr>
      <th width="17%", style="text-align:center">\({@x10@}\)</th>
       <th width="22%", style="text-align:center">\({@x10^2@}\)</th>
       <th width="17%", style="text-align:center">\({@y10@}\)</th>
       <th width="22%", style="text-align:center">\({@y10^2@}\)</th>
      <th width="22%", style="text-align:center">\({@x10*y10@}\)</th>
<th width="15%", style="text-align:center">\({@rank_x[10]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[10]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[10]-rank_y[10]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[10]-rank_y[10])^2@}\)</th>
  </tr>
 
  <tr>
      <th width="17%", style="text-align:center">\(\sum x={@sum_x@}\)</th>
       <th width="22%", style="text-align:center">\(\sum x^2={@sum_x_sq@}\)</th>
       <th width="17%", style="text-align:center">\(\sum y={@sum_y@}\)</th>
       <th width="22%", style="text-align:center">\(\sum y^2={@sum_y_sq@}\)</th>
      <th width="22%", style="text-align:center">\(\sum xy={@sum_xy@}\)</th>
  
  </tr>  
 </table>
<p>a)</p>
<p>\(n=10\)</p>
<p> Product moment correlation coefficient \(\displaystyle r=\frac{n\left ( \sum xy \right )-\left ( \sum x \right )\left ( \sum y\right )}{\sqrt{\left [ n\left ( \sum x^{2} \right ) -\left ( \sum x \right )^{2}\right ]\left [  n\left ( \sum y^{2} \right ) -\left ( \sum y \right )^{2} \right ]}}\)</p>
<p>\(\displaystyle \qquad \qquad \qquad \qquad \qquad \qquad \qquad=\frac{10\left ( {@sum_xy@} \right )-\left ( {@sum_x@} \right )\left ( {@sum_y@}\right )}{\sqrt{\left [ 10\left ( {@sum_x_sq@}\right ) -\left ( {@sum_x@} \right )^{2}\right ]\left [  10\left ({@sum_y_sq@} \right ) -\left ( {@sum_y@} \right )^{2} \right ]}}\)</p>
<p>\(\displaystyle \qquad \qquad \qquad \qquad \qquad \qquad \qquad=\frac{{@10*sum_xy@} -{@sum_x*sum_y@}}{\sqrt{\left [ {@10*sum_x_sq@} -{@sum_x^2@}\right ]\left [  {@10*sum_y_sq@}  - {@sum_y^2@} \right ]}}\)</p>
<p>\(\displaystyle \qquad \qquad \qquad \qquad \qquad \qquad \qquad=\frac{{@(10*sum_xy)-(sum_x*sum_y)@}}{\sqrt{\left ( {@(10*sum_x_sq)-sum_x^2@} \right )\left ( {@(10*sum_y_sq)-sum_y^2@}  \right )}}\)</p>
<p>\(\displaystyle \qquad \qquad \qquad \qquad \qquad \qquad \qquad=\frac{{@(10*sum_xy)-(sum_x*sum_y)@}}{\sqrt{{@((10*sum_x_sq)-(sum_x^2))*((10*sum_y_sq)-(sum_y^2))@} }}\)</p>

<p>\(\displaystyle \qquad \qquad \qquad \qquad \qquad \qquad \qquad={@ta_1@}\)</p>
<p>b) Spearman rank correlation  coefficient \(\displaystyle \rho=1- \frac{ 6\sum d^2}{n(n^2-1)}\)</p>
<p>\( \displaystyle r=1- \frac{6({@total_d_squared@})}{10(100-1)}\)</p>
<p>\( r=1-\frac{{@6*total_d_squared@}}{990}=1-{@(6*total_d_squared)/990@}\)</p>\(={@ta_2@}\)</p>



<p>c) The equation of regression line \(y\) on \(x\) is given by \(\hat{y}=\hat{\alpha}+\hat{\beta} x\).</p>
<p>where \(\displaystyle \hat{\alpha}=\frac{\sum y \sum x^2-\sum x \sum xy}{n \sum x^2-(\sum x)^2}=\frac{({@sum_y@}) ({@sum_x_sq@})-({@sum_x@})({@sum_xy@})}{12({@sum_x_sq@})-({@sum_x@})^2}={@alpha@}={@decimalplaces(alpha,4)@}\)</p>
<p>and</p>
<p>\(\displaystyle \hat{\beta}=\frac{n\sum xy -\sum x \sum y}{n \sum x^2-(\sum x)^2}=\frac{12({@sum_xy@})-({@sum_x@})({@sum_y@})}{12({@sum_x_sq@})-({@sum_x@})^2}={@beta@}={@decimalplaces(beta,4)@}\)</p>
<p>Therefore, the equation of regression line is given by: \(\hat{y}={@decimalplaces(alpha,4)@}+{@decimalplaces(beta,4)@} x\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>x1:rand_with_step(30,80,1);
x2:rand_with_prohib(30,80,[x1]);
x3:rand_with_prohib(30,80,[x1,x2]);
x4:rand_with_prohib(30,80,[x1,x2,x3]);
x5:rand_with_prohib(30,80,[x1,x2,x3,x4]);
x6:rand_with_prohib(30,80,[x1,x2,x3,x4,x5]);
x7:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6]);
x8:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6,x7]);
x9:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6,x7,x8]);
x10:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6,x7,x8,x9]);
x_coord_list:[x1,x2,x3,x4,x5,x6,x7,x8];
x_coord_list_ord:sort(x_coord_list);
x_coord_list:[x1,x2,x3,x4,x5,x6,x7,x8,x9,x10];
x_coord_list_ord:sort(x_coord_list);

rank_x:makelist(0,i,1,10);
 
for i:1 thru 10 do (
    for j:1 thru 10 do (
        if x_coord_list[i]=x_coord_list_ord[j] then rank_x[i]:j
    )
);

y1:rand_with_step(30,60,1);
y2:rand_with_prohib(30,60,[y1]);
y3:rand_with_prohib(30,60,[y1,y2]);
y4:rand_with_prohib(30,60,[y1,y2,y3]);
y5:rand_with_prohib(30,60,[y1,y2,y3,y4]);
y6:rand_with_prohib(30,60,[y1,y2,y3,y4,y5]);
y7:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6]);
y8:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6,y7]);
y9:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6,y7,y8]);
y10:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6,y7,y8,y9]);
y_coord_list:[y1,y2,y3,y4,y5,y6,y7,y8,y9,y10];
y_coord_list_ord:sort(y_coord_list);
rank_y:makelist(0,i,1,10);
 
for i:1 thru 10 do (
    for j:1 thru 10 do (
        if y_coord_list[i]=y_coord_list_ord[j] then rank_y[i]:j
    )
);
r_1:(rank_x[1]-rank_y[1])^2;
r_2:(rank_x[2]-rank_y[2])^2;
r_3:(rank_x[3]-rank_y[3])^2;
r_4:(rank_x[4]-rank_y[4])^2;
r_5:(rank_x[5]-rank_y[5])^2;
r_6:(rank_x[6]-rank_y[6])^2;
r_7:(rank_x[7]-rank_y[7])^2;
r_8:(rank_x[8]-rank_y[8])^2;
r_9:(rank_x[9]-rank_y[9])^2;
r_10:(rank_x[10]-rank_y[10])^2;

d_squared:[r_1,r_2,r_3,r_4,r_5,r_6,r_7,r_8,r_9,r_10]; 
/*sum(d_squared)*/
n_1:sum(d_squared[i],i,1,10); 
total_d_squared:n_1;
n:10;
sum_x:x1+x2+x3+x4+x5+x6+x7+x8+x9+x10;
sum_y:y1+y2+y3+y4+y5+y6+y7+y8+y9+y10;
n:10;

sum_xy:(x1*y1)+(x2*y2)+(x3*y3)+(x4*y4)+(x5*y5)+(x6*y6)+(x7*y7)+(x8*y8)+(x9*y9)+(x10*y10);
sum_x_sq:x1^2+x2^2+x3^2+x4^2+x5^2+x6^2+x7^2+x8^2+x9^2+x10^2;
sum_y_sq:y1^2+y2^2+y3^2+y4^2+y5^2+y6^2+y7^2+y8^2+y9^2+y10^2;
alpha:(sum_y*sum_x_sq-sum_x*sum_xy)/(12*sum_x_sq-(sum_x)^2);
beta:(12*sum_xy-sum_x*sum_y)/(12*sum_x_sq-sum_x^2);


ta_1a:(n*sum_xy)-(sum_x*sum_y);
ta_1b:sqrt(((n*sum_x_sq)-sum_x^2)*((n*sum_y_sq)-sum_y^2));
ta_1:decimalplaces((ta_1a/ta_1b),2);
ta_2:decimalplaces(1-((6*total_d_squared)/990),2);
ta_3:decimalplaces(beta,2);
ta_4:decimalplaces(alpha,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Product moment correlation coefficient \(r={@ta_1@}\)</p>
<p>Spearman’s coefficient of rank correlation \(\rho= {@ta_2@}\)</p>
<p>The equation of regression line \(y\) on \(x\) is given by.</p>
<p>\(y={@ta_3*x+ta_4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">This question requires the student to calculate the product moment correlation coefficient , spearman’s coefficient of rank correlation and the equation of the regression line based on a  dataset,</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta_4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1837314822</deployedseed>
    <deployedseed>731942106</deployedseed>
    <deployedseed>239796888</deployedseed>
    <deployedseed>1579221385</deployedseed>
    <deployedseed>1615979120</deployedseed>
    <deployedseed>127152485</deployedseed>
    <deployedseed>979456558</deployedseed>
    <deployedseed>93315998</deployedseed>
    <deployedseed>1155861483</deployedseed>
    <deployedseed>129114381</deployedseed>
    <deployedseed>714217370</deployedseed>
    <deployedseed>682176564</deployedseed>
    <deployedseed>49053088</deployedseed>
    <deployedseed>567261546</deployedseed>
    <deployedseed>1295780418</deployedseed>
    <deployedseed>188897039</deployedseed>
    <deployedseed>809717668</deployedseed>
    <deployedseed>137119141</deployedseed>
    <deployedseed>1284728059</deployedseed>
    <deployedseed>527832986</deployedseed>
    <deployedseed>1383474876</deployedseed>
    <deployedseed>1410631787</deployedseed>
    <deployedseed>95230115</deployedseed>
    <deployedseed>2059193095</deployedseed>
    <deployedseed>1737115731</deployedseed>
    <deployedseed>1199862307</deployedseed>
    <deployedseed>191874849</deployedseed>
    <deployedseed>1465239737</deployedseed>
    <deployedseed>126680560</deployedseed>
    <deployedseed>814850464</deployedseed>
    <deployedseed>1702782711</deployedseed>
    <deployedseed>469014418</deployedseed>
    <deployedseed>1937283064</deployedseed>
    <deployedseed>553897458</deployedseed>
    <deployedseed>965978439</deployedseed>
    <deployedseed>1827479405</deployedseed>
    <deployedseed>626241460</deployedseed>
    <deployedseed>1483153536</deployedseed>
    <deployedseed>1539808420</deployedseed>
    <deployedseed>975402332</deployedseed>
    <deployedseed>953322689</deployedseed>
    <deployedseed>162304965</deployedseed>
    <deployedseed>960958919</deployedseed>
    <deployedseed>1145291025</deployedseed>
    <deployedseed>416446156</deployedseed>
    <deployedseed>2017734300</deployedseed>
    <deployedseed>124779846</deployedseed>
    <deployedseed>2120855289</deployedseed>
    <deployedseed>1512868266</deployedseed>
    <deployedseed>1879547728</deployedseed>
    <deployedseed>302532676</deployedseed>
    <deployedseed>1737937640</deployedseed>
    <deployedseed>709754845</deployedseed>
    <deployedseed>1902119927</deployedseed>
    <deployedseed>1173254929</deployedseed>
    <deployedseed>1318509700</deployedseed>
    <deployedseed>1264156841</deployedseed>
    <deployedseed>800434502</deployedseed>
    <deployedseed>960802323</deployedseed>
    <deployedseed>2108093503</deployedseed>
    <deployedseed>649664543</deployedseed>
    <deployedseed>1828726767</deployedseed>
    <deployedseed>433089181</deployedseed>
    <deployedseed>342935020</deployedseed>
    <deployedseed>1279102917</deployedseed>
    <deployedseed>43566671</deployedseed>
    <deployedseed>1825044993</deployedseed>
    <deployedseed>1015345332</deployedseed>
    <deployedseed>188434999</deployedseed>
    <deployedseed>258214119</deployedseed>
    <deployedseed>679976473</deployedseed>
    <deployedseed>1061448125</deployedseed>
    <deployedseed>721027735</deployedseed>
    <deployedseed>263551481</deployedseed>
    <deployedseed>1753251863</deployedseed>
    <deployedseed>146919405</deployedseed>
    <deployedseed>1265413364</deployedseed>
    <deployedseed>2060935693</deployedseed>
    <deployedseed>744434716</deployedseed>
    <deployedseed>1819745939</deployedseed>
    <deployedseed>407673684</deployedseed>
    <deployedseed>918564364</deployedseed>
    <deployedseed>1331934444</deployedseed>
    <deployedseed>1297528094</deployedseed>
    <deployedseed>1681812063</deployedseed>
    <deployedseed>479991888</deployedseed>
    <deployedseed>1228493372</deployedseed>
    <deployedseed>1202070454</deployedseed>
    <deployedseed>1340998744</deployedseed>
    <deployedseed>1965618550</deployedseed>
    <deployedseed>112953684</deployedseed>
    <deployedseed>618745063</deployedseed>
    <deployedseed>1834842076</deployedseed>
    <deployedseed>594711701</deployedseed>
    <deployedseed>90999968</deployedseed>
    <deployedseed>2114232643</deployedseed>
    <deployedseed>2129241653</deployedseed>
    <deployedseed>1933622854</deployedseed>
    <deployedseed>1184982419</deployedseed>
    <deployedseed>1883847666</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta_4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>