<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Pearson correlation coefficient-Q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Using the data in the table below answer the following questions.</p>
<br>
<table border="solid"  style="width: 80%;">
  <tr>
<th > \(\textbf{X}  \,\text{values}\)</th>
<th >\({@x_1@}\)</th>
<th >\({@x_2@}\)</th>
<th >\({@x_3@}\)</th>
 <th >\({@x_4@}\)</th>
 <th >\({@x_5@}\)</th>
 </tr>
  <tr>
  <th >\(\textbf{Y}  \,\text{values}\)</th>
    <td >\({@y_1@}\)</td>
    <td >\({@y_2@}\)</td>
    <td >\({@y_3@}\)</td>
    <td >\({@y_4@}\)</td>
    <td >\({@y_5@}\)</td>  
  </tr>
</table>
<br>

<table border="solid"  style="width: 80%;">
  <tr>
    <th>\(x-M_x\)</th>
    <th >\(y-M_y\)</th>
    <th >\((x-M_x)^2\)</th>
    <th>\((y-M_y)^2\)</th> 
    <th  style="width: 100%;">\((x-M_x)(y-M_y)\)</th> 
  </tr>
  <tr>
    <td >\({@x1_mx@}\)</td>
    <td >[[input:ans1]] [[validation:ans1]][[feedback:prt1]] </td>
    <td >\({@x1_mx_squared@}\)</td>
    <td >{@decimalplaces(y1_my,2)^2@}</td>
    <td >{@x1_mx*decimalplaces(y1_my,2)@}</td>
  
  </tr>
   <tr>
    <td >[[input:ans2]] [[validation:ans2]][[feedback:prt2]] </td>
    <td >{@decimalplaces(y2_my,2)@}</td>
    <td >{@x2_mx_squared@}</td>
    <td >{@(decimalplaces(y2_my,2))^2@}</td>
    <td >{@x2_mx*decimalplaces(y2_my,2)@}</td>
  
  </tr>
   <tr>
    <td >\({@x3_mx@}\)</td>
    <td >\({@decimalplaces(y3_my,2)@}\)</td>
    <td >[[input:ans3]] [[validation:ans3]][[feedback:prt3]] </td>
    <td >\({@decimalplaces(y3_my,2)^2@}\)</td>
   <td >\({@x3_mx*decimalplaces(y3_my,2)@}\)</td>
  
  </tr>
   <tr>
    <td >\({@x4_mx@}\)</td>
    <td >\({@decimalplaces(y4_my,2)@}\)</td>
    <td >\({@x4_mx_squared@}\)</td>
    <td >[[input:ans4]] [[validation:ans4]][[feedback:prt4]] </td>
    <td >\({@x4_mx*decimalplaces(y4_my,2)@}\)</td>
  
  </tr>
   <tr>
    <td >\({@x5_mx@}\)</td>
    <td >\({@decimalplaces(y5_my,2)@}\)</td>
    <td >\({@x5_mx_squared@}\)</td>
    <td >\({@decimalplaces(y5_my,2)^2@}\)</td>
   <td >[[input:ans5]] [[validation:ans5]][[feedback:prt5]] </td>
  
  <tr>
    <td ></td>
    <td ></td>
    <td>\(\sum (X-M_x)^2={@sum_x_mxsquared@}\)</td>
    <td>\(\sum (Y-M_y)^2=\) [[input:ans7]] [[validation:ans7]][[feedback:prt7]] </td>
    <td >\(\sum ((X-M_x)(Y-M_y))=\) [[input:ans8]] [[validation:ans8]][[feedback:prt8]] </td>
  </tr>
</table>
<br><br>
<p>Taking the formula of finding the pearons correlation denoted as \(\textbf{r}=\displaystyle \frac{\sum ((X-M_x)(Y-M_y))}{\sqrt{ \sum (X-M_x)^2 \cdot \sum (Y-M_Y)^2 }}\)</p>
<p>\(\textbf{r}=\) [[input:ans9]] [[validation:ans9]][[feedback:prt9]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<br>
<table border="solid"  style="width: 80%;">
  <tr>
<th > \(\textbf{X}  \,\text{values}\)</th>
<th >\({@x_1@}\)</th>
<th >\({@x_2@}\)</th>
<th >\({@x_3@}\)</th>
 <th >\({@x_4@}\)</th>
 <th >\({@x_5@}\)</th>
 </tr>
  <tr>
  <th >\(\textbf{Y}  \,\text{values}\)</th>
    <td >\({@y_1@}\)</td>
    <td >\({@y_2@}\)</td>
    <td >\({@y_3@}\)</td>
    <td >\({@y_4@}\)</td>
    <td >\({@y_5@}\)</td>  
  </tr>
</table>

<br>
<table border="solid" style="width: 100%; table-layout: fixed;">
  <tr>
    <th>\( x-M_x \)</th>
    <th>\( y-M_y\)</th>
    <th>\( \left(x-M_x\right)^2\)</th>
    <th>\( \left(y-M_y\right)^2\)</th> 
    <th>\(\left(x-M_x \right) \left(y-M_y\right)\)</th> 
  </tr>
  
  <tr>
    <td>\({@x1_mx@}\)</td>
    <td>\({@decimalplaces(y1_my,2)@}\)</td>
    <td>\({@x1_mx_squared@}\)</td>
    <td>{@decimalplaces(y1_my,2)^2@}</td>
    <td>{@x1_mx*decimalplaces(y1_my,2)@}</td>
  
  </tr>
   <tr>
   <td>\({@x2_mx@}\)</td>
    <td>\({@decimalplaces(y2_my,2)@}\)</td>
    <td>\({@x2_mx_squared@}\)</td>
    <td>{@decimalplaces(y2_my,2)^2@}</td>
    <td>{@x2_mx*decimalplaces(y2_my,2)@}</td>
  
  </tr>
   <tr>
    <td>\({@x3_mx@}\)</td>
    <td>\({@decimalplaces(y3_my,2)@}\)</td>
    <td>\({@x3_mx_squared@}\)</td>
    <td>\({@decimalplaces(y3_my,2)^2@}\)</td>
    <td>\({@x3_mx*decimalplaces(y3_my,2)@}\)</td>
  
  </tr>
   <tr>
    <td>\({@x4_mx@}\)</td>
    <td>\({@decimalplaces(y4_my,2)@}\)</td>
    <td>\({@x4_mx_squared@}\)</td>
    <td>\({@decimalplaces(y4_my,2)^2@}\)</td>
    <td>({@x4_mx*decimalplaces(y4_my,2)@}\)</td>
  
  </tr>
   <tr>
    <td>\({@x5_mx@}\)</td>
    <td>\({@decimalplaces(y5_my,2)@}\)</td>
    <td>\({@x5_mx_squared@}\)</td>
    <td>\({@decimalplaces(y5_my,2)^2@}\)</td>
    <td>\({@x5_mx*decimalplaces(y5_my,2)@}\)</td>
  
  <tr>
    <td></td>
    <td></td>
    <td>\(\sum (x-M_x)^2={@sum_x_mxsquared@}\)</td>
    <td>\(\sum (y-M_y)^2={@sum_y_mysquared@}\) </td>
    <td>\(\sum ((x-M_x)(y-M_y))={@sum_x_y@}\) </td>
  </tr>
</table>

<br><br>
<p>Taking the formula of finding the pearons correlation denoted as \(\textbf{r}=\displaystyle \frac{\sum ((X-M_x)(Y-M_y))}{\sqrt{ \sum (X-M_x)^2 \cdot \sum (Y-M_Y)^2 }}\)</p>
<p>Where,</p>
<p>\(((X-M_x)(Y-M_y))={@sum_x_y@}\)</p>
<p>\(\sum (X-M_x)^2 ={@sum_x_mxsquared@}\)</p>
<p>\(\sum (Y-M_Y)^2={@sum_y_mysquared@}\)</p>
<p> \(\textbf{r}=\displaystyle \frac{({@sum_x_y@})}{\sqrt{ ({@sum_x_mxsquared@}) ({@sum_y_mysquared@}) }}\)</p>
<p> \(\textbf{r}={@r@}\)</p>
<p>Which is a strong positive correlation</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*pearson correlation*/
/*x values*/
x_1:rand(5)+60;
x_2:x_1+2;
x_3:x_2+2;
x_4:x_3+2;
x_5:x_4+2;
/*y values*/
y_1:rand([2.4,2.6,2.8]);
y_2:y_1+0.1;
y_3:y_2+0.2;
y_4:y_3+0.4;
y_5:y_4+0.4;
/*mean of x and y*/
m_x:(x_1+x_2+x_3+x_4+x_5)/5;
m_y:(y_1+y_2+y_3+y_4+y_5)/5;
/*deviations of x*/
x1_mx:x_1-m_x;
x2_mx:x_2-m_x;
x3_mx:x_3-m_x;
x4_mx:x_4-m_x;
x5_mx:x_5-m_x;
/*deviations of y*/
y1_my:y_1-m_y;
y2_my:y_2-m_y;
y3_my:y_3-m_y;
y4_my:y_4-m_y;
y5_my:y_5-m_y;
/*squares of deviation of x*/
x1_mx_squared:(x1_mx)^2;
x2_mx_squared:(x2_mx)^2;
x3_mx_squared:(x3_mx)^2;
x4_mx_squared:(x4_mx)^2;
x5_mx_squared:(x5_mx)^2;
sum_x_mxsquared:x1_mx_squared+x2_mx_squared+x3_mx_squared+x4_mx_squared+x5_mx_squared;
/*squares od deviation of y*/
y1_my_squared:(y1_my)^2;
y2_my_squared:(y2_my)^2;
y3_my_squared:(y3_my)^2;
y4_my_squared:(y4_my)^2;
y5_my_squared:(y5_my)^2;
sum_y_mysquared:decimalplaces(y1_my,2)^2+decimalplaces(y2_my,2)^2+decimalplaces(y3_my,2)^2+decimalplaces(y4_my,2)^2+decimalplaces(y5_my,2)^2;
/*deviations squared*/
x1_y1:(x1_mx)*(y1_my);
x2_y2:(x2_mx)*(y2_my);
x3_y3:(x3_mx)*(y3_my);
x4_y4:(x4_mx)*(y4_my);
x5_y5:(x5_mx)*(y5_my);
sum_x_y:x1_mx*decimalplaces(y1_my,2)+x2_mx*decimalplaces(y2_my,2)+x3_mx*decimalplaces(y3_my,2)+x4_mx*decimalplaces(y4_my,2)+x5_mx*decimalplaces(y5_my,2);
/*model answers*/
sol1:sum_x_mxsquared*sum_y_mysquared;
r:decimalplaces(sum_x_y/(sqrt(sol1)),4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>x values</p>
<ul>
  <li>\({@x_1@}\)</li>
  <li>\({@x_2@}\)</li>
  <li>\({@x_3@}\)</li>
  <li>\({@x_4@}\)</li>
  <li>\({@x_5@}\)</li>
</ul>
<p>y values</p>
<ul>
  <li>\({@y_1@}\)</li>
  <li>\({@y_2@}\)</li>
  <li>\({@y_3@}\)</li>
  <li>\({@y_4@}\)</li>
  <li>\({@y_5@}\)</li>
</ul>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer"> Emmaculate Odhiambo</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">The question focuses on finding the pearsons correlation of given data of x and y values</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>decimalplaces(y1_my,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>x2_mx</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>x3_mx_squared</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>decimalplaces(y4_my,2)^2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>x5_mx*decimalplaces(y5_my,2)</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>sum_y_mysquared</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>sum_x_y</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>r</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>decimalplaces(y1_my,2)</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>x2_mx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>x3_mx_squared</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>decimalplaces(y4_my,2)^2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>x5_mx*decimalplaces(y5_my,2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans7</sans>
        <tans>sum_y_mysquared</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>sum_x_y</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt9</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>r</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt9-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt9-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>427576035</deployedseed>
    <deployedseed>1333584726</deployedseed>
    <deployedseed>659132362</deployedseed>
    <deployedseed>262884923</deployedseed>
    <deployedseed>1913971281</deployedseed>
    <deployedseed>357426967</deployedseed>
    <deployedseed>1457205300</deployedseed>
    <deployedseed>791253620</deployedseed>
    <deployedseed>1688831220</deployedseed>
    <deployedseed>1442762160</deployedseed>
    <deployedseed>1728542541</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>decimalplaces(y1_my,2)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>x2_mx</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>x3_mx_squared</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>decimalplaces(y4_my,2)^2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>x5_mx*decimalplaces(y5_my,2)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>sum_y_mysquared</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>sum_x_y</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>r</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt9</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt9-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>