<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Probability of Selecting Paperback Books from a Shelf</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>\({@n_select@}\) books are taken from a shelf containing \({@n@}\) books of which \({@n_paperback@}\) are paper-back and \({@n_hardcopy@}\) are hard-copy. <br>Find the probability that the set of \({@n_select@}\) books consists of at least \({@2@}\) paper-back.</p>
<p><strong>Note:</strong> Give you answer in exact form or correct to \(3\) decimal places</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the probability that the set of four books consists of at least two paperbacks, we can calculate the probability of selecting sets with two, three, or four paperbacks and then sum them up.</p>
<p>First, let's calculate the total number of ways to choose {@n_select@}  books out of {@n@} , which is given by the combination formula:</p>
<p>\( \displaystyle \text{Total number of selection} = ^{@n@} C_{@n_select@} ={@total_outcomes@} \\
\text{Number of selections at least} \, 2 \, \text{paper back} =\left( ^{@n_paperback@} C_2 \times ^{@n_hardcopy@} C_2\right) + \left(^{@n_paperback@} C_3 \times ^{@n_hardcopy@} C_1 \right) + \left(^{@n_paperback@} C_4 \times ^{@n_hardcopy@} C_0  \right)={@favorable_outcomes@} \\ \displaystyle
\text{P(Number of selections at least} \, 2 \, \text{paper back})= \frac{{@favorable_outcomes@} }{{@total_outcomes@}}={@ta@} \approx {@decimalplaces(ta,3)@}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the combinations function */ 
n_select:4;   /* Number of books to select */ 
aa:rand(5)+2;
n_paperback: n_select+aa;   /* Number of paperback books */ 
bb:rand(5)+4;
n_hardcopy: n_paperback+bb;   /* Number of hard-copy books */ 

n: n_paperback+n_hardcopy; /* Total number of books */
/* Define the total number of outcomes */ 
total_outcomes: binomial(n, n_select);  
/* Define the number of favorable outcomes */ 
favorable_outcomes: binomial(n_paperback, 2) * binomial(n_hardcopy, 2) + binomial(n_paperback, 3) * binomial(n_hardcopy, 1) + binomial(n_paperback, 4)*binomial(n_hardcopy, 0);  
/* Calculate the probability */ 
probability: favorable_outcomes / total_outcomes; 
ta:probability;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@n_select@}\) selected books 
<br>Books in total \({@n@}\) 
<br> \({@n_paperback@}\) are paper-back
<br> \({@n_hardcopy@}\) are hard-copy. 
<p>\(\displaystyle \text{Total number of selection} = ^{@n@} C_{@n_select@} ={@total_outcomes@}  \\
\text{P(Number of selections at least} \, 2 \, \text{paper back})= \frac{{@favorable_outcomes@} }{{@total_outcomes@}}={@ta@} \approx {@decimalplaces(ta,3)@}
\)</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/1bFL6JpUU4VJyNe5cqXcc3xunBxNvigtInAp_ufkruvg/edit</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">This question requires the student to calculate the probability of selecting at least two paperback books from a shelf containing sixteen books, with six paperbacks and ten hardcopies.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct answer 3dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Well done for giving you answer correct to 3 decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>answer with an error margin of 0.0005</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>809924932</deployedseed>
    <deployedseed>693709055</deployedseed>
    <deployedseed>136602816</deployedseed>
    <deployedseed>1316150234</deployedseed>
    <deployedseed>1208316517</deployedseed>
    <deployedseed>647367287</deployedseed>
    <deployedseed>1565333439</deployedseed>
    <deployedseed>183935767</deployedseed>
    <deployedseed>1892108110</deployedseed>
    <deployedseed>786062631</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>