<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Probability Analysis of Drawing White Apples from Urns</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Urn \(I\) contains \({@white_I@}\) white and \({@red_I@}\) red apples and urn \(II\) contains \({@white_II@}\) white and \({@red_II@}\) red apples. One apple is
drawn at random from one of the urns and found to be white. Find the probability that it was drawn
from urn \(I\)</p>
<p><b>Note:</b> Give your answer as a fraction.</p>
<p>\(P(U_1|w)=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>We need to get \(P(U_I|w)\)</p>
<p>By Bayes Theorem, \(\displaystyle P(U_I|w)=\frac{P(U_I)P(w|U_I)}{P(U_I)P(w|U_I)+P(U_{II})P(w|U_{II})}\)</p>
<p>A priori, the two urns are equally likely to be selected. Thus \(P(U_I)=P(U_{II})=\frac{1}{2}.\)</p>
<p>\(P(w|U_I)=P(\)white apple from urn \(\displaystyle I)={@white_I/(white_I+red_I)@}\) </p>
<p>\(P(w|U_{II})=P(\)white apple from urn \(\displaystyle II)={@white_II/(white_II+red_II)@}\) </p>
<p>\(\displaystyle P(U_I|w)= \frac{ \frac{1}{2} \times {@white_I/(white_I+red_I)@}}{(\frac{1}{2} \times {@white_I/(white_I+red_I)@})+(\frac{1}{2} \times {@white_II/(white_II+red_II)@})}={@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>white_I:rand_with_step(2,8,1);
red_I:rand_with_prohib(2,8,[white_I]);
white_II:rand_with_prohib(2,8,[white_I,red_I]);
red_II:rand_with_prohib(2,8,[white_I,red_I,white_II]);
ta:((1/2)*(white_I/(white_I+red_I)))/(((1/2)*(white_I/(white_I+red_I)))+((1/2)*(white_II/(white_II+red_II))));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Urn \(I\) contains \({@white_I@}\) white and \({@red_I@}\) red apples and urn \(II\) contains \({@white_II@}\) white and \({@red_II@}\) red apples.</p>
<p>\(P(U_1|w)={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/1Y_962hRAywGPoE-glatkmbcuOOKxz9UP/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Georg Osang</p>
<p id="description">This question requires the student to calculate the probability of drawing a white apple from Urn I given specific compositions of two urns containing white and red apples respectively.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1155577763</deployedseed>
    <deployedseed>636849754</deployedseed>
    <deployedseed>231798081</deployedseed>
    <deployedseed>1843742595</deployedseed>
    <deployedseed>324091382</deployedseed>
    <deployedseed>1460324934</deployedseed>
    <deployedseed>1394251052</deployedseed>
    <deployedseed>579462721</deployedseed>
    <deployedseed>17030037</deployedseed>
    <deployedseed>558652644</deployedseed>
    <deployedseed>1418932797</deployedseed>
    <deployedseed>450250329</deployedseed>
    <deployedseed>36757050</deployedseed>
    <deployedseed>607567873</deployedseed>
    <deployedseed>1989337168</deployedseed>
    <deployedseed>640565937</deployedseed>
    <deployedseed>600209327</deployedseed>
    <deployedseed>2080662479</deployedseed>
    <deployedseed>965007637</deployedseed>
    <deployedseed>152045102</deployedseed>
    <deployedseed>933270494</deployedseed>
    <deployedseed>714322453</deployedseed>
    <deployedseed>442508298</deployedseed>
    <deployedseed>1288062526</deployedseed>
    <deployedseed>37940691</deployedseed>
    <deployedseed>1104686088</deployedseed>
    <deployedseed>1249513398</deployedseed>
    <deployedseed>1096599581</deployedseed>
    <deployedseed>248781555</deployedseed>
    <deployedseed>1982584264</deployedseed>
    <deployedseed>585510248</deployedseed>
    <deployedseed>1759317786</deployedseed>
    <deployedseed>1846505790</deployedseed>
    <deployedseed>237694389</deployedseed>
    <deployedseed>1929192227</deployedseed>
    <deployedseed>2039541696</deployedseed>
    <deployedseed>370673313</deployedseed>
    <deployedseed>242195716</deployedseed>
    <deployedseed>1783708264</deployedseed>
    <deployedseed>578368202</deployedseed>
    <deployedseed>1414718753</deployedseed>
    <deployedseed>1456248028</deployedseed>
    <deployedseed>317691019</deployedseed>
    <deployedseed>1094038530</deployedseed>
    <deployedseed>1385683467</deployedseed>
    <deployedseed>44203811</deployedseed>
    <deployedseed>1116200352</deployedseed>
    <deployedseed>773237383</deployedseed>
    <deployedseed>1226722004</deployedseed>
    <deployedseed>149182225</deployedseed>
    <deployedseed>2117370418</deployedseed>
    <deployedseed>1861531146</deployedseed>
    <deployedseed>825852902</deployedseed>
    <deployedseed>1777128044</deployedseed>
    <deployedseed>933445499</deployedseed>
    <deployedseed>324328847</deployedseed>
    <deployedseed>757989884</deployedseed>
    <deployedseed>516345038</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>