<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Q2_Vital Statistical forms</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Vital statistics form perhaps the most important branch of statistics as they deal with mankind in the aggregate. </p>
<p>Which among the following is <b>NOT</b> a vital statistic? Select all that apply:</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

[[ if test="member(1, ta_index)" ]]
<p>Warfare rate - This is not a vital statistic. Vital statistics typically focus on events related to births, deaths, marriages, and divorces, while the warfare rate is more related to conflicts and military activities.</p>

[[/ if ]]

[[ if test="member(2, ta_index)" ]]
<p>Elegance rate - This is not a recognized vital statistic. "Elegance rate" seems to be a term that is not commonly used in the context of vital statistics.</p>
[[/ if ]]

[[ if test="member(3, ta_index)" ]]
<p>Sleeping rate - This is not a vital statistic. "Sleeping rate" does not typically refer to a demographic or vital event and seems to be an unrelated concept.</p>
[[/ if ]]

[[ if test="member(4, ta_index)" ]]
<p>Data collection: This is not a vital statistic - "Data collection" is a process, not a vital statistic. Vital statistics are specific numerical measures related to vital events in a population.</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a:2;
correct_answer:rand_selection([
[1, true, "Warfare rate"],
[2, true, "Elegance rate"],
[3, true, "Sleeping rate"],
[4, true, "Data collection"]],a);

b:2;
incorrect_answers:rand_selection([
[5, false, "Birth rate"],
[6, false, "Divorce rate"],
[7, false, "Marriage rate"]], b);

ta:random_permutation(append(correct_answer,incorrect_answers));
ta_index: sort(makelist(x[1], x, ta));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@mcq_correct(ta)@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire<p>
<p id="concept">https://docs.google.com/document/d/1ioECBsnRwApAxz2ZzavQKR2YDOXMKT3ezRA5r5d8Cr4/edit</p>
<p id="reviewer">1</p>
<p id="reviewer">2</p>
<p id="description">The question focuses on outlining the none vital statistics.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>checkbox</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1964495402</deployedseed>
    <deployedseed>795893936</deployedseed>
    <deployedseed>1750505753</deployedseed>
    <deployedseed>2125595196</deployedseed>
    <deployedseed>33538748</deployedseed>
    <deployedseed>114501915</deployedseed>
    <deployedseed>802139450</deployedseed>
    <deployedseed>1186061354</deployedseed>
    <deployedseed>476527268</deployedseed>
    <deployedseed>840744087</deployedseed>
    <deployedseed>78108257</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>