<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>Selecting Balls from Bags</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The content of three bags \({@ball_I_white@}, {@ball_I_red@}\) and \({@ball_I_black@}\) as follows</p>
<p>
<table border="solid">
  <tr>
    <th width="15%", style="text-align:center">Balls</th>
    <th width="15%", style="text-align:center">White</th>
     <th width="15%", style="text-align:center">Black </th>
     <th width="15%", style="text-align:center">Red </th
  </tr>
  <tr>
    <td style="text-align:center"> Bag \(I\)</td>
    <td style="text-align:center">\({@ball_I_white@}\)</td>
    <td style="text-align:center">\({@ball_I_black@}\)</td>
  <td style="text-align:center">\({@ball_I_red@}\)</td>
  </tr>
<tr>
   <td style="text-align:center">Bag \(II\)</td>
    <td style="text-align:center">\({@ball_2_white@}\)</td>
     <td style="text-align:center">\({@ball_2_black@}\)</td>
    <td style="text-align:center">\({@ball_2_red@}\)</td>
  </tr>
  <tr>
    <td style="text-align:center">Bag \(III\)</td>
    <td style="text-align:center">\({@ball_3_white@}\)</td>
    <td style="text-align:center">\({@ball_3_black@}\)</td>
    <td style="text-align:center">\({@ball_3_red@}\)</td>
  </tr>
</table>
</p>
<br>
<p>Note: Give your answer as a fraction, or to \(2\) decimal places.</p>
<p>A bag is chosen at random and from it two balls are drawn at random. The two balls are found
to be one white and one black. What is the probability that they are from the second bag?</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Let \(B_1,B_2,B_3\) denote the events that the bags \(I, II\) and \(III\) are chosen respectively.</p>
<p>Let \(A\) be the event that the two balls taken from the selected bag are white and black.</p>
<p>Then,</p>
<p>\(P(B_1)=P(B_2)=P(B_3)=\displaystyle \frac{1}{3}\)</p>
<p>Then \(P(A|B_1)=\displaystyle\frac {^{@ball_I_white@}C_1 \times ^{@ball_I_black@}C_1}{^{@is_sum@}C_2}={@p_a_b1@}\):<br>
We get \(^{@ball_I_white@}C_1 \times ^{@ball_I_black@}C_1\) because there are \({@ball_I_white@}\) white balls in bag \(I\) to choose \(1\) ball from for \(^{@ball_I_white@}C_1\) choices to achieve that, and \({@ball_I_black@}\) black balls in bag \(I\) to choose \(1\) from for \(^{@ball_I_black@}C_1\) choices to achieve that, out of a total of \(^{@is_sum@}C_2\) choices to choose 2 balls.</p>
<p>The same procedure is repeated for \(P(A|B_2)=\displaystyle\frac {^{@ball_2_white@}C_1 \times ^{@ball_2_black@}C_1}{^{@is_sum@}C_2}={@p_a_b2@}, P(A|B_3)=\displaystyle\frac {^{@ball_3_white@}C_1 \times ^{@ball_3_black@}C_1}{^{@is_sum@}C_2}={@p_a_b3@}\)</p>
<br>
<p>Finally, \(P(B_2|A)=\displaystyle\frac {P(B_2)P(A|B_2)}{\sum_{i=1}^{3} P(B_i)P(A|B_i)}=\displaystyle\frac {\displaystyle \frac{1}{3}\times {@p_a_b2@}}{(\displaystyle \frac{1}{3}\times {@p_a_b1@})+(\displaystyle \frac{1}{3}\times {@p_a_b2@})+(\displaystyle \frac{1}{3}\times {@p_a_b3@})}=\displaystyle{@p_b2_a@}\) or \({@sol1@}\) in float form</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>0.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*choosing random numbers that give equal sum*/
is_sum: rand(7) + 5; 
/*is_sum: 6;*/ 
list_num: []; 
sum_nums: 0;  
while sum_nums # is_sum do (    
    list_num: makelist((rand(4) + 1), i, 1, 3),     
    sum_nums: sum(list_num[i], i, 1, 3) );  
list_num;

/*items in the bags*/
ball_I_white:list_num[1];
ball_I_red:list_num[2];
ball_I_black:list_num[3];
/*phase 2*/
ball_2_white:list_num[2];
ball_2_red:list_num[3];
ball_2_black:list_num[1];
/*phase 3*/
ball_3_white:list_num[3];
ball_3_red:list_num[1];
ball_3_black:list_num[2];
/*worked solution variables*/
sum_combination:binomial(is_sum, 2);
p_a_b1:(binomial(ball_I_white,1)*binomial(ball_I_black,1))/sum_combination;
p_a_b2:(binomial(ball_2_white,1)*binomial(ball_2_black,1))/sum_combination;
p_a_b3:(binomial(ball_3_white,1)*binomial(ball_3_black,1))/sum_combination;

s1:((1/3)*p_a_b2);
s2:(((1/3)*p_a_b1)+((1/3)*p_a_b2)+((1/3)*p_a_b3));
p_b2_a:s1/s2;
sol1:decimalplaces(p_b2_a,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>sum: {@is_sum@}</p>
<p> list of balls: {@list_num@}</p>
<p>\(P(B_2|A)=\displaystyle\frac {P(B_2)P(A|B_2)}{\sum_{i=1}^{3} P(B_i)P(A|B_i)}=\displaystyle\frac {\displaystyle \frac{1}{3}\times {@p_a_b2@}}{(\displaystyle \frac{1}{3}\times {@p_a_b1@})+(\displaystyle \frac{1}{3}\times {@p_a_b2@})+(\displaystyle \frac{1}{3}\times {@p_a_b3@})}=\displaystyle{@p_b2_a@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Georg Osang</p>
<p id="description">There are 3 bags with various numbers of white, black and red balls. A bag is chosen at random and from it two balls are drawn at random. The two balls are found to be one white and one black, the questions requires you to find probability that they are from the second bag.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans> sol1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>p_b2_a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for 2dp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sol1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks if the answer has more dp</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>p_b2_a</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is partially correct because it's not up to the correct number of dp in the instructions</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1442715049</deployedseed>
    <deployedseed>716055251</deployedseed>
    <deployedseed>1061854559</deployedseed>
    <deployedseed>1140866566</deployedseed>
    <deployedseed>1114854360</deployedseed>
    <deployedseed>522015057</deployedseed>
    <deployedseed>29834417</deployedseed>
    <deployedseed>642197661</deployedseed>
    <deployedseed>440811653</deployedseed>
    <deployedseed>584108329</deployedseed>
    <deployedseed>1259874611</deployedseed>
    <deployedseed>652038407</deployedseed>
    <deployedseed>769305834</deployedseed>
    <deployedseed>1022490689</deployedseed>
    <deployedseed>643603141</deployedseed>
    <deployedseed>2124752837</deployedseed>
    <deployedseed>421525700</deployedseed>
    <deployedseed>245407616</deployedseed>
    <deployedseed>770301874</deployedseed>
    <deployedseed>532592790</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value> sol1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>