<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>test- Box Plot matched with dotplots</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The following exercise shows stacked dot plots for four data sets . Match each dotplot to its respective boxplot.</p>

<p>min={@lim1@}, max={@lim2@}, median={@med@}, LQ={@lq@}, UQ={@uq@}</p>
<p>min1={@lim11@}, max1={@lim21@}, median1={@med11@}, LQ1={@lq1@}, UQ1={@uq1@}</p>
<p>min3={@lim33@}, max3={@lim31@}, median3={@med33@}, LQ3={@lq3@}, UQ3={@uq3@}</p>
<p>{@lx@}</p><p>{@lx1@}</p><p>{@lx3@}</p>


<p>[[jsxgraph]]var lx = {#lx#};var lim1 ={#lim1#};var lim2 = {#lim2#};var med = {#med#};var lq = {#lq#};var uq = {#uq#};
var p = [];
//create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#},{#ymax#},{#xmax#},{#ymin#}],
  keepaspectratio: true,axis: true,pan:{enabled:false}
});

//Create points of boxplot1
var d1 = board.create('line', [[lim1, 1],[lim1, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d2 = board.create('line', [[lim2, 1],[lim2, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d3 = board.create('line', [[med, 1],[med, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d4 = board.create('line', [[lq, 1],[lq, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,
fixed:true
})
var d5 = board.create('line', [[uq, 1],[uq, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d6 = board.create('line', [[uq, 1],[lq, 1]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d7 = board.create('line', [[uq, 2],[lq, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d8 = board.create('line', [[lim1, 1.5],[lq, 1.5]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})
var d9 = board.create('line', [[lim2, 1.5],[uq, 1.5]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true
})

//create points for boxplot2
var lx1 = {#lx1#};
var lim11 ={#lim11#};
var lim21 = {#lim21#};
var med11 = {#med11#};
var lq1 = {#lq1#};
var uq1 = {#uq1#};

var d11 = board.create('line', [
[lim11, 1],
[lim11, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})

var d22 = board.create('line', [
[lim21, 1],
[lim21, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})

var d33 = board.create('line', [
[med11, 1],
[med11, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})

var d44 = board.create('line', [
[lq1, 1],
[lq1, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})
var d55 = board.create('line', [
[uq1, 1],
[uq1, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})
var d66 = board.create('line', [
[uq1, 1],
[lq1, 1]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})
var d77 = board.create('line', [
[uq1, 2],
[lq1, 2]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})
var d88 = board.create('line', [
[lim11, 1.5],
[lq1, 1.5]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})
var d99 = board.create('line', [
[lim21, 1.5],
[uq1, 1.5]
], {
straightFirst: false,
straightLast: false,
strokeWidth: 2,
fixed:true
})

//create boxplot3 
var lx3 = {#lx3#};var lim33 ={#lim33#};var lim31 = {#lim31#};var med33 = {#med33#};var lq3 = {#lq3#};var uq3 = {#uq3#};

var d113= board.create('line', [[lim33, 1],[lim33, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d223 = board.create('line', [[lim31, 1],[lim31, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d333 = board.create('line', [[med33, 1],[med33, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d443= board.create('line', [[lq3, 1],[lq3, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d553= board.create('line', [[uq3, 1],[uq3, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d663= board.create('line', [[uq3, 1],[lq3, 1]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d773= board.create('line', [[uq3, 2],[lq3, 2]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d883= board.create('line', [[lim33, 1.5],[lq3, 1.5]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

var d993= board.create('line', [[lim31, 1.5],[uq3, 1.5]], {
straightFirst: false,straightLast: false,strokeWidth: 2,fixed:true})

//create stacked dotplots
//Create points of Data Set 1
var lx2 = {#dorx2#};var ly2 = {#dory2#};
var i = 0;
for (i = 0; i <= lx2.length - 1; i++) {
  p[i] = board.create('point', [lx2[i], ly2[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

//Create Label for Data 2
var txt = board.create('text', [{#xmin#}+1, {#maxfreq2#}/2+2, "B"], {
  fontSize: 30,fixed: true
});

//Create points of Data Set 1
var lx1 = {#dorx1#};var ly1 = {#dory1#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',name: '',size: 2,fixed:true });
}

//Create Label for Data 1
var txt = board.create('text', [{#xmin#}+1, {#maxfreq1#}/2+{#maxfreq2#} + 5, "A"], {
  fontSize: 30,fixed: true
});
[[/jsxgraph]]
</p>
<p>DotplotA matches with what boxplot?Enter the number of the matching boxplot=[[input:ans1]] [[validation:ans1]]</p>
<p>DotplotB matches with what boxplot?Enter the number of the matching boxplot=[[input:ans2]] [[validation:ans2]]</p>
<p>DotplotC matches with what boxplot?Enter the number of the matching boxplot=[[input:ans3]] [[validation:ans3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[mean1:rand(5)+10;
sd1:rand(3)+2;
ds1:rand(900)+100;
list:makelist(float(quantile_normal(rand(1.0),mean1,sd1)),i,1,ds1);
lx:sort(list);
lim1:lx[1];
lim2:lx[ds1];
med:median(lx);
lq:quantile (lx,1/4),numer;
uq:quantile (lx, 3/4),numer;
xmin:lx[1]-2;
xmax:lx[ds1]+2;
ymin:-4;
ymax:8;

aa:rand(2.0);
bb:rand(3.0);
mean2:mean1+(rand(5)+10)*aa;
sd2:sd1*bb;
ds2:ds1;
list2:makelist(float(quantile_normal(rand(1.0),mean2,sd2)),i,1,ds2);
lx1:sort(list2);
lim11:lx1[1];
lim21:lx1[ds2];
med11:median(lx1);
lq1:quantile (lx1,1/4),numer;
uq1:quantile (lx1, 3/4),numer;
xmin:lx[1]-2;
xmax:lx[ds2]+2;
ymin:-4;
ymax:8;

aa1:rand(3.0);bb1:rand(3.0);mean2:mean1+(rand(5)+10)*aa1;sd2:sd1*bb1;ds2:ds1;
list2:makelist(float(quantile_normal(rand(1.0),mean2,sd2)),i,1,ds2);lx1:sort(list2);
lim11:lx1[1];lim21:lx1[ds2];med11:median(lx1);lq1:quantile (lx1,1/4),numer;
uq1:quantile (lx1, 3/4),numer;xmin:lx[1]-2;xmax:lx[ds2]+2;ymin:-1;ymax:8;

aa3:rand(4.0);bb3:rand(6.0);mean3:mean2+(rand(7)+10)*aa3;sd3:sd1*bb3;ds3:ds2;
list3:makelist(float(quantile_normal(rand(1.0),mean3,sd3)),i,1,ds3);lx3:sort(list3);
lim33:lx3[1];lim31:lx3[ds3];med33:median(lx3);lq3:quantile (lx3,1/4),numer;
uq3:quantile (lx3, 3/4),numer;xmin3:lx3[1]-2;xmax3:lx3[ds3]+2;ymin:-1;ymax:8;

s_aa:rand([-1,0,1]);
s_bb:rand([1/3,1/2,1,2,3]);
dun1:makelist(round(float(quantile_normal(rand(1.0),mean1,sd1))),i,1,ds1);
mdun1:float(mean(dun1));
r_mdun1:round(mdun1);
sddun1:round(float(std1(dun1)));
dorx1:sort(dun1);
dun2:makelist(round(float(quantile_normal(rand(1.0),mean2,sd2))),i,1,ds2);
mdun2:float(mean(dun2));
mdiff:float(mdun1-mdun2);
dun2: if s_aa =0 then dun2+mdiff else dun2;
mdun2:float(mean(dun2));
r_mdun2:round(mdun2);
sddun2:round(float(std1(dun2)));
dorx2:sort(dun2);

maxfreq1:0;
k:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then maxfreq1:k) else k:1;
maxfreq2:0;
k:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (k:k+1, if maxfreq2 < k then maxfreq2:k) else k:1;

xmin:if dorx1[1]<=dorx2[1] then dorx1[1]-2 else dorx2[1]-2;
xmax:if dorx1[ds1]>=dorx2[ds2] then dorx1[ds1]+2+(dorx1[ds1]-dorx1[1])/1.5 else dorx2[ds2]+2+(dorx2[ds2]-dorx2[1])/1.5;
ymin:-4;
ymax: maxfreq1+maxfreq2+5;

j:1;
dory1:makelist(0,i,1,ds1);
dory1[1]:1+maxfreq2+3;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory1[i]:j+maxfreq2+3) else (j:1, dory1[i]:j+maxfreq2+3);
j:1;
dory2:makelist(0,i,1,ds2);
dory2[1]:1;
for i:2 thru ds2 do if dorx2[i]=dorx2[i-1] then (j:j+1, dory2[i]:j) else (j:1, dory2[i]:j);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@list@}
min={@lim1@}, max={@lim2@}, median={@med@}, LQ={@lq@}, UQ={@uq@}

{@lx@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.5</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>