<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>The Spearman’s rank correlation coefficient_Q2 (copy)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The table below shows marks obtained by eight students in a mathematics and physics exam. Calculate the Spearman’s coefficient of rank correlation.</p>
<p><b>Note:</b> Give your answers correct to \(2\) decimal places.</p>
<table border="solid">
  <tr>
      <th width="16%", style="text-align:center">Mathematics</th>
       <th width="10%", style="text-align:center">\({@x1@}\)</th>
       <th width="10%", style="text-align:center">\({@x2@}\)</th>
       <th width="10%", style="text-align:center">\({@x3@}\)</th>
      <th width="10%", style="text-align:center">\({@x4@}\)</th>
       <th width="10%", style="text-align:center">\({@x5@}\)</th>
      <th width="10%", style="text-align:center">\({@x6@}\)</th>
       <th width="10%", style="text-align:center">\({@x7@}\)</th>
       <th width="10%", style="text-align:center">\({@x8@}\)</th>
  </tr>
   <tr>
      <th width="16%", style="text-align:center">
Physics 

</th>
       <th width="10%", style="text-align:center">\({@y1@}\)</th>
       <th width="10%", style="text-align:center">\({@y2@}\)</th>
       <th width="10%", style="text-align:center">\({@y3@}\)</th>
      <th width="10%", style="text-align:center">\({@y4@}\)</th>
       <th width="10%", style="text-align:center">\({@y5@}\)</th>
      <th width="10%", style="text-align:center">\({@y6@}\)</th>
       <th width="10%", style="text-align:center">\({@y7@}\)</th>
       <th width="10%", style="text-align:center">\({@y8@}\)</th>
  </tr>
  
 </table>
<p></p>

<p>a. Spearman’s coefficient of rank correlation \(\rho=\) [[input:ans1]]. [[validation:ans1]][[feedback:prt1]]</p>
<p>b. Examine the calculated Spearman’s coefficient of rank correlation \((\rho )\) for the given data in mathematics and physics. How would you interpret this correlation coefficient in the context of the marks obtained by the eight students?</p><p> [[input:ans2]].[[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
a. 
<table border="solid">
  <tr>
      <th width="10%", style="text-align:center">Mathematics</th>
       <th width="10%", style="text-align:center">Physics</th>
       <th width="15%", style="text-align:center">Rank(Mathematics) </th>
       <th width="15%", style="text-align:center">Rank(Physics)</th>
      <th width="40%", style="text-align:center">d \(=\) Rank(Mathematics)\(-\)Rank(Physics)</th>
       <th width="10%", style="text-align:center">\(\text{d}^2\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x1@}\)</th>
    <th width="10%", style="text-align:center">\({@y1@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[1]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[1]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[1]-rank_y[1]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[1]-rank_y[1])^2@}\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x2@}\)</th>
    <th width="10%", style="text-align:center">\({@y2@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[2]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[2]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[2]-rank_y[2]@}\)</th>
    <th width="10%", style="text-align:center">\({@(rank_x[2]-rank_y[2])^2@}\)</th>
  </tr>
  <tr>
    <th width="10%", style="text-align:center">\({@x3@}\)</th>
    <th width="10%", style="text-align:center">\({@y3@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_x[3]@}\)</th>
    <th width="15%", style="text-align:center">\({@rank_y[3]@}\)</th>
    <th width="40%", style="text-align:center">\({@rank_x[3]-rank_y[3]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[3]-rank_y[3])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x4@}\)</th>
       <th width="10%", style="text-align:center">\({@y4@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[4]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[4]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[4]-rank_y[4]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[4]-rank_y[4])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x5@}\)</th>
       <th width="10%", style="text-align:center">\({@y5@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[5]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[5]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[5]-rank_y[5]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[5]-rank_y[5])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x6@}\)</th>
       <th width="10%", style="text-align:center">\({@y6@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[6]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[6]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[6]-rank_y[6]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[6]-rank_y[6])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x7@}\)</th>
       <th width="10%", style="text-align:center">\({@y7@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[7]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[7]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[7]-rank_y[7]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[7]-rank_y[7])^2@}\)</th>
  </tr>
  <tr>
       <th width="10%", style="text-align:center">\({@x8@}\)</th>
       <th width="10%", style="text-align:center">\({@y8@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_x[8]@}\)</th>
       <th width="15%", style="text-align:center">\({@rank_y[8]@}\)</th>
       <th width="40%", style="text-align:center">\({@rank_x[8]-rank_y[8]@}\)</th>
       <th width="10%", style="text-align:center">\({@(rank_x[8]-rank_y[8])^2@}\)</th>
  </tr>
  
 </table>

<p>\(\displaystyle \sum d^{2}={@total_d_squared@}\)</p>

<p>\(n=\) no. of ranks \(=8\)</p>
<p>Spearman rank correlation  coefficient \(\displaystyle r=1- \frac{ 6\sum d^2}{n(n^2-1)}\)</p>
<p>\( \displaystyle r=1- \frac{6({@total_d_squared@})}{8(64-1)}\)</p>
<p>\( r=1-\frac{{@6*total_d_squared@}}{504}=1-{@(6*total_d_squared)/504@}\)</p>\(={@ta_1@}\)</p>
b.
<p>
[[ if test='godfrey' ]]
<br>Correct to two decimal places, the value of the coefficient is \({@ta_1@}\) which is \(< 0\). So we can say that  <b>there is a moderate negative rank correlation between the marks in mathematics and physics</b>. Thus, we can conclude that higher values of the variable  Mathematics do not tend to be associated with higher values of the variable physics and vice versa.

[[else]]

<br>Correct to two decimal places, the value of the coefficient is \({@ta_1@}\) which is \(>  0\). So we can say that  <b>there is a strong positive rank correlation between the marks in mathematics and physics</b>. Thus, we can conclude that higher values of the variable  Mathematics  tend to be associated with higher values of the variable physics and vice versa.
[[/if]]
[[ if test='godfrey' ]]
<br>Correct to two decimal places, the value of the coefficient is \({@ta_1@}\) which is \(>  0\). So we can say that  <b>there is a strong positive rank correlation between the marks in mathematics and physics</b>. Thus, we can conclude that higher values of the variable  Mathematics  tend to be associated with higher values of the variable physics and vice versa.
[[else]]
<br>Correct to two decimal places, the value of the coefficient is \({@ta_1@}\) which is \( =0\). So we can say that  <b>there is no significant rank correlation between the marks in mathematics and physics</b>.
[[/if]]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[x1:rand_with_step(30,80,1);
x2:rand_with_prohib(30,80,[x1]);
x3:rand_with_prohib(30,80,[x1,x2]);
x4:rand_with_prohib(30,80,[x1,x2,x3]);
x5:rand_with_prohib(30,80,[x1,x2,x3,x4]);
x6:rand_with_prohib(30,80,[x1,x2,x3,x4,x5]);
x7:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6]);
x8:rand_with_prohib(30,80,[x1,x2,x3,x4,x5,x6,x7]);
x_coord_list:[x1,x2,x3,x4,x5,x6,x7,x8];
x_coord_list_ord:sort(x_coord_list);

rank_x:makelist(0,i,1,8);
 
for i:1 thru 8 do (
    for j:1 thru 8 do (
        if x_coord_list[i]=x_coord_list_ord[j] then rank_x[i]:j
    )
);
y1:rand_with_step(30,60,1);
y2:rand_with_prohib(30,60,[y1]);
y3:rand_with_prohib(30,60,[y1,y2]);
y4:rand_with_prohib(30,60,[y1,y2,y3]);
y5:rand_with_prohib(30,60,[y1,y2,y3,y4]);
y6:rand_with_prohib(30,60,[y1,y2,y3,y4,y5]);
y7:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6]);
y8:rand_with_prohib(30,60,[y1,y2,y3,y4,y5,y6,y7]);
y_coord_list:[y1,y2,y3,y4,y5,y6,y7,y8];
y_coord_list_ord:sort(y_coord_list);
rank_y:makelist(0,i,1,8);
 
for i:1 thru 8 do (
    for j:1 thru 8 do (
        if y_coord_list[i]=y_coord_list_ord[j] then rank_y[i]:j
    )
);
r_1:(rank_x[1]-rank_y[1])^2;
r_2:(rank_x[2]-rank_y[2])^2;
r_3:(rank_x[3]-rank_y[3])^2;
r_4:(rank_x[4]-rank_y[4])^2;
r_5:(rank_x[5]-rank_y[5])^2;
r_6:(rank_x[6]-rank_y[6])^2;
r_7:(rank_x[7]-rank_y[7])^2;
r_8:(rank_x[8]-rank_y[8])^2;

d_squared:[r_1,r_2,r_3,r_4,r_5,r_6,r_7,r_8]; 


/*sum(d_squared)*/
n:sum(d_squared[i],i,1,8); 
total_d_squared:n;
ta_1:decimalplaces(1-((6*total_d_squared)/504),2);
ta_2a:ta_1>0;
ta_2b:ta_1<0;
ta_2c:ta_1=0;

ta_2: if ta_2a then "there is a strong positive rank correlation between the marks in mathematics and physics" elseif ta_2b then "there is a moderate negative rank correlation between the marks in mathematics and physics" elseif ta_2c then "there is no significant rank correlation between the marks in mathematics and physics";
godfrey:oddp(rand(5));

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Spearman’s coefficient of rank correlation \(r={@ta_1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">This question requires the student to calculate the Spearman's rank correlation coefficient between two sets of variables (X and Y) and determine the coefficient of determination. The question also asks the student to interpret the results in terms of the ranks of the variables.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>40</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1497565443</deployedseed>
    <deployedseed>836901860</deployedseed>
    <deployedseed>1442893574</deployedseed>
    <deployedseed>1806828755</deployedseed>
    <deployedseed>412643948</deployedseed>
    <deployedseed>1017840458</deployedseed>
    <deployedseed>986238665</deployedseed>
    <deployedseed>199324972</deployedseed>
    <deployedseed>1726838093</deployedseed>
    <deployedseed>433947272</deployedseed>
    <deployedseed>1317161982</deployedseed>
    <deployedseed>197570647</deployedseed>
    <deployedseed>526600430</deployedseed>
    <deployedseed>1515048809</deployedseed>
    <deployedseed>1523409664</deployedseed>
    <deployedseed>864009752</deployedseed>
    <deployedseed>1395265209</deployedseed>
    <deployedseed>505656370</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>