<?xml version='1.0' encoding='UTF-8'?>
<quiz><question type="stack">
    <name>
      <text>tomato game</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>In each block below, choose six treatments you want to include in that block and whether you want to randomise your choices in each block. If you are using a phone, you may need to slide your finger to the right to see the second block.</p>

<p>When you've finished designing your experiment, click the 'Check' button. You'll then receive feedback and will have the opportunity to export your data.</p>

<p>Do you want to randomise your treatment choices in each block? [[input:ans13]] [[validation:ans13]]</p>

<br>

<p>
<table border="1">
  <tr align="center">
    <th width="300px">{@factors_info[2][1][2]@}</th>
    <th width="300px">{@factors_info[2][2][2]@}</th>
  </tr>
  <tr>
    <th>[[input:ans1]]</th>
    <th>[[input:ans7]]</th>
  </tr>
  <tr>
    <th>[[input:ans2]]</th>
    <th>[[input:ans8]]</th>
  </tr>
  <tr>
    <th>[[input:ans3]]</th>
    <th>[[input:ans9]]</th>
  </tr>
  <tr>
    <th>[[input:ans4]]</th>
    <th>[[input:ans10]]</th>
  </tr>
  <tr>
    <th>[[input:ans5]]</th>
    <th>[[input:ans11]]</th>
  </tr>
  <tr>
    <th>[[input:ans6]]</th>
    <th>[[input:ans12]]</th>
  </tr>
</table>
</p>

<p>[[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans4]] [[validation:ans5]] [[validation:ans6]] [[validation:ans7]] [[validation:ans8]] [[validation:ans9]] [[validation:ans10]] [[validation:ans11]] [[validation:ans12]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* ------------- INFORMATION FROM QUESTION AUTHORS */

/* ["Year", "Block", "Optional factor 1", "Optional factor 2", ...] */
factor_names:["Year", "Block", "Variety", "Fertiliser", "Weeding"];

/* Factor Info Structure: [
    [F1-V1 effect, "Factor 1 - Value 1"], 
    [F1-V2 effect, "Factor 1 - Value 2"], 
    ...
] */
factors_info:[
    [
        [0, "1"], 
        [200, "2"]
    ],
    [
        [0, "Block A"], 
        [100, "Block B"]
    ],
    [
        [10, "Short Head and Compact"], 
        [-10, "Long Head"]
    ],
    [
        [200, "Yes"], 
        [-200, "No"]
    ],
    [
        [100, "Partial"], 
        [-100, "Complete"]
    ]
];

/* Interactions Info Structure: [
    "Factor 1 name (m varieties)", 
    "Factor 2 name (n varieties)", 
    Factor interaction matrix (n x m matrix)
] */
interactions_info:[
    ["Fertiliser", "Weeding", [
        [0,200],
        [0,0]
    ]]
];

/* INITIAL CONDITIONS */
year_coord:1;
base_yield:800;
error_mean:0;
error_sd:75;

/* ------------- END AUTHOR INFO */


/* DATA SET DIMENSIONS */
number_factors: length(factors_info);
number_treatments: product(length(factors_info[i]), i, 3, length(factors_info));


/* CONTRIBUTION TO POSITION FROM YEAR */
number_blocks: length(factors_info[2]);
year_contribution: number_blocks*number_treatments*(year_coord-1);


/* FUNCTIONS */
/* Function to calculate the expected yield for a given data row */
/* Input: list of coordinates for each factor */
/* Output: expected yield for the given combination of factors with interactions */
effect_calculator(coords):=block(
    [
        factors_info:factors_info,
        interactions_info:interactions_info,
        counter_1,
        effects_sum,
        interaction_effect_sum,
        interaction,
        factor1_index,
        factor2_index,
        factor1_var_index,
        factor2_var_index,
        interaction_effect
    ],
    single_effects_sum: sum(factors_info[counter_1][coords[counter_1]][1], counter_1, 1, length(factors_info)),
    interaction_effect_sum: 0,
    for interaction in interactions_info do block(       
        factor1_index: sublist_indices(factor_names, lambda([x], x=interaction[1]))[1],
        factor2_index: sublist_indices(factor_names, lambda([x], x=interaction[2]))[1],
        factor1_var_index: coords[factor1_index],
        factor2_var_index: coords[factor2_index],
        interaction_effect: interaction[3][factor1_var_index][factor2_var_index],
        interaction_effect_sum: interaction_effect_sum + interaction_effect
    ),
    return(single_effects_sum + interaction_effect_sum)
);

/* Function to create the labels for a given data row */
/* Input: List of coordinates for each factor */
/* Output: list of labels for the given combination of factors without interactions */
create_label(coords):=block(
    [
        factors_info:factors_info,
        number_factors:number_factors,
        i,
        label_maker
    ],
    label_maker:makelist(factors_info[i][coords[i]][2],i,1,number_factors),
    return(label_maker)
);

/* Function to map position of item in list to coordinates in nested data list */
/* Input: position of item in list, number between 1 and the number of all possible environments
(i.e. *number of years* x *number of blocks* x *number of student specified treatments*) */
/* Output: relevant coordinates in factors_info list */
pos_to_coords(num):=block(
    [factors_info:factors_info, 
    number_factors:number_factors,
    number_rows:number_rows,
    factor_lengths,
    clock],
    factor_lengths: makelist(length(factors_info[i]),i,1,number_factors),
    /* How regularly is each factor required to change value? */
    clock:reverse(makelist(
        product(factor_lengths[j], j, number_factors+1-i, number_factors),
        i,
        0,
        number_factors-1
    )),
    return(makelist(mod(floor((num-1)/clock[i]),factor_lengths[i])+1, i, 1, number_factors))
);

/* Function to convert a list of strings into a single string listing each entry, broken up by "-" */
/* Input: List of string values */
/* Output: Single string */
list_to_str(list):=block(
    sentence:pop(list),
    for word in list do (
        sentence: sconcat(sentence, " - ", word)
    ),
    return(sentence)
);


/* QUESTION VARIABLES */
/* Randomize the order of the packets within the blocks? */
shuffle: [
    [1, true, "Yes"],
    [2, false, "No"]
];
/* Treatment dropdown list generation */
/* Generate base info for options in dropdown */
treatment_index_list: makelist(i, i, 1, number_treatments);
treatment_ma_list: makelist(true, i, 1, number_treatments);
treatment_name_list: makelist(
    list_to_str(append([sconcat("T", i)], makelist(sconcat(
                factor_names[j],
                ": ",
                factors_info[j][pos_to_coords(i)[j]][2]
                )
            , j, 3, number_factors))), 
    i, 1, number_treatments
);
/* Generate final list of options for dropdown */
treatments: makelist(
    [treatment_index_list[i], treatment_ma_list[i], treatment_name_list[i]], 
    i, 1, number_treatments
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Dan Kelly</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="description"></p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text/>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text/>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text/>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>dropdown</type>
      <tans>shuffle</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>dropdown</type>
      <tans>treatments</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* ---------- CALCULATIONS FOR EXPORTING DATA */
/* Gather answer list of treatments chosen from table of dropdowns - Hard-coded */
block_a_ans: [ans1,ans2,ans3,ans4,ans5,ans6];
if ans13 = 1 then block_a_ans: random_permutation(block_a_ans);
block_b_ans: [ans7,ans8,ans9,ans10,ans11,ans12];
if ans13 = 1 then block_b_ans: random_permutation(block_b_ans);

/* List the chosen treatements:
Each possible growing environment (including year and block) is indexed sequentially
The student specifies the treatment, however there are 4 possible growing environments for any given treatment
(assuming 2 years and 2 blocks).
To ensure the chosen block is taken into account within this indexing, 
we add the total number of user specified "treatments" to the growing environments in block B 
The contribution of the year is dealt with seperately, since it applies to all the packets
(see the definition of "treatment_coords") */
chosen_treatments: append(
    block_a_ans,
    block_b_ans + number_treatments*[1,1,1,1,1,1]
);

/* Answer list with all factor and block combinations for testing */
/*
chosen_treatments: append(
    makelist(i, i, 1, number_blocks*number_treatments)
)
*/

/* Calculate yield of each option with randomised error */
data_set: makelist(0, i, 1, length(chosen_treatments));
for i: 1 step 1 thru length(chosen_treatments) do block(
    treatment_coords: pos_to_coords(year_contribution + chosen_treatments[i]),   
    data_set[i]: append(create_label(treatment_coords), [max(float(base_yield + effect_calculator(treatment_coords) + error_sd*quantile_normal(rand(1.0),error_mean,1)), 0)])
);

/* Add headers */
data_set: append([append(factor_names, ["Yield"])], data_set);


/* ---------- CALCULATIONS FOR SPECIFIC FEEDBACK NODES */
/* Node 3: Blocking */
blocking: if setify(block_a_ans) = setify(block_b_ans) then 1 else 0;

/* Node 4: Full Factorial */
full_factorial: if cardinality(setify(append(block_a_ans,block_b_ans))) = number_treatments then 1 else 0;

/* Node 5: Interaction */
/* Determine the coordinates of the treatments chosen in blocks A and B */
coord_list_a: makelist(pos_to_coords(block_a_ans[x]), x, 1, length(block_a_ans));
coord_list_b: makelist(pos_to_coords(block_b_ans[x]), x, 1, length(block_b_ans));
full_interactions:0;

/* Parse through the possible pairs of factors and determine whether either block A or block B (or both) contain all possible
pairs of values of these two factors, if so increase the recorded number of "full_interactions" */
for i: 3 thru number_factors-1 step 1 do (
    for j: i+1 thru number_factors step 1 do (
        /* Find the list of interactions between factors i and j tested in block A, if this has full cardinality then increase "full_interactions" */
        interactions_tested_a: setify(makelist([coord_list_a[x][i], coord_list_a[x][j]], x, 1, length(coord_list_a))),
        if cardinality(interactions_tested_a) = length(factors_info[i])*length(factors_info[j]) then full_interactions: full_interactions + 1,
        /* Repeat for block B*/
        interactions_tested_b: setify(makelist([coord_list_b[x][i], coord_list_b[x][j]], x, 1, length(coord_list_b))),
        if cardinality(interactions_tested_b) = length(factors_info[i])*length(factors_info[j]) then full_interactions: full_interactions + 1
    )
);

/* Test whether more than 1 "full interaction" has been achieved */
interactions: if full_interactions > 0 then 1 else 0;

/*Node 6: Internal Variability */
internal_variability: if cardinality(setify(block_a_ans))<length(block_a_ans) or cardinality(setify(block_b_ans))<length(block_b_ans) then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Export data button</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Click 'Export data as CSV' below to obtain a .csv file with treatment details and yields.</p>

<button onclick="exportToCsv()">Export data as CSV</button>

<script>
var Results= {#data_set#};

exportToCsv = function() {
  var CsvString = "";
  Results.forEach(function(RowItem, RowIndex) {
    RowItem.forEach(function(ColItem, ColIndex) {
      CsvString += ColItem + ',';
    });
    CsvString += "\r\n";
  });
  CsvString = "data:application/csv," + encodeURIComponent(CsvString);
 var x = document.createElement("A");
 x.setAttribute("href", CsvString );
 x.setAttribute("download","details_list.csv");
 document.body.appendChild(x);
 x.click();
}
</script>

<br>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Feedback introduction statement</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>There are no right or wrong answers to this question. Your choice of design favours certain principles over others. The analysis you can do and the accuracy of your results will depend on the design you have chosen.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check for same treatments in both blocks</description>
        <answertest>AlgEquiv</answertest>
        <sans>blocking</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>Your design prioritises "Blocking" as a design principle because both blocks contain the same treatments. This means that in each comparison there are at least two repetitions. This increases the reliability of your results.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>

<p>Your design compromises "Blocking" as a design principle because your two blocks do not have the same treatments. The differences between the blocks can therefore confuse the effects when you compare treatments. Your analysis may need to take this into account.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for use of all treatments across both blocks</description>
        <answertest>AlgEquiv</answertest>
        <sans>full_factorial</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>Your design gives priority to "Full Factorial", i.e. it includes all treatments as a design principle. Each of the {@number_treatments@} possible treatments appears at least once. This way, you have data on all possible combinations and can consider all possible comparisons in your analysis.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>

<p>Your design compromises "Full Factorial" as a design principle because there is at least one treatment that you have not included. This means that you have no information about how this treatment may have been carried out and cannot include it in comparisons.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks if a block contains all possible interactions of two variables</description>
        <answertest>AlgEquiv</answertest>
        <sans>interactions</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>Your design prioritises "Interaction" as a design principle because in one block there are four comparable treatments highlighting interactions between two factors for a fixed value of the third factor.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>

<p>Your design compromises "Interaction" as a design principle because none of the blocks contain four comparable treatments highlighting interactions between two factors for a fixed value of the third factor.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Checks of a repeated treatment in either block</description>
        <answertest>AlgEquiv</answertest>
        <sans>internal_variability</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<br>

<p>Your design prioritises "Internal Variability" as a design principle as in a block you have at least one repeated treatment. As blocks are considered homogeneous, repetitions in a block can provide information about internal variability. But you are therefore comparing fewer treatments in at least one of the blocks.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<br>

<p>Your design compromises "Internal Variability" as a design principle because there is no repetition of treatments in either block. But it did allow you to compare 6 different treatments in each block.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>